<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>10.17. Exercises &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.fafafbd97df8a0d1.css?v=B5361BC8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.d062db89aaf1dddd.css?v=B02DFDA" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/runtime.b6a5f441f397ace5.bundle.js?v=3F299E7"></script>
    <script defer="" src="../_static/637.d54be67956c5c660.bundle.js?v=915B1670"></script>
    <script defer="" src="../_static/runestone.f2a5bb1b0f5cce0b.bundle.js?v=871C0DAB"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="10.18. Chapter Assessment" href="ChapterAssessment.html" />
    <link rel="prev" title="10.16. Glossary" href="Glossary.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'COMP1000UofM';
    eBookConfig.basecourse = 'COMP1000UofM';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = false;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '6.0.10';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          COMP1000UofM
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=COMP1000UofM&page=Exercises'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">10.16. </span>Glossary" data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">10.18. </span>Chapter Assessment' data-toggle="tooltip" >              
    <a href="ChapterAssessment.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <p></p>
<div class="section" id="exercises">
<h1><span class="section-number">10.17. </span>Exercises<a class="headerlink" href="#exercises" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='q1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_1 data-question_label="10.17.1">
<div id=ac11_14_1_question class="ac_question col-md-12">
<p>Write a function named <code class="docutils literal notranslate"><span class="pre">num_test</span></code> that takes a number as input. If the number is greater than 10, the function should return “Greater than 10.” If the number is less than 10, the function should return “Less than 10.” If the number is equal to 10, the function should return “Equal to 10.”</p>

</div>
<textarea data-lang="python" id="ac11_14_1_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">


====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(num_test(5), "Less than 10.", "Testing the num_test function on input 5.")
      self.assertEqual(num_test(0), "Less than 10.", "Testing the num_test function on input 0.")
      self.assertEqual(num_test(12.99), "Greater than 10.", "Testing the num_test function on input 12.99.")
      self.assertEqual(num_test(10.00), "Equal to 10.", "Testing the num_test function on input 10.00.")



myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q2' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_2 data-question_label="10.17.2">
<div id=ac11_14_2_question class="ac_question col-md-12">
<p>Write a function that will return the number of digits in an integer.</p>

</div>
<textarea data-lang="python" id="ac11_14_2_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def numDigits(n):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(numDigits(2),1,"Tested numDigits on input of 2")
      self.assertEqual(numDigits(55),2,"Tested numDigits on input of 55")
      self.assertEqual(numDigits(1352),4,"Tested numDigits on input of 1352")
      self.assertEqual(numDigits(444),3,"Tested numDigits on input of 444")



myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=answer11_14_2 data-question_label="10.17.3">
<div id=answer11_14_2_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="answer11_14_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def numDigits(n):
    n_str = str(n)
    return len(n_str)


print(numDigits(50))
print(numDigits(20000))
print(numDigits(1))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function disqus_bfd6f74a183c4682b29c72c4411200fb_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_bfd6f74a183c4682b29c72c4411200fb_' + eBookConfig.course;
                    this.page.url = 'https://disqus_bfd6f74a183c4682b29c72c4411200fb_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_bfd6f74a183c4682b29c72c4411200fb_' + eBookConfig.course;
            disqus_url = 'http://disqus_bfd6f74a183c4682b29c72c4411200fb_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="disqus_bfd6f74a183c4682b29c72c4411200fb" onclick="disqus_bfd6f74a183c4682b29c72c4411200fb_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='disqus_bfd6f74a183c4682b29c72c4411200fb']").attr('data-disqus-identifier', 'disqus_bfd6f74a183c4682b29c72c4411200fb_' + eBookConfig.course);
  });
</script>
</div>
    </div>

</li>
<li><div id='q3' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_3 data-question_label="10.17.4">
<div id=ac11_14_3_question class="ac_question col-md-12">
<p>Write a function that reverses its string argument.</p>

</div>
<textarea data-lang="python" id="ac11_14_3_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def reverse(astring):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(reverse("happy"),"yppah","Tested reverse on input of 'happy'")
        self.assertEqual(reverse("Python"),"nohtyP","Tested reverse on input of 'Python'")
        self.assertEqual(reverse(""),"","Tested reverse on input of ''")




myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q4' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_4 data-question_label="10.17.5">
<div id=ac11_14_4_question class="ac_question col-md-12">
<p>Write a function that mirrors its string argument,
generating a string containing the original string and the string backwards.</p>

</div>
<textarea data-lang="python" id="ac11_14_4_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

def mirror(mystr):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(mirror("good"),"gooddoog","Tested mirror on input of 'good'")
        self.assertEqual(mirror("Python"),"PythonnohtyP","Tested mirror on input of 'Python'")
        self.assertEqual(mirror(""),"","Tested mirror on input of ''")
        self.assertEqual(mirror("a"),"aa","Tested mirror on input of 'a'")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=answer11_14_4 data-question_label="10.17.6">
<div id=answer11_14_4_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="answer11_14_4_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def reverse(mystr):
    reversed = ''
    for char in mystr:
        reversed = char + reversed
    return reversed

def mirror(mystr):
    return mystr + reverse(mystr)

assert mirror('good') == 'gooddoog'
assert mirror('Python') == 'PythonnohtyP'
assert mirror('') == ''
assert mirror('a') == 'aa'
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function disqus_70b7ac515456497c952a2de5caa27ab9_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_70b7ac515456497c952a2de5caa27ab9_' + eBookConfig.course;
                    this.page.url = 'https://disqus_70b7ac515456497c952a2de5caa27ab9_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_70b7ac515456497c952a2de5caa27ab9_' + eBookConfig.course;
            disqus_url = 'http://disqus_70b7ac515456497c952a2de5caa27ab9_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="disqus_70b7ac515456497c952a2de5caa27ab9" onclick="disqus_70b7ac515456497c952a2de5caa27ab9_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='disqus_70b7ac515456497c952a2de5caa27ab9']").attr('data-disqus-identifier', 'disqus_70b7ac515456497c952a2de5caa27ab9_' + eBookConfig.course);
  });
</script>
</div>
    </div>

</li>
<li><div id='q5' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_5 data-question_label="10.17.7">
<div id=ac11_14_5_question class="ac_question col-md-12">
<p>Write a function that removes all occurrences of a given letter from a string.</p>

</div>
<textarea data-lang="python" id="ac11_14_5_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def remove_letter(theLetter, theString):
    # your code here

====


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(remove_letter("a","apple"),"pple","Tested remove_letter on inputs of 'a' and 'apple'")
        self.assertEqual(remove_letter("a","banana"),"bnn","Tested remove_letter on inputs of 'a' and 'banana'")
        self.assertEqual(remove_letter("z","banana"),"banana","Tested remove_letter on inputs of 'z' and 'banana'")



myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q6' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_6 data-question_label="10.17.8">
<div id=ac11_14_6_question class="ac_question col-md-12">
<p>Although Python provides us with many list methods, it is good practice and very instructive to think about how they are implemented.  Implement a Python function that works like the following:</p>
<ol class="loweralpha simple">
<li><p>count</p></li>
<li><p>in</p></li>
<li><p>reverse</p></li>
<li><p>index</p></li>
<li><p>insert</p></li>
</ol>

</div>
<textarea data-lang="python" id="ac11_14_6_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=answer11_14_6 data-question_label="10.17.9">
<div id=answer11_14_6_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="answer11_14_6_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def count(obj, lst):
    count = 0
    for e in lst:
        if e == obj:
            count = count + 1
    return count

def is_in(obj, lst):  # cannot be called in() because in is a reserved keyword
    for e in lst:
        if e == obj:
            return True
    return False

def reverse(lst):
    reversed = []
    for i in range(len(lst)-1, -1, -1): # step through the original list backwards
        reversed.append(lst[i])
    return reversed

def index(obj, lst):
    for i in range(len(lst)):
        if lst[i] == obj:
            return i
    return -1

def insert(obj, index, lst):
    newlst = []
    for i in range(len(lst)):
        if i == index:
            newlst.append(obj)
        newlst.append(lst[i])
    return newlst

lst = [0, 1, 1, 2, 2, 3, 4, 5, 6, 7, 8, 9]
print(count(1, lst))
print(is_in(4, lst))
print(reverse(lst))
print(index(2, lst))
print(insert('cat', 4, lst))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function disqus_39ee0274e51d4c888cc20b6fefa4069c_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_39ee0274e51d4c888cc20b6fefa4069c_' + eBookConfig.course;
                    this.page.url = 'https://disqus_39ee0274e51d4c888cc20b6fefa4069c_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_39ee0274e51d4c888cc20b6fefa4069c_' + eBookConfig.course;
            disqus_url = 'http://disqus_39ee0274e51d4c888cc20b6fefa4069c_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="disqus_39ee0274e51d4c888cc20b6fefa4069c" onclick="disqus_39ee0274e51d4c888cc20b6fefa4069c_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='disqus_39ee0274e51d4c888cc20b6fefa4069c']").attr('data-disqus-identifier', 'disqus_39ee0274e51d4c888cc20b6fefa4069c_' + eBookConfig.course);
  });
</script>
</div>
    </div>

</li>
<li><div id='q7' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_7 data-question_label="10.17.10">
<div id=ac11_14_7_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">replace(s,</span> <span class="pre">old,</span> <span class="pre">new)</span></code> that replaces all occurences of
<code class="docutils literal notranslate"><span class="pre">old</span></code> with <code class="docutils literal notranslate"><span class="pre">new</span></code> in a string <code class="docutils literal notranslate"><span class="pre">s</span></code>:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">test</span><span class="p">(</span><span class="n">replace</span><span class="p">(</span><span class="s1">&#39;Mississippi&#39;</span><span class="p">,</span> <span class="s1">&#39;i&#39;</span><span class="p">,</span> <span class="s1">&#39;I&#39;</span><span class="p">),</span> <span class="s1">&#39;MIssIssIppI&#39;</span><span class="p">)</span>

<span class="n">s</span> <span class="o">=</span> <span class="s1">&#39;I love spom!  Spom is my favorite food.  Spom, spom, spom, yum!&#39;</span>
<span class="n">test</span><span class="p">(</span><span class="n">replace</span><span class="p">(</span><span class="n">s</span><span class="p">,</span> <span class="s1">&#39;om&#39;</span><span class="p">,</span> <span class="s1">&#39;am&#39;</span><span class="p">),</span>
       <span class="s1">&#39;I love spam!  Spam is my favorite food.  Spam, spam, spam, yum!&#39;</span><span class="p">)</span>

<span class="n">test</span><span class="p">(</span><span class="n">replace</span><span class="p">(</span><span class="n">s</span><span class="p">,</span> <span class="s1">&#39;o&#39;</span><span class="p">,</span> <span class="s1">&#39;a&#39;</span><span class="p">),</span>
       <span class="s1">&#39;I lave spam!  Spam is my favarite faad.  Spam, spam, spam, yum!&#39;</span><span class="p">)</span>
</pre></div>
</div>
<p><em>Hint</em>: use the <code class="docutils literal notranslate"><span class="pre">split</span></code> and <code class="docutils literal notranslate"><span class="pre">join</span></code> methods.</p>

</div>
<textarea data-lang="python" id="ac11_14_7_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def replace(s, old, new):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(replace('Mississippi','i','I'),'MIssIssIppI',"Tested replace on input 'Mississippi','i','I'")
        self.assertEqual(replace('Bookkeeper','e','A'),'BookkAApAr',"Tested failed on input 'Bookkeeper','e','A'")
        self.assertEqual(replace('Deeded','e','q'),'Dqqdqd',"Tested failed on input 'Deeded','e','q'")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q8' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_8 data-question_label="10.17.11">
<div id=ac11_14_8_question class="ac_question col-md-12">
<p>Write a Python function that will take a the list of 100 random integers between 0 and 1000 and return the maximum value.  (Note: there is a builtin function named <code class="docutils literal notranslate"><span class="pre">max</span></code> but pretend you cannot use it.)</p>

</div>
<textarea data-lang="python" id="ac11_14_8_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
import random as r
lst = []

for i in range(100):
    num = r.randint(1, 1000)
    lst.append(num)

def largest(lst):
    #your code here

====
from unittest.gui import TestCaseGui
import re
class myTests(TestCaseGui):
    def testOne(self):
        output = self.getOutput().split('\n')
        editor = self.getEditorText().split('\n')
        float_re = r'[-+]?(\d+(\.\d*)?|\.\d+)([eE][-+]?\d+)?'

        self.assertEqual(largest(lst), max(lst), 'Checking for the list'+ str(lst))

        # hardcode check
        self.assertFalse(re.search(r'max', self.getEditorText()), 'Checking for max')
myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=answer11_14_8 data-question_label="10.17.12">
<div id=answer11_14_8_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="answer11_14_8_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import random

def max(lst):
    max = 0
    for e in lst:
        if e > max:
            max = e
    return max

lst = []
for i in range(100):
    lst.append(random.randint(0, 1000))

print(max(lst))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function disqus_714fd5537ebf41189ce5fb6fb16d1d26_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_714fd5537ebf41189ce5fb6fb16d1d26_' + eBookConfig.course;
                    this.page.url = 'https://disqus_714fd5537ebf41189ce5fb6fb16d1d26_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_714fd5537ebf41189ce5fb6fb16d1d26_' + eBookConfig.course;
            disqus_url = 'http://disqus_714fd5537ebf41189ce5fb6fb16d1d26_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="disqus_714fd5537ebf41189ce5fb6fb16d1d26" onclick="disqus_714fd5537ebf41189ce5fb6fb16d1d26_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='disqus_714fd5537ebf41189ce5fb6fb16d1d26']").attr('data-disqus-identifier', 'disqus_714fd5537ebf41189ce5fb6fb16d1d26_' + eBookConfig.course);
  });
</script>
</div>
    </div>

</li>
<li><div id='q9' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_9 data-question_label="10.17.13">
<div id=ac11_14_9_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">sum_of_squares(xs)</span></code> that computes the sum
of the squares of the numbers in the list <code class="docutils literal notranslate"><span class="pre">xs</span></code>.  For example,
<code class="docutils literal notranslate"><span class="pre">sum_of_squares([2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> should return 4+9+16 which is 29:</p>

</div>
<textarea data-lang="python" id="ac11_14_9_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def sum_of_squares(xs):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum_of_squares([2,3,4]),29,"Tested sum_of_squares on input [2,3,4]")
        self.assertEqual(sum_of_squares([0,1,-1]),2,"Tested sum_of_squares on input [0,1,-1]")
        self.assertEqual(sum_of_squares([5,12,14]),365,"Tested sum_of_squares on input [5,12,14]")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q10' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_10 data-question_label="10.17.14">
<div id=ac11_14_10_question class="ac_question col-md-12">
<p>Write a function to count how many odd numbers are in a list.</p>

</div>
<textarea data-lang="python" id="ac11_14_10_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def countOdd(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(countOdd([1,3,5,7,9]),5,"Tested countOdd on input [1,3,5,7,9]")
        self.assertEqual(countOdd([1,2,3,4,5]),3,"Tested countOdd on input [-1,-2,-3,-4,-5]")
        self.assertEqual(countOdd([2,4,6,8,10]),0,"Tested countOdd on input [2,4,6,8,10]")
        self.assertEqual(countOdd([0,-1,12,-33]),2,"Tested countOdd on input [0,-1,12,-33]")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=answer11_14_10 data-question_label="10.17.15">
<div id=answer11_14_10_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="answer11_14_10_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import random

def countOdd(lst):
    odd = 0
    for e in lst:
        if e % 2 != 0:
            odd = odd + 1
    return odd

# make a random list to test the function
lst = []
for i in range(100):
    lst.append(random.randint(0, 1000))

print(countOdd(lst))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function disqus_fdd366b1b4c8494082a385e1e1197844_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_fdd366b1b4c8494082a385e1e1197844_' + eBookConfig.course;
                    this.page.url = 'https://disqus_fdd366b1b4c8494082a385e1e1197844_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_fdd366b1b4c8494082a385e1e1197844_' + eBookConfig.course;
            disqus_url = 'http://disqus_fdd366b1b4c8494082a385e1e1197844_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="disqus_fdd366b1b4c8494082a385e1e1197844" onclick="disqus_fdd366b1b4c8494082a385e1e1197844_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='disqus_fdd366b1b4c8494082a385e1e1197844']").attr('data-disqus-identifier', 'disqus_fdd366b1b4c8494082a385e1e1197844_' + eBookConfig.course);
  });
</script>
</div>
    </div>

</li>
<li><div id='q11' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_11 data-question_label="10.17.16">
<div id=ac11_14_11_question class="ac_question col-md-12">
<p>Sum up all the even numbers in a list.</p>

</div>
<textarea data-lang="python" id="ac11_14_11_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def sumEven(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(sumEven([1,3,5,7,9]),0,"Tested sumEven on input [1,3,5,7,9]")
      self.assertEqual(sumEven([-1,-2,-3,-4,-5]),-6,"Tested sumEven on input [-1,-2,-3,-4,-5]")
      self.assertEqual(sumEven([2,4,6,7,9]),12,"Tested sumEven on input [2,4,6,7,9]")
      self.assertEqual(sumEven([0,1,12,33]),12,"Tested sumEven on input [0,1,12,33]")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q12' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_12 data-question_label="10.17.17">
<div id=ac11_14_12_question class="ac_question col-md-12">
<p>Sum up all the negative numbers in a list.</p>

</div>
<textarea data-lang="python" id="ac11_14_12_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def sumNegatives(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumNegatives([-1,-2,-3,-4,-5]),-15,"Tested sumNegatives on input [-1,-2,-3,-4,-5]")
        self.assertEqual(sumNegatives([1,-3,5,-7,9]),-10,"Tested sumNegatives on input [1,-3,5,-7,9]")
        self.assertEqual(sumNegatives([-2,-4,6,-7,9]),-13,"Tested sumNegatives on input [-2,-4,6,-7,9]")
        self.assertEqual(sumNegatives([0,1,2,3,4]),0,"Tested sumNegatives on input [0,1,2,3,4]")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=answer11_14_12 data-question_label="10.17.18">
<div id=answer11_14_12_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="answer11_14_12_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import random

def sumNegative(lst):
    sum = 0
    for e in lst:
        if e < 0:
            sum = sum + e
    return sum

lst = []
for i in range(100):
    lst.append(random.randrange(-1000, 1000))

print(sumNegative(lst))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function disqus_bfe671ac1e0942f2be4de7179921f83f_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_bfe671ac1e0942f2be4de7179921f83f_' + eBookConfig.course;
                    this.page.url = 'https://disqus_bfe671ac1e0942f2be4de7179921f83f_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_bfe671ac1e0942f2be4de7179921f83f_' + eBookConfig.course;
            disqus_url = 'http://disqus_bfe671ac1e0942f2be4de7179921f83f_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="disqus_bfe671ac1e0942f2be4de7179921f83f" onclick="disqus_bfe671ac1e0942f2be4de7179921f83f_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='disqus_bfe671ac1e0942f2be4de7179921f83f']").attr('data-disqus-identifier', 'disqus_bfe671ac1e0942f2be4de7179921f83f_' + eBookConfig.course);
  });
</script>
</div>
    </div>

</li>
<li><div id='q13' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_13 data-question_label="10.17.19">
<div id=ac11_14_13_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">findHypot</span></code>.  The function will be given the length of two sides of a right-angled triangle and it should return the length of the hypotenuse. (Hint:  <code class="docutils literal notranslate"><span class="pre">x</span> <span class="pre">**</span> <span class="pre">0.5</span></code> will return the square root, or use <code class="docutils literal notranslate"><span class="pre">sqrt</span></code> from the math module)</p>

</div>
<textarea data-lang="python" id="ac11_14_13_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

def findHypot(a,b):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(findHypot(12.0,5.0),13.0,"Tested findHypot on inputs of 12.0 and 5.0")
        self.assertEqual(findHypot(14.0,48.0),50.0,"Tested findHypot on inputs of 14.0 and 48.0")
        self.assertEqual(findHypot(21.0,72.0),75.0,"Tested findHypot on inputs of 21.0 and 72.0")
        self.assertAlmostEqual(findHypot(1,1.73205),1.999999,2,"Tested findHypot on inputs of 1 and 1.73205")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q14' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_14 data-question_label="10.17.20">
<div id=ac11_14_14_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">is_even(n)</span></code> that takes an integer as an argument and returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if the argument is an <strong>even number</strong> and <code class="docutils literal notranslate"><span class="pre">False</span></code> if it is <strong>odd</strong>.</p>

</div>
<textarea data-lang="python" id="ac11_14_14_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def is_even(n):
    #your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
     def testOne(self):
         self.assertEqual(is_even(10),True,"Tested is_even on input of 10")
         self.assertEqual(is_even(5),False,"Tested is_even on input of 5")
         self.assertEqual(is_even(1),False,"Tested is_even on input of 1")
         self.assertEqual(is_even(0),True,"Tested is_even on input of 0")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q15' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_15 data-question_label="10.17.21">
<div id=ac11_14_15_question class="ac_question col-md-12">
<p>Now write the function <code class="docutils literal notranslate"><span class="pre">is_odd(n)</span></code> that returns <code class="docutils literal notranslate"><span class="pre">True</span></code> when <code class="docutils literal notranslate"><span class="pre">n</span></code> is odd and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise.</p>

</div>
<textarea data-lang="python" id="ac11_14_15_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

def is_odd(n):
    # your code here


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(is_odd(10),False,"Tested is_odd on input of 10")
        self.assertEqual(is_odd(5),True,"Tested is_odd on input of 5")
        self.assertEqual(is_odd(1),True,"Tested is_odd on input of 1")
        self.assertEqual(is_odd(0),False,"Tested is_odd on input of 0")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q16' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_14_16 data-question_label="10.17.22">
<div id=ac11_14_16_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">is_rightangled</span></code> which, given the length of three sides of a triangle, will determine whether the triangle is right-angled. Assume that the third argument to the function is always the longest side. It will return <code class="docutils literal notranslate"><span class="pre">True</span></code> if the triangle is right-angled, or <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise.</p>
<p>Hint: floating point arithmetic is not always exactly accurate,
so it is not safe to test floating point numbers for equality.
If a good programmer wants to know whether
<code class="docutils literal notranslate"><span class="pre">x</span></code> is equal or close enough to <code class="docutils literal notranslate"><span class="pre">y</span></code>, they would probably code it up as</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">if</span>  <span class="nb">abs</span><span class="p">(</span><span class="n">x</span> <span class="o">-</span> <span class="n">y</span><span class="p">)</span> <span class="o">&lt;</span> <span class="mf">0.001</span><span class="p">:</span>      <span class="c1"># if x is approximately equal to y</span>
    <span class="o">...</span>
</pre></div>
</div>

</div>
<textarea data-lang="python" id="ac11_14_16_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def is_rightangled(a, b, c):
    # your code here


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(is_rightangled(1.5,2.0,2.5),True,"Tested is_rightangled on inputs of 1.5, 2.0 and 2.5")
        self.assertEqual(is_rightangled(4.0,8.0,16.0),False,"Tested is_rightangled on inputs of 4.0, 8.0 and 16.0")
        self.assertEqual(is_rightangled(4.1,8.2,9.1678787077),True,"Tested is_rightangled on inputs of 4.1, 8.2 and 9.1678787077")
        self.assertEqual(is_rightangled(4.1,8.2,9.16787),True,"Tested is_rightangled on inputs of 4.1, 8.2, and 9.16787")
        self.assertEqual(is_rightangled(4.1,8.2,9.168),False,"Tested is_rightangled on inputs of 4.1, 8.2 and 9.168")
        self.assertEqual(is_rightangled(0.5,0.4,0.64031),True,"Tested is_rightangled on inputs of 0.5, 0.4 and 0.64031")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
</ol>
<div class="section" id="contributed-exercises">
<h2><span class="section-number">10.17.1. </span>Contributed Exercises<a class="headerlink" href="#contributed-exercises" title="Permalink to this headline">¶</a></h2>
{% for q in questions: %}
    <div class='oneq full-width'>
        {{ q['htmlsrc']|safe }}
    </div>
{% endfor %}</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 10.16. Glossary' data-toggle="tooltip">
  <a href="Glossary.html" aria-label="Previous - 10.16. Glossary">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 10.18. Chapter Assessment' data-toggle="tooltip" >
    <a href="ChapterAssessment.html" aria-label="Next - 10.18. Chapter Assessment">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 6.0.10.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>