<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>10.5. Returning a value from a function &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.fafafbd97df8a0d1.css?v=B5361BC8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.d062db89aaf1dddd.css?v=B02DFDA" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/runtime.b6a5f441f397ace5.bundle.js?v=3F299E7"></script>
    <script defer="" src="../_static/637.d54be67956c5c660.bundle.js?v=915B1670"></script>
    <script defer="" src="../_static/runestone.f2a5bb1b0f5cce0b.bundle.js?v=871C0DAB"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="10.6. 👩‍💻 Decoding a Function" href="DecodingaFunction.html" />
    <link rel="prev" title="10.4. Function Parameters" href="FunctionParameters.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'COMP1000UofM';
    eBookConfig.basecourse = 'COMP1000UofM';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = false;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '6.0.10';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          COMP1000UofM
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=COMP1000UofM&page=Returningavaluefromafunction'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">10.4. </span>Function Parameters" data-toggle="tooltip">
    <a href="FunctionParameters.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">10.6. </span>👩‍💻 Decoding a Function' data-toggle="tooltip" >              
    <a href="DecodingaFunction.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="returning-a-value-from-a-function">
<h1><span class="section-number">10.5. </span>Returning a value from a function<a class="headerlink" href="#returning-a-value-from-a-function" title="Permalink to this headline">¶</a></h1>
<img alt="gif of a box labeled function with three spaces on the top for input and a space on the bottom for output. Three arrows enter the top and are labeled as input or arguments. The function box shakes, and then one arrow leaves the bottom of the function box." src="../_images/function_call.gif" />
<p>Not only can you pass a parameter value into a function, a function can also produce a value. You have already
seen this in some previous functions that you have used. For example, <code class="docutils literal notranslate"><span class="pre">len</span></code> takes a list or string as a parameter
value and returns a number, the length of that list or string. <code class="docutils literal notranslate"><span class="pre">range</span></code> takes an integer as a parameter value and
returns a list containing all the numbers from 0 up to that parameter value.</p>
<p>Functions that return values are sometimes called <strong>fruitful functions</strong>. In many other languages, a function that
doesn’t return a value is called a <strong>procedure</strong>, but we will stick here with the Python way of also calling it a
function, or if we want to stress it, a <em>non-fruitful</em> function.</p>

    <div class="runestone" style="margin-left: auto; margin-right:auto">
        <div id="goog_return_values" data-component="youtube" class="align-left youtube-video"
            data-video-height="315" data-question_label="10.5.1"
            data-video-width="560" data-video-videoid="LGOZyrRCJ1o"
            data-video-divid="goog_return_values" data-video-start="0"
            data-video-end="-1" ></div>
    </div>
    <img alt="../_images/blackboxfun.png" src="../_images/blackboxfun.png" />
<p>How do we write our own fruitful function? Let’s start by creating a very simple mathematical function that we will
call <code class="docutils literal notranslate"><span class="pre">square</span></code>. The square function will take one number as a parameter and return the result of squaring that
number. Here is the black-box diagram with the Python code following.</p>
<img alt="../_images/squarefun.png" src="../_images/squarefun.png" />

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_4_1 data-question_label="10.5.2">
<div id=ac11_4_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac11_4_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def square(x):
    y = x * x
    return y

toSquare = 10
result = square(toSquare)
print("The result of {} squared is {}.".format(toSquare, result))
</textarea>
</div>
</div>
<p>The <strong>return</strong> statement is followed by an expression which is evaluated. Its result is returned to the caller as the
“fruit” of calling this function. Because the return statement can contain any Python expression we could have
avoided creating the <strong>temporary variable</strong> <code class="docutils literal notranslate"><span class="pre">y</span></code> and simply used <code class="docutils literal notranslate"><span class="pre">return</span> <span class="pre">x*x</span></code>. Try modifying the square function
above to see that this works just the same. On the other hand, using <strong>temporary variables</strong> like <code class="docutils literal notranslate"><span class="pre">y</span></code> in the program
above makes debugging easier. These temporary variables are referred to as <strong>local variables</strong>.</p>
<p>Notice something important here. The name of the variable we pass as an argument — <code class="docutils literal notranslate"><span class="pre">toSquare</span></code> — has nothing to
do with the name of the formal parameter — <code class="docutils literal notranslate"><span class="pre">x</span></code>.  It is as if  <code class="docutils literal notranslate"><span class="pre">x</span> <span class="pre">=</span> <span class="pre">toSquare</span></code> is executed when <code class="docutils literal notranslate"><span class="pre">square</span></code> is
called. It doesn’t matter what the value was named in the caller (the place where the function was invoked). Inside
<code class="docutils literal notranslate"><span class="pre">square</span></code>, it’s name is <code class="docutils literal notranslate"><span class="pre">x</span></code>.  You can see this very clearly in codelens, where the global variables and the local
variables for the square function are in separate boxes.</p>

<div class="runestone" style="max-width: none;">
    <div class="alert alert-warning cd_section" data-component="codelens" data-question_label="10.5.3">
        <div class="pytutorVisualizer" id="clens11_4_1"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 10.5.3 (clens11_4_1)</span> </p>
    </div>

<script>
var clens11_4_1_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["clens11_4_1"] = {"code": "def square(x):\n    y = x * x\n    return y\n\ntoSquare = 10\nsquareResult = square(toSquare)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1]}, "ordered_globals": ["square"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "toSquare": 10}, "ordered_globals": ["square", "toSquare"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "square", "globals": {"square": ["REF", 1], "toSquare": 10}, "ordered_globals": ["square", "toSquare"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": 10}, "ordered_varnames": ["x"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "toSquare": 10}, "ordered_globals": ["square", "toSquare"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": 10}, "ordered_varnames": ["x"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "toSquare": 10}, "ordered_globals": ["square", "toSquare"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": 10, "y": 100}, "ordered_varnames": ["x", "y"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 3, "event": "return", "func_name": "square", "globals": {"square": ["REF", 1], "toSquare": 10}, "ordered_globals": ["square", "toSquare"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": 10, "y": 100, "__return__": 100}, "ordered_varnames": ["x", "y", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 6, "event": "return", "func_name": "<module>", "globals": {"square": ["REF", 1], "toSquare": 10, "squareResult": 100}, "ordered_globals": ["square", "toSquare", "squareResult"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}]};
</script>
</div>
<p>There is one more aspect of function return values that should be noted. All Python functions return the special value
<code class="docutils literal notranslate"><span class="pre">None</span></code> unless there is an explicit return statement with a value other than <code class="docutils literal notranslate"><span class="pre">None</span></code>. Consider the following common
mistake made by beginning Python programmers. As you step through this example, pay very close attention to the return
value in the local variables listing. Then look at what is printed when the function is over.</p>

<div class="runestone" style="max-width: none;">
    <div class="alert alert-warning cd_section" data-component="codelens" data-question_label="10.5.4">
        <div class="pytutorVisualizer" id="clens11_4_2"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 10.5.4 (clens11_4_2)</span> </p>
    </div>

<script>
var clens11_4_2_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["clens11_4_2"] = {"code": "def square(x):\n    y = x * x\n    print(y)   # Bad! This is confusing! Should use return instead!\n\ntoSquare = 10\nsquareResult = square(toSquare)\nprint(\"The result of {} squared is {}.\".format(toSquare, squareResult))", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1]}, "ordered_globals": ["square"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "toSquare": 10}, "ordered_globals": ["square", "toSquare"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "square", "globals": {"square": ["REF", 1], "toSquare": 10}, "ordered_globals": ["square", "toSquare"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": 10}, "ordered_varnames": ["x"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "toSquare": 10}, "ordered_globals": ["square", "toSquare"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": 10}, "ordered_varnames": ["x"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "toSquare": 10}, "ordered_globals": ["square", "toSquare"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": 10, "y": 100}, "ordered_varnames": ["x", "y"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 3, "event": "return", "func_name": "square", "globals": {"square": ["REF", 1], "toSquare": 10}, "ordered_globals": ["square", "toSquare"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": 10, "y": 100, "__return__": null}, "ordered_varnames": ["x", "y", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": "100\n"}, {"line": 7, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "toSquare": 10, "squareResult": null}, "ordered_globals": ["square", "toSquare", "squareResult"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": "100\n"}, {"line": 7, "event": "return", "func_name": "<module>", "globals": {"square": ["REF", 1], "toSquare": 10, "squareResult": null}, "ordered_globals": ["square", "toSquare", "squareResult"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": "100\nThe result of 10 squared is None.\n"}]};
</script>
</div>
<p>The problem with this function is that even though it prints the value of the squared input, that value will not be
returned to the place where the call was done. Instead, the value <code class="docutils literal notranslate"><span class="pre">None</span></code> will be returned. Since line 6 uses the
return value as the right hand side of an assignment statement, <code class="docutils literal notranslate"><span class="pre">squareResult</span></code> will have <code class="docutils literal notranslate"><span class="pre">None</span></code> as its value and
the result printed in line 7 is incorrect. Typically, functions will return values that can be printed or processed in
some other way by the caller.</p>
<p>A return statement, once executed, immediately terminates execution of a function, even if it is not the last
statement in the function. In the following code, when line 3 executes, the value 5 is returned and assigned to the
variable x, then printed. Lines 4 and 5 never execute. Run the following code and try making some modifications of
it to make sure you understand why “there” and 10 never print out.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_4_2 data-question_label="10.5.5">
<div id=ac11_4_2_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac11_4_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def weird():
    print("here")
    return 5
    print("there")
    return 10

x = weird()
print(x)
</textarea>
</div>
</div>
<p>The fact that a return statement immediately ends execution of the code block inside a function is important to
understand for writing complex programs, and it can also be very useful. The following example is a situation where
you can use this to your advantage – and understanding this will help you understand other people’s code better, and
be able to walk through code more confidently.</p>
<p>Consider a situation where you want to write a function to find out, from a class attendance list, whether anyone’s
first name is longer than five letters, called <code class="docutils literal notranslate"><span class="pre">longer_than_five</span></code>. If there is anyone in class whose first name is
longer than 5 letters, the function should return <code class="docutils literal notranslate"><span class="pre">True</span></code>. Otherwise, it should return <code class="docutils literal notranslate"><span class="pre">False</span></code>.</p>
<p>In this case, you’ll be using conditional statements in the code that exists in the <strong>function body</strong>, the code block
indented underneath the function definition statement (just like the code that starts with the line <code class="docutils literal notranslate"><span class="pre">print(&quot;here&quot;)</span></code>
in the example above – that’s the body of the function <code class="docutils literal notranslate"><span class="pre">weird</span></code>, above).</p>
<p><strong>Bonus challenge for studying:</strong> After you look at the explanation below, stop looking at the code – just the
description of the function above it, and try to write the code yourself! Then test it on different lists and make
sure that it works. But read the explanation first, so you can be sure you have a solid grasp on these function
mechanics.</p>
<p>First, an English plan for this new function to define called <code class="docutils literal notranslate"><span class="pre">longer_than_five</span></code>:</p>
<ul class="simple">
<li><p>You’ll want to pass in a list of strings (representing people’s first names) to the function.</p></li>
<li><p>You’ll want to iterate over all the items in the list, each of the strings.</p></li>
<li><p>As soon as you get to one name that is longer than five letters, you know the function should return <code class="docutils literal notranslate"><span class="pre">True</span></code> – yes, there is at least one name longer than five letters!</p></li>
<li><p>And if you go through the whole list and there was no name longer than five letters, then the function should return <code class="docutils literal notranslate"><span class="pre">False</span></code>.</p></li>
</ul>
<p>Now, the code:</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_4_3 data-question_label="10.5.6">
<div id=ac11_4_3_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac11_4_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def longer_than_five(list_of_names):
    for name in list_of_names: # iterate over the list to look at each name
        if len(name) > 5: # as soon as you see a name longer than 5 letters,
            return True # then return True!
            # If Python executes that return statement, the function is over and the rest of the code will not run -- you already have your answer!
    return False # You will only get to this line if you
    # iterated over the whole list and did not get a name where
    # the if expression evaluated to True, so at this point, it's correct to return False!

# Here are a couple sample calls to the function with different lists of names. Try running this code in Codelens a few times and make sure you understand exactly what is happening.

list1 = ["Sam","Tera","Sal","Amita"]
list2 = ["Rey","Ayo","Lauren","Natalie"]

print(longer_than_five(list1))
print(longer_than_five(list2))
</textarea>
</div>
</div>
<p>So far, we have just seen return values being assigned to variables. For example, we had the line
<code class="docutils literal notranslate"><span class="pre">squareResult</span> <span class="pre">=</span> <span class="pre">square(toSquare)</span></code>. As with all assignment statements, the right hand side is executed first. It
invokes the <code class="docutils literal notranslate"><span class="pre">square</span></code> function, passing in a parameter value 10 (the current value of <code class="docutils literal notranslate"><span class="pre">toSquare</span></code>). That returns a
value 100, which completes the evaluation of the right-hand side of the assignment. 100 is then assigned to the
variable <code class="docutils literal notranslate"><span class="pre">squareResult</span></code>. In this case, the function invocation was the entire expression that was evaluated.</p>
<p>Function invocations, however, can also be used as part of more complicated expressions. For example,
<code class="docutils literal notranslate"><span class="pre">squareResult</span> <span class="pre">=</span> <span class="pre">2</span> <span class="pre">*</span> <span class="pre">square(toSquare)</span></code>. In this case, the value 100 is returned and is then multiplied by 2 to
produce the value 200. When python evaluates an expression like <code class="docutils literal notranslate"><span class="pre">x</span> <span class="pre">*</span> <span class="pre">3</span></code>, it substitutes the current value of x into
the expression and then does the multiplication. When python evaluates an expression like <code class="docutils literal notranslate"><span class="pre">2</span> <span class="pre">*</span> <span class="pre">square(toSquare)</span></code>, it
substitutes the return value 100 for entire function invocation and then does the multiplication.</p>
<p>To reiterate, when executing a line of code <code class="docutils literal notranslate"><span class="pre">squareResult</span> <span class="pre">=</span> <span class="pre">2</span> <span class="pre">*</span> <span class="pre">square(toSquare)</span></code>, the python
interpreter does these steps:</p>
<ol class="arabic simple">
<li><p>It’s an assignment statement, so evaluate the right-hand side expression <code class="docutils literal notranslate"><span class="pre">2</span> <span class="pre">*</span> <span class="pre">square(toSquare)</span></code>.</p></li>
<li><p>Look up the values of the variables square and toSquare: square is a function object and toSquare is 10</p></li>
<li><p>Pass 10 as a parameter value to the function, get back the return value 100</p></li>
<li><p>Substitute 100 for square(toSquare), so that the expression now reads <code class="docutils literal notranslate"><span class="pre">2</span> <span class="pre">*</span> <span class="pre">100</span></code></p></li>
<li><p>Assign 200 to variable <code class="docutils literal notranslate"><span class="pre">squareResult</span></code></p></li>
</ol>
<p><strong>Check your understanding</strong></p>

            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="10.5.7" data-multipleanswers="false"  id="question11_4_1"  style="visibility: hidden;">
            <p>func-4-7: What is wrong with the following function definition:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">addEm</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="n">y</span><span class="p">,</span> <span class="n">z</span><span class="p">):</span>
    <span class="k">return</span> <span class="n">x</span><span class="o">+</span><span class="n">y</span><span class="o">+</span><span class="n">z</span>
    <span class="nb">print</span><span class="p">(</span><span class="s1">&#39;the answer is&#39;</span><span class="p">,</span> <span class="n">x</span><span class="o">+</span><span class="n">y</span><span class="o">+</span><span class="n">z</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer"  id="question11_4_1_opt_a">You should never use a print statement in a function definition.</li><li data-component="feedback">Although you should not mistake print for return, you may include print statements inside your functions.</li>
            
            <li data-component="answer" data-correct id="question11_4_1_opt_b">You should not have any statements in a function after the return statement.  Once the function gets to the return statement it will immediately stop executing the function.</li><li data-component="feedback">This is a very common mistake so be sure to watch out for it when you write your code!</li>
            
            <li data-component="answer"  id="question11_4_1_opt_c">You must calculate the value of x+y+z before you return it.</li><li data-component="feedback">Python will automatically calculate the value x+y+z and then return it in the statement as it is written</li>
            
            <li data-component="answer"  id="question11_4_1_opt_d">A function cannot return a number.</li><li data-component="feedback">Functions can return any legal data, including (but not limited to) numbers, strings, lists, dictionaries, etc.</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="10.5.8" data-multipleanswers="false"  id="question11_4_2"  style="visibility: hidden;">
            <p>func-4-8: What will the following function return?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">addEm</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="n">y</span><span class="p">,</span> <span class="n">z</span><span class="p">):</span>
    <span class="nb">print</span><span class="p">(</span><span class="n">x</span><span class="o">+</span><span class="n">y</span><span class="o">+</span><span class="n">z</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="question11_4_2_opt_a">The value None</li><li data-component="feedback">We have accidentally used print where we mean return.  Therefore, the function will return the value None by default.  This is a VERY COMMON mistake so watch out!  This mistake is also particularly difficult to find because when you run the function the output looks the same.  It is not until you try to assign its value to a variable that you can notice a difference.</li>
            
            <li data-component="answer"  id="question11_4_2_opt_b">The value of x+y+z</li><li data-component="feedback">Careful!  This is a very common mistake.  Here we have printed the value x+y+z but we have not returned it.  To return a value we MUST use the return keyword.</li>
            
            <li data-component="answer"  id="question11_4_2_opt_c">The string 'x+y+z'</li><li data-component="feedback">x+y+z calculates a number (assuming x+y+z are numbers) which represents the sum of the values x, y and z.</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="10.5.9" data-multipleanswers="false"  id="question11_4_3"  style="visibility: hidden;">
            <p>func-4-9: What will the following code output?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">square</span><span class="p">(</span><span class="n">x</span><span class="p">):</span>
    <span class="n">y</span> <span class="o">=</span> <span class="n">x</span> <span class="o">*</span> <span class="n">x</span>
    <span class="k">return</span> <span class="n">y</span>

<span class="nb">print</span><span class="p">(</span><span class="n">square</span><span class="p">(</span><span class="mi">5</span><span class="p">)</span> <span class="o">+</span> <span class="n">square</span><span class="p">(</span><span class="mi">5</span><span class="p">))</span>
</pre></div>
</div>

            <li data-component="answer"  id="question11_4_3_opt_a">25</li><li data-component="feedback">It squares 5 twice, and adds them together.</li>
            
            <li data-component="answer" data-correct id="question11_4_3_opt_b">50</li><li data-component="feedback">The two return values are added together.</li>
            
            <li data-component="answer"  id="question11_4_3_opt_c">25 + 25</li><li data-component="feedback">The two results are substituted into the expression and then it is evaluated. The returned values are integers in this case, not strings.</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="10.5.10" data-multipleanswers="false"  id="question11_4_4"  style="visibility: hidden;">
            <p>func-4-10: What will the following code output?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">square</span><span class="p">(</span><span class="n">x</span><span class="p">):</span>
    <span class="n">y</span> <span class="o">=</span> <span class="n">x</span> <span class="o">*</span> <span class="n">x</span>
    <span class="k">return</span> <span class="n">y</span>

<span class="nb">print</span><span class="p">(</span><span class="n">square</span><span class="p">(</span><span class="n">square</span><span class="p">(</span><span class="mi">2</span><span class="p">)))</span>
</pre></div>
</div>

            <li data-component="answer"  id="question11_4_4_opt_a">8</li><li data-component="feedback">It squares 2, yielding the value 4. But that doesn't mean the next value multiplies 2 and 4.</li>
            
            <li data-component="answer" data-correct id="question11_4_4_opt_b">16</li><li data-component="feedback">It squares 2, yielding the value 4. 4 is then passed as a value to square again, yeilding 16.</li>
            
            <li data-component="answer"  id="question11_4_4_opt_c">Error: can't put a function invocation inside parentheses</li><li data-component="feedback">This is a more complicated expression, but still valid. The expression square(2) is evaluated, and the return value 4 substitutes for square(2) in the expression.</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="10.5.11" data-multipleanswers="false"  id="question11_4_5"  style="visibility: hidden;">
            <p>func-4-11: What will the following code output?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">cyu2</span><span class="p">(</span><span class="n">s1</span><span class="p">,</span> <span class="n">s2</span><span class="p">):</span>
    <span class="n">x</span> <span class="o">=</span> <span class="nb">len</span><span class="p">(</span><span class="n">s1</span><span class="p">)</span>
    <span class="n">y</span> <span class="o">=</span> <span class="nb">len</span><span class="p">(</span><span class="n">s2</span><span class="p">)</span>
    <span class="k">return</span> <span class="n">x</span><span class="o">-</span><span class="n">y</span>

<span class="n">z</span> <span class="o">=</span> <span class="n">cyu2</span><span class="p">(</span><span class="s2">&quot;Yes&quot;</span><span class="p">,</span> <span class="s2">&quot;no&quot;</span><span class="p">)</span>
<span class="k">if</span> <span class="n">z</span> <span class="o">&gt;</span> <span class="mi">0</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;First one was longer&quot;</span><span class="p">)</span>
<span class="k">else</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;Second one was at least as long&quot;</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer"  id="question11_4_5_opt_a">1</li><li data-component="feedback">cyu2 returns the value 1, but that's not what prints.</li>
            
            <li data-component="answer"  id="question11_4_5_opt_b">Yes</li><li data-component="feedback">"Yes" is longer, but that's not what prints.</li>
            
            <li data-component="answer" data-correct id="question11_4_5_opt_c">First one was longer</li><li data-component="feedback">cyu2 returns the value 1, which is assigned to z.</li>
            
            <li data-component="answer"  id="question11_4_5_opt_d">Second one was at least as long</li><li data-component="feedback">cyu2 returns the value 1, which is assigned to z.</li>
            
            <li data-component="answer"  id="question11_4_5_opt_e">Error</li><li data-component="feedback">what do you think will cause an error.</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="10.5.12" data-multipleanswers="false"  id="question11_4_6"  style="visibility: hidden;">
            <p>func-4-12: Which will print out first, square, g, or a number?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">square</span><span class="p">(</span><span class="n">x</span><span class="p">):</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;square&quot;</span><span class="p">)</span>
    <span class="k">return</span> <span class="n">x</span><span class="o">*</span><span class="n">x</span>

<span class="k">def</span> <span class="nf">g</span><span class="p">(</span><span class="n">y</span><span class="p">):</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;g&quot;</span><span class="p">)</span>
    <span class="k">return</span> <span class="n">y</span> <span class="o">+</span> <span class="mi">3</span>

<span class="nb">print</span><span class="p">(</span><span class="n">square</span><span class="p">(</span><span class="n">g</span><span class="p">(</span><span class="mi">2</span><span class="p">)))</span>
</pre></div>
</div>

            <li data-component="answer"  id="question11_4_6_opt_a">square</li><li data-component="feedback">Before executing square, it has to figure out what value to pass in, so g is executed first</li>
            
            <li data-component="answer" data-correct id="question11_4_6_opt_b">g</li><li data-component="feedback">g has to be executed and return a value in order to know what paramater value to provide to x.</li>
            
            <li data-component="answer"  id="question11_4_6_opt_c">a number</li><li data-component="feedback">square and g both have to execute before the number is printed.</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="10.5.13" data-multipleanswers="false"  id="question11_4_7"  style="visibility: hidden;">
            <p>func-4-13: How many lines will the following code print?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">show_me_numbers</span><span class="p">(</span><span class="n">list_of_ints</span><span class="p">):</span>
    <span class="nb">print</span><span class="p">(</span><span class="mi">10</span><span class="p">)</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;Next we&#39;ll accumulate the sum&quot;</span><span class="p">)</span>
    <span class="n">accum</span> <span class="o">=</span> <span class="mi">0</span>
    <span class="k">for</span> <span class="n">num</span> <span class="ow">in</span> <span class="n">list_of_ints</span><span class="p">:</span>
        <span class="n">accum</span> <span class="o">=</span> <span class="n">accum</span> <span class="o">+</span> <span class="n">num</span>
    <span class="k">return</span> <span class="n">accum</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;All done with accumulation!&quot;</span><span class="p">)</span>

<span class="n">show_me_numbers</span><span class="p">([</span><span class="mi">4</span><span class="p">,</span><span class="mi">2</span><span class="p">,</span><span class="mi">3</span><span class="p">])</span>
</pre></div>
</div>

            <li data-component="answer"  id="question11_4_7_opt_a">3</li><li data-component="feedback">The function gets to a return statement after 2 lines are printed, so the third print statement will not run.</li>
            
            <li data-component="answer" data-correct id="question11_4_7_opt_b">2</li><li data-component="feedback">Yes! Two printed lines, and then the function body execution reaches a return statement.</li>
            
            <li data-component="answer"  id="question11_4_7_opt_c">None</li><li data-component="feedback">The function returns an integer value! However, this code does not print out the result of the function invocation, so you can't see it (print is for people). The only lines you see printed are the ones that occur in the print statements before the return statement.</li>
            

            </ul>
            </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_4_4 data-question_label="10.5.14">
<div id=ac11_4_4_question class="ac_question col-md-12">
<p><strong>8.</strong> Write a function named <code class="docutils literal notranslate"><span class="pre">same</span></code> that takes a string as input, and simply returns that string.</p>

</div>
<textarea data-lang="python" id="ac11_4_4_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(same('hello'), 'hello', "Testing the same function on input 'hello'.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_4_5 data-question_label="10.5.15">
<div id=ac11_4_5_question class="ac_question col-md-12">
<p><strong>9.</strong> Write a function called <code class="docutils literal notranslate"><span class="pre">same_thing</span></code> that returns the parameter, unchanged.</p>

</div>
<textarea data-lang="python" id="ac11_4_5_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(same_thing(5), 5,"Testing the function same_thing with input 5")
      self.assertEqual(same_thing("Welcome"), "Welcome", "Testing the function same_thing with input 'Welcome'")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_4_6 data-question_label="10.5.16">
<div id=ac11_4_6_question class="ac_question col-md-12">
<p><strong>10.</strong> Write a function called <code class="docutils literal notranslate"><span class="pre">subtract_three</span></code> that takes an integer or any number as input, and returns that number minus three.</p>

</div>
<textarea data-lang="python" id="ac11_4_6_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(subtract_three(9), 6, "Testing the subtract_three function on input 9.")
      self.assertEqual(subtract_three(-5), -8, "Testing the subtract_three function on input -5.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_4_7 data-question_label="10.5.17">
<div id=ac11_4_7_question class="ac_question col-md-12">
<p><strong>11.</strong> Write a function called <code class="docutils literal notranslate"><span class="pre">change</span></code> that takes one number as its input and returns that number, plus 7.</p>

</div>
<textarea data-lang="python" id="ac11_4_7_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(change(5), 12,"Testing the function change with input 5")
      self.assertEqual(change(-10), -3, "Testing the function change with input -10")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_4_8 data-question_label="10.5.18">
<div id=ac11_4_8_question class="ac_question col-md-12">
<p><strong>12.</strong> Write a function named <code class="docutils literal notranslate"><span class="pre">intro</span></code> that takes a string as input. This string ist intended to be a person’s name and the output is a standardized greeting. For example, given the string “Becky” as input, the function should return: “Hello, my name is Becky and I love SI 106.”</p>

</div>
<textarea data-lang="python" id="ac11_4_8_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(intro("Mike"), "Hello, my name is Mike and I love SI 106.", "Testing the intro function on input 'Mike'.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_4_9 data-question_label="10.5.19">
<div id=ac11_4_9_question class="ac_question col-md-12">
<p><strong>13.</strong> Write a function called <code class="docutils literal notranslate"><span class="pre">s_change</span></code> that takes one string as input and returns that string, concatenated with the string ” for fun.”.</p>

</div>
<textarea data-lang="python" id="ac11_4_9_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(s_change("Coding"), "Coding for fun." ,"Testing the function s_change with input coding")
      self.assertEqual(s_change("We go to the beach"), "We go to the beach for fun." , "Testing the function s_change with input We go to the beach")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac11_4_10 data-question_label="10.5.20">
<div id=ac11_4_10_question class="ac_question col-md-12">
<p><strong>14.</strong> Write a function called <code class="docutils literal notranslate"><span class="pre">decision</span></code> that takes a string as input, and then checks the number of characters. If it has over 17 characters, return “This is a long string”, if it is shorter or has 17 characters, return “This is a short string”.</p>

</div>
<textarea data-lang="python" id="ac11_4_10_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(decision("Well hello dolly"), "This is a short string", "Testing the function decision with input 'Well hello dolly'")
      self.assertEqual(decision("In olden days a glimps of stocking was looked on a something shocking but heaven knows, anything goes"), "This is a long string", "Testing the function decision with input 'In olden days a glimps of stocking was looked on a something shocking but heaven knows, anything goes'")
      self.assertEqual(decision("how do you do sir"), "This is a short string", "Testing the function decision with input 'how do you do sir'")

myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 10.4. Function Parameters' data-toggle="tooltip">
  <a href="FunctionParameters.html" aria-label="Previous - 10.4. Function Parameters">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 10.6. 👩‍💻 Decoding a Function' data-toggle="tooltip" >
    <a href="DecodingaFunction.html" aria-label="Next - 10.6. 👩‍💻 Decoding a Function">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 6.0.10.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>