<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>8.8. Function Composition (functions calling functions) &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.675cfab3e021b8f9.css?v=AC6C1FF6" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/runtime.15a489335c5c6109.bundle.js?v=8DB56D48"></script>
    <script defer="" src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script defer="" src="../_static/runestone.a1906624d1effde5.bundle.js?v=919C0428"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="8.9. Flow of Execution Summary" href="FlowofExecutionSummary.html" />
    <link rel="prev" title="8.7. Scope - Local vs Global Variables" href="LocalGlobalVariablesScope.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'COMP1000UofM';
    eBookConfig.basecourse = 'COMP1000UofM';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = false;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '6.0.9';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          COMP1000UofM
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=COMP1000UofM&page=FunctionComposition'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">8.7. </span>Scope - Local vs Global Variables" data-toggle="tooltip">
    <a href="LocalGlobalVariablesScope.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">8.9. </span>Flow of Execution Summary' data-toggle="tooltip" >              
    <a href="FlowofExecutionSummary.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="function-composition-functions-calling-functions">
<span id="index-0"></span><h1><span class="section-number">8.8. </span>Function Composition (functions calling functions)<a class="headerlink" href="#function-composition-functions-calling-functions" title="Permalink to this headline">¶</a></h1>
<p>It is important to understand that each of the functions we write can be used and called from other functions we
write. This is one of the most important ways that computer programmers take a large problem and break it down into a
group of smaller problems. This process of breaking a problem into smaller subproblems is called <strong>functional decomposition</strong>.</p>
<p>Here’s a simple example of functional decomposition using two functions. The first function called <code class="docutils literal notranslate"><span class="pre">square</span></code> simply
computes the square of a given number. The second function called <code class="docutils literal notranslate"><span class="pre">sum_of_squares</span></code> makes use of square to compute
the sum of three numbers that have been squared.</p>

<div class="runestone codelens ">
    <div class="cd_section" data-component="codelens" data-question_label="8.8.1">
        <div id=clens8_8_1_question class="ac_question">
        </div>
        <div class="pytutorVisualizer" id="clens8_8_1"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 8.8.1 (clens8_8_1)</span> </p>
    </div>

<script>
var clens8_8_1_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["clens8_8_1"] = {"code": "def square(x):\n    y = x * x\n    return y\n\ndef sum_of_squares(x,y,z):\n    a = square(x)\n    b = square(y)\n    c = square(z)\n\n    return a+b+c\n\na = -5\nb = 2\nc = 10\nresult = sum_of_squares(a,b,c)\nprint(result)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1]}, "ordered_globals": ["square"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null]}, "stdout": ""}, {"line": 12, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2]}, "ordered_globals": ["square", "sum_of_squares"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 13, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5}, "ordered_globals": ["square", "sum_of_squares", "a"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 14, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2}, "ordered_globals": ["square", "sum_of_squares", "a", "b"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 15, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 5, "event": "call", "func_name": "sum_of_squares", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10}, "ordered_varnames": ["x", "y", "z"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sum_of_squares_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "sum_of_squares", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10}, "ordered_varnames": ["x", "y", "z"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sum_of_squares_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10}, "ordered_varnames": ["x", "y", "z"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 2, "parent_frame_id_list": [], "encoded_locals": {"x": -5}, "ordered_varnames": ["x"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f2"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10}, "ordered_varnames": ["x", "y", "z"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 2, "parent_frame_id_list": [], "encoded_locals": {"x": -5}, "ordered_varnames": ["x"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f2"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10}, "ordered_varnames": ["x", "y", "z"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 2, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 25}, "ordered_varnames": ["x", "y"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f2"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 3, "event": "return", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10}, "ordered_varnames": ["x", "y", "z"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 2, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 25, "__return__": 25}, "ordered_varnames": ["x", "y", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f2"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 7, "event": "step_line", "func_name": "sum_of_squares", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25}, "ordered_varnames": ["x", "y", "z", "a"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sum_of_squares_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25}, "ordered_varnames": ["x", "y", "z", "a"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 3, "parent_frame_id_list": [], "encoded_locals": {"x": 2}, "ordered_varnames": ["x"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f3"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25}, "ordered_varnames": ["x", "y", "z", "a"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 3, "parent_frame_id_list": [], "encoded_locals": {"x": 2}, "ordered_varnames": ["x"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f3"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25}, "ordered_varnames": ["x", "y", "z", "a"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 3, "parent_frame_id_list": [], "encoded_locals": {"x": 2, "y": 4}, "ordered_varnames": ["x", "y"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f3"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 3, "event": "return", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25}, "ordered_varnames": ["x", "y", "z", "a"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 3, "parent_frame_id_list": [], "encoded_locals": {"x": 2, "y": 4, "__return__": 4}, "ordered_varnames": ["x", "y", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f3"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 8, "event": "step_line", "func_name": "sum_of_squares", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25, "b": 4}, "ordered_varnames": ["x", "y", "z", "a", "b"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sum_of_squares_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25, "b": 4}, "ordered_varnames": ["x", "y", "z", "a", "b"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 4, "parent_frame_id_list": [], "encoded_locals": {"x": 10}, "ordered_varnames": ["x"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f4"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25, "b": 4}, "ordered_varnames": ["x", "y", "z", "a", "b"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 4, "parent_frame_id_list": [], "encoded_locals": {"x": 10}, "ordered_varnames": ["x"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f4"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25, "b": 4}, "ordered_varnames": ["x", "y", "z", "a", "b"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 4, "parent_frame_id_list": [], "encoded_locals": {"x": 10, "y": 100}, "ordered_varnames": ["x", "y"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f4"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 3, "event": "return", "func_name": "square", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25, "b": 4}, "ordered_varnames": ["x", "y", "z", "a", "b"], "is_zombie": false, "is_highlighted": false, "unique_hash": "sum_of_squares_f1"}, {"func_name": "square", "is_parent": false, "frame_id": 4, "parent_frame_id_list": [], "encoded_locals": {"x": 10, "y": 100, "__return__": 100}, "ordered_varnames": ["x", "y", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f4"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 10, "event": "step_line", "func_name": "sum_of_squares", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25, "b": 4, "c": 100}, "ordered_varnames": ["x", "y", "z", "a", "b", "c"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sum_of_squares_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 10, "event": "return", "func_name": "sum_of_squares", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c"], "stack_to_render": [{"func_name": "sum_of_squares", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"x": -5, "y": 2, "z": 10, "a": 25, "b": 4, "c": 100, "__return__": 129}, "ordered_varnames": ["x", "y", "z", "a", "b", "c", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sum_of_squares_f1"}], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 16, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10, "result": 129}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": ""}, {"line": 16, "event": "return", "func_name": "<module>", "globals": {"square": ["REF", 1], "sum_of_squares": ["REF", 2], "a": -5, "b": 2, "c": 10, "result": 129}, "ordered_globals": ["square", "sum_of_squares", "a", "b", "c", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(x)", null], "2": ["FUNCTION", "sum_of_squares(x, y, z)", null]}, "stdout": "129\n"}]};
</script>
</div>
<p>Even though this is a pretty simple idea, in practice this example illustrates many very important Python concepts,
including local and global variables along with parameter passing. Note that the body of <code class="docutils literal notranslate"><span class="pre">square</span></code> is not executed
until it is called from inside the <code class="docutils literal notranslate"><span class="pre">sum_of_squares</span></code> function for the first time on line 6.</p>
<p>Also notice that when <code class="docutils literal notranslate"><span class="pre">square</span></code> is called (at Step 8, for example), there are two groups of local variables, one for
<code class="docutils literal notranslate"><span class="pre">square</span></code> and one for <code class="docutils literal notranslate"><span class="pre">sum_of_squares</span></code>.  Each group of local variables is called a <strong>stack frame</strong>. The variables
<code class="docutils literal notranslate"><span class="pre">x</span></code>, and <code class="docutils literal notranslate"><span class="pre">y</span></code> are local variables in both functions. These are completely different variables, even though they
have the same name. Each function invocation creates a new frame, and variables are looked up in that frame. Notice
that at step 9, y has the value 25 in one frame and 2 in the other.</p>
<p>What happens when you to refer to variable y on line 3? Python looks up the value of y in the stack frame for the
<code class="docutils literal notranslate"><span class="pre">square</span></code> function. If it didn’t find it there, it would go look in the global frame.</p>
<p>In the example below, we create a turtle drawing program that uses function composition.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac8_8_1 data-question_label="8.8.2">
<div id=ac8_8_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac8_8_1_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import turtle
import random

def random_col(tur):
     """ Sets turtle to a random color """
     red = random.random()
     green = random.random()
     blue = random.random()
     tur.color(red, green, blue)


def random_location(tur):
     """ Take turtle to a random location on canvas
     Pre-conditions: Assume canvas is 400 x 400 """
     x = random.randrange(-180, 180) # get random x location
     y = random.randrange(-180, 180) # get random y location
     tur.penup()
     tur.goto(x,y) # move to location without drawing
     tur.pendown()

def draw_triangle(tur, side_length):
    """ draws a triangle in current position and color"""
    for _ in range(3):
        tur.forward(side_length)
        tur.left(120)


def draw_square(tur, side_length):
    """ draws a square in current position and color"""
    for _ in range(4):
        tur.forward(side_length)
        tur.left(90)

def draw_design(tur, side_length):
    """ draws a square with two internal triangles """
    random_location(tur)
    random_col(tur)
    draw_square(tur, side_length)
    draw_triangle(tur, side_length)
    tur.forward(side_length)
    tur.left(90)
    tur.forward(side_length)
    tur.left(90)
    draw_triangle(tur, side_length)


wn = turtle.Screen()
yan = turtle.Turtle()
yan.speed(10)

for _ in range(5):
     side = random.randrange(30, 100)
     draw_design(yan, side)
</textarea>
</div>
</div>
<p>The turtle example above has 5 different functions. Two of them are <strong>helper functions</strong> that move a turtle to a random location and set a turtle to draw in a random color. The other three are drawing functions that cause the turtle to draw things. The global part of the code (lines 47-53) create the window and the turtle and sets the turtle speed to 10 so the turtle draws quickly. Then there is a for loop that iterates 10 times. Each time through the loop a new random size is generated and then the draw_design() function is called with the random size. The draw_design function calls the other four functions, ane even calls the draw_triangle function twice.</p>
<p><strong>Check your Understanding</strong></p>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/Users/clatulip/Documents/RunestoneBooks/ThinkLikeComputer/_sources/FunctionsParametersScope/FunctionComposition.rst</span>, line 127)</p>
<p>Duplicate ID – see FunctionsParametersScope/FunctionComposition, line 69</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: ac8_8_1
   :language: python
   :autograde: unittest
   :practice: T

   **1.** Write two functions, one called ``addit`` and one called ``mult``. ``addit`` takes one number as an input and adds 5. ``mult`` takes one number as an input, and multiplies that input by whatever is returned by ``addit``, and then returns the result.
   ~~~~

   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

      def testOne(self):
         self.assertEqual(mult(1), 6,&quot;Testing the function mult with input 1 (should be 6)&quot;)
         self.assertEqual(mult(-2), -6, &quot;Testing the function mult with input -2 (should be -6)&quot;)
         self.assertEqual(mult(0), 0, &quot;Testing the function mult with input 0 (should be 0)&quot;)

      def testTwo(self):
         self.assertEqual(addit(1), 6, &quot;Testing the function addit with input 1 (should be 6)&quot;)
         self.assertEqual(addit(-2), 3, &quot;Testing the function addit with input -2 (should be 3)&quot;)
         self.assertEqual(addit(0), 5, &quot;Testing the function addit with input 0 (should be 5)&quot;)

   myTests().main()
</pre></div>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 8.7. Scope - Local vs Global Variables' data-toggle="tooltip">
  <a href="LocalGlobalVariablesScope.html" aria-label="Previous - 8.7. Scope - Local vs Global Variables">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 8.9. Flow of Execution Summary' data-toggle="tooltip" >
    <a href="FlowofExecutionSummary.html" aria-label="Next - 8.9. Flow of Execution Summary">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 6.0.9.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>