<?xml version="1.0"?>
<section xml:id="python-swap_post-test">
  <title>Post Test</title>
  <p>Please answer
            the following problems to the best of your ability without any
            outside help. You can stop working on a problem after you worked
            on it for about five minutes without solving it.</p>
  <subsection xml:id="python-swap_problems">
    <title>Problems</title>
    <exercise label="ps-swap1-ac-v2">
      <statement>
        <p>Finish writing the code to swap the values in x and y (so that x ends up with y's initial value and y ends up with x's initial value).</p>
      </statement>
      <program label="ps-swap1-ac-v2_editor" interactive="activecode" language="python">
        <input>

x = 6
y = 2
temp = 0

# print the values
print(x)
print(y)

# swap the values of x and y
# write your code here

# print the values
print(x)
print(y)
        </input>
        <tests>
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(x, 2, "value of x after swap")
        self.assertEqual(y, 6, "value of y after swap")

myTests().main()
</tests>
      </program>
    </exercise>
    <exercise label="ps-swap2-ac">
      <statement>
        <p>Finish writing the code to swap the values in a and b (so that a ends up with b's initial value and b ends up with a's initial value).</p>
      </statement>
      <program label="ps-swap2-ac_editor" interactive="activecode" language="python">
        <input>

a = -3
b = 5
temp = 0

# print the values
print(a)
print(b)

# swap the values of a and b
# write your code here

# print the values
print(a)
print(b)
        </input>
        <tests>
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(a, 5, "value of a after swap")
        self.assertEqual(b, -3, "value of b after swap")

myTests().main()
</tests>
      </program>
    </exercise>
  </subsection>
  <subsection xml:id="python-swap_feedback">
    <title>Feedback</title>
    <exercise label="ps-posttest-sa">
      <statement>
        <p>Please provide feedback here. Please share any comments, problems, or suggestions.</p>
      </statement>
    </exercise>
  </subsection>
  <subsection xml:id="python-swap_what-to-do-next">
    <title>What to do next:</title>
    <raw format="html" xml:space="preserve">&lt;p&gt;Click on the following link to go to the post survey: &lt;b&gt;&lt;a id="ps-postsurvey"&gt;&lt;font size="+2"&gt;Post Survey&lt;/font&gt;&lt;/a&gt;&lt;/b&gt;&lt;/p&gt;</raw>
    <raw format="html" xml:space="preserve">&lt;script type="text/javascript" &gt;

  window.onload = function() {

    a = document.getElementById("ps-postsurvey")
    a.href = "ps-postsurvey.html"
  };

&lt;/script&gt;</raw>
  </subsection>
</section>
