<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html  lang="en" class="runestone-sphinx">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>8.15. Chapter Assessment &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=E2DC0A34" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=4182B24D" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-347.f9add1ca35d5ad93.css?v=306CBB19" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-runestone.f64dcc3632d7a0c3.css?v=F4F91F21" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script src="../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script src="../_static/sphinx_highlight.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/prefix-runtime.2438a8bdd6d43b5c.bundle.js?v=7D29D67B"></script>
    <script defer="" src="../_static/prefix-347.8328b60515045466.bundle.js?v=2CF09D3E"></script>
    <script defer="" src="../_static/prefix-runestone.aeac7cf9ad5c07ef.bundle.js?v=4FF65C30"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
<script src="https://media.ethicalads.io/media/client/ethicalads.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="9. Sequences" href="../Sequences/toctree.html" />
    <link rel="prev" title="8.14. Exercises" href="Exercises.html" />

  
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'FOPP-PIE';
    eBookConfig.basecourse = 'FOPP-PIE';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '7.2.11';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" id="rs-book-toc-link" href="../index.html" aria-label="index-page">
          
          FOPP-PIE
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="ad_warning"></span></li>
        <li><span id="browsing_warning"></span></li>

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://guide.runestone.academy' target="_blank" rel="noopener noreferrer" >Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=FOPP-PIE&page=ChapterAssessment'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">8.14. </span>Exercises" data-toggle="tooltip">
    <a href="Exercises.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">9. </span>Sequences' data-toggle="tooltip" >              
    <a href="../Sequences/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <section id="chapter-assessment">
<h1><span class="section-number">8.15. </span>Chapter Assessment<a class="headerlink" href="#chapter-assessment" title="Permalink to this heading">¶</a></h1>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac11_15_1 data-question_label="8.15.1">
<div id=ac11_15_1_question class="ac_question">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">int_return</span></code> that takes an integer as input and returns the same integer.</p>

</div>
<textarea data-lang="python" id="ac11_15_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(int_return(10), 10, "Testing that function int_return(10) returns 10")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac11_15_2 data-question_label="8.15.2">
<div id=ac11_15_2_question class="ac_question">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">add</span></code> that takes any number as its input and returns that sum with 2 added.</p>

</div>
<textarea data-lang="python" id="ac11_15_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testTwo(self):
      self.assertEqual(add(-2), 0, "Testing that add(-2) returns 0")
      self.assertEqual(add(6), 8, "Testing that add(6) returns 8")
      self.assertEqual(add(4), 6, "Testing that add(4) returns 6")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac11_15_3 data-question_label="8.15.3">
<div id=ac11_15_3_question class="ac_question">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">change</span></code> that takes any string, adds “Nice to meet you!” to the end of the argument given, and returns that new string.</p>

</div>
<textarea data-lang="python" id="ac11_15_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testThree(self):
      self.assertEqual(change("I'm Bob. "), "I'm Bob. Nice to meet you!", "Tests that change('I'm Bob. '') returns 'I'm Bob. Nice to meet you!'")
      self.assertEqual(change(""), "Nice to meet you!", "Tests that change() returns 'Nice to meet you!'")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac11_15_6 data-question_label="8.15.4">
<div id=ac11_15_6_question class="ac_question">
<p>You will need to write two functions for this problem. The first function, <code class="docutils literal notranslate"><span class="pre">divide</span></code> that takes in any number and returns that same number divided by 2. The second function called <code class="docutils literal notranslate"><span class="pre">sum</span></code> should take any number, divide it by 2, and add 6. It should return this new number. You should call the <code class="docutils literal notranslate"><span class="pre">divide</span></code> function within the <code class="docutils literal notranslate"><span class="pre">sum</span></code> function. Do not worry about decimals.</p>

</div>
<textarea data-lang="python" id="ac11_15_6_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testSixA(self):
      self.assertEqual(divide(4), 2, "Tests that divide(4) returns 2")
   def testSixB(self):
      self.assertEqual(sum(4), 8, "Tests that sum(4) returns 8")
      self.assertEqual(sum(2), 7, "Tests that sum(2) returns 7")
      self.assertEqual(sum(-6), 3, "Tests that sum(-6) returns 3")
      self.assertEqual(sum(0), 6, "Tests that sum(0) returns 6")

myTests().main()
</textarea>
</div>
</div>

        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="8.15.5" id="assess_ch8_1_lh"  style="visibility: hidden;">
            <p>:</p>
<a class="reference internal image-reference" href="../_images/Chapter_8_1.PNG"><img alt="Finding the maximum number" class="align-center" src="../_images/Chapter_8_1.PNG" style="width: 1200px;" /></a>
<p>The screenshot above depicts a CodeLens program mid-execution. We have a function called maximum() that is in the middle of its execution. The purpose of the function is to find the highest value in the list and then return it. Looking at the diagram, What should be the return value? <input type="text"></p>

        <script type="application/json">
            [[{"number": [97, 97], "feedback": "<p>Correct.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Incorrect. Our function is looking for and returning the highest element in the class_grades list</p>\n"}]]
        </script>

        </div>
        </div>
            
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="8.15.6" data-multipleanswers="true"  id="assess_ch8_2_lh"  style="visibility: hidden;">
    <p>:</p>
<a class="reference internal image-reference" href="../_images/Chapter_8_2.PNG"><img alt="local vs global variables" class="align-center" src="../_images/Chapter_8_2.PNG" style="width: 1200px;" /></a>
<p>Using the Codelens diagram screenshot, select all the variables that are local to the function generate_coords</p>

    <li data-component="answer"  id="assess_ch8_2_lh_opt_a">NUM_COORDS</li><li data-component="feedback">Incorrect, NUM_COORDS is under the global frame in the Codelens screenshot. Look at the generate_coords frame</li>
    
    <li data-component="answer" data-correct='yes' id="assess_ch8_2_lh_opt_b">MAX</li><li data-component="feedback">Correct, we are using the local version of MAX declared inside the function</li>
    
    <li data-component="answer" data-correct='yes' id="assess_ch8_2_lh_opt_c">MIN</li><li data-component="feedback">Correct, we are using the local version of MIN declared inside the function</li>
    
    <li data-component="answer"  id="assess_ch8_2_lh_opt_d">x_coord</li><li data-component="feedback">Incorrect, the x_coord append line is using a reference of the global frame list x_coord</li>
    
    <li data-component="answer" data-correct='yes' id="assess_ch8_2_lh_opt_e">coord</li><li data-component="feedback">Correct, the iterator variable is still a variable and in this case it is created locally inside the function</li>
    

    </ul>
    </div>
    
        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="8.15.7" id="assess_ch8_3_lh"  style="visibility: hidden;">
            <p>:</p>
<a class="reference internal image-reference" href="../_images/Chapter_8_4.PNG"><img alt="average function" class="align-center" src="../_images/Chapter_8_4.PNG" style="width: 1200px;" /></a>
<p>Consider the above program. Which Codelens screenshot shows the correct program output? (The options are labelled 1,2,3 in the top right corner of each screenshot) <input type="text"></p>
<a class="reference internal image-reference" href="../_images/Chapter_8_4_A.PNG"><img alt="average function output" class="align-center" src="../_images/Chapter_8_4_A.PNG" style="width: 900px;" /></a>

        <script type="application/json">
            [[{"number": [1, 1], "feedback": "<p>Correct.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Incorrect, take another look at the grades inside the list and the number of students</p>\n"}]]
        </script>

        </div>
        </div>
            </section>


  

  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 8.14. Exercises' data-toggle="tooltip">
  <a href="Exercises.html" aria-label="Previous - 8.14. Exercises">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 9. Sequences' data-toggle="tooltip" >
    <a href="../Sequences/toctree.html" aria-label="Next - 9. Sequences">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 7.2.11.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>