<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>8.13. Exercises &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.675cfab3e021b8f9.css?v=AC6C1FF6" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/runtime.15a489335c5c6109.bundle.js?v=8DB56D48"></script>
    <script defer="" src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script defer="" src="../_static/runestone.a1906624d1effde5.bundle.js?v=919C0428"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="8.14. Chapter Assessment" href="week3a1.html" />
    <link rel="prev" title="8.12. Glossary" href="Glossary.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'COMP1000UofM';
    eBookConfig.basecourse = 'COMP1000UofM';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = false;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '6.0.9';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          COMP1000UofM
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=COMP1000UofM&page=Exercises'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">8.12. </span>Glossary" data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">8.14. </span>Chapter Assessment' data-toggle="tooltip" >              
    <a href="week3a1.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <p></p>
<div class="section" id="exercises">
<h1><span class="section-number">8.13. </span>Exercises<a class="headerlink" href="#exercises" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='q1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac7_14_1 data-question_label="8.13.1">
<div id=ac7_14_1_question class="ac_question col-md-12">
<p>Write code that asks the user to enter a numeric score (0-100). In response, it should print out the score and corresponding letter grade, according to the table below.</p>
<table class="docutils align-default">
<colgroup>
<col style="width: 58%" />
<col style="width: 42%" />
</colgroup>
<thead>
<tr class="row-odd"><th class="head"><p>Score</p></th>
<th class="head"><p>Grade</p></th>
</tr>
</thead>
<tbody>
<tr class="row-even"><td><p>&gt;= 90</p></td>
<td><p>A</p></td>
</tr>
<tr class="row-odd"><td><p>[80-90)</p></td>
<td><p>B</p></td>
</tr>
<tr class="row-even"><td><p>[70-80)</p></td>
<td><p>C</p></td>
</tr>
<tr class="row-odd"><td><p>[60-70)</p></td>
<td><p>D</p></td>
</tr>
<tr class="row-even"><td><p>&lt; 60</p></td>
<td><p>F</p></td>
</tr>
</tbody>
</table>
<p>The square and round brackets denote closed and open intervals.
A closed interval includes the number, and open interval excludes it. So 79.99999 gets grade C , but 80 gets grade B.</p>

</div>
<textarea data-lang="python" id="ac7_14_1_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
scores = [77.51, 92.86, 98.01, 69.71, 78.52, 59.69, 60.49, 85.04, 87.33, 91.04]

====
from unittest.gui import TestCaseGui
import re
import math
class myTests(TestCaseGui):
    def testOne(self):
        output = self.getOutput().split('\n')
        editor = self.getEditorText().split('\n')
        float_re = r'[-+]?(\d+(\.\d*)?|\.\d+)([eE][-+]?\d+)?'
        _grades=[]
        for score in scores:
            _score = float(score)
            if _score >= 90:
                _grades.append('A')
            elif _score >= 80:
                _grades.append('B')
            elif _score >= 70:
                _grades.append('C')
            elif _score >= 60:
                _grades.append('D')
            else:
                _grades.append('F')

        self.assertEqual(grades, _grades)
        # hardcode check
        # LOOK FOR IF STATEMENTS
        outer_ifs = re.findall(r'^(if[ (].*: *)$', self.getEditorText(), re.M)
        outer_elifs = re.findall(r'^(elif[ (].*: *)$', self.getEditorText(), re.M)
        outer_elses = re.findall(r'^(else *: *)$', self.getEditorText(), re.M)
        inner_ifs = re.findall(r'^( +if[ (].*: *)$', self.getEditorText(), re.M)
        inner_elifs = re.findall(r'^( +elif[ (].*: *)$', self.getEditorText(), re.M)
        inner_elses = re.findall(r'^( +else *: *)$', self.getEditorText(), re.M)
        self.assertTrue(len(outer_ifs)==0 and len(outer_elifs)==0 and len(outer_elses)==0 and
                        len(inner_ifs)==1 and len(inner_elifs)==3 and len(inner_elses)==1,
                'Checking if-statements')
        # LOOK FOR for STATEMENTS
        outer_loops = re.findall(r'^(for[ (].* in.*: *)$', self.getEditorText(), re.M)
        inner_loops = re.findall(r'^( +for[ (].* in.*: *)$', self.getEditorText(), re.M)
        self.assertTrue(len(outer_loops)==1 and len(inner_loops)==0, 'Checking for-statements')
myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ans7_14_1 data-question_label="8.13.2">
<div id=ans7_14_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ans7_14_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
sc = input("Enter a score from 0 to 100 (decimal points are allowed)")
fl_sc = float(sc)

if fl_sc < 60:
    gr = "F"
elif fl_sc < 70:
    gr = "D"
elif fl_sc < 80:
    gr = "C"
elif fl_sc < 90:
    gr = "B"
else:
    gr = "A"

print("Score", fl_sc, "gets a grade of", gr)
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q2' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac7_14_2 data-question_label="8.13.3">
<div id=ac7_14_2_question class="ac_question col-md-12">
<p>A year is a <strong>leap year</strong> if it is divisible by 4; however, if the year can be evenly divided by 100, it is NOT a leap year, unless the year is <strong>also</strong> evenly divisible by 400 then it is a leap year. Write code that asks the user to input a year and output True if it’s a leap year, or False otherwise. Use if statements.</p>
<table class="docutils align-default">
<colgroup>
<col style="width: 58%" />
<col style="width: 42%" />
</colgroup>
<thead>
<tr class="row-odd"><th class="head"><p>Year</p></th>
<th class="head"><p>Leap?</p></th>
</tr>
</thead>
<tbody>
<tr class="row-even"><td><p>1944</p></td>
<td><p>True</p></td>
</tr>
<tr class="row-odd"><td><p>2011</p></td>
<td><p>False</p></td>
</tr>
<tr class="row-even"><td><p>1986</p></td>
<td><p>False</p></td>
</tr>
<tr class="row-odd"><td><p>1800</p></td>
<td><p>False</p></td>
</tr>
<tr class="row-even"><td><p>1900</p></td>
<td><p>False</p></td>
</tr>
<tr class="row-odd"><td><p>2000</p></td>
<td><p>True</p></td>
</tr>
<tr class="row-even"><td><p>2056</p></td>
<td><p>True</p></td>
</tr>
</tbody>
</table>
<p>Above are some examples of what the output should be for various inputs.</p>

</div>
<textarea data-lang="python" id="ac7_14_2_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

years = [1967, 1900, 1400, 1628, 1701, 1217, 1359, 1300, 2000, 1054,
1724, 1000, 1800, 1100, 2100, 1023, 1600, 1500, 1358, 1160,
1700, 1744, 2009, 1200]

====
from unittest.gui import TestCaseGui
import re
import math
class myTests(TestCaseGui):
    def testOne(self):
        output = self.getOutput().split('\n')
        editor = self.getEditorText().split('\n')
        float_re = r'[-+]?(\d+(\.\d*)?|\.\d+)([eE][-+]?\d+)?'

        _is_leap_year = []
        for _year in years:
            if _year % 4 == 0 :
                if _year % 100 == 0:
                    _is_leap_year.append( _year % 400 == 0)
                else:
                    _is_leap_year.append(True)
            else:
                _is_leap_year.append(False)

        self.assertEqual(is_leap_year, _is_leap_year, 'Checking answer')

        # hardcode check
        float_re = r'[-+]?(\d+(\.\d*)?|\.\d+)([eE][-+]?\d+)?'
        print_float_re = r'print\( *'+float_re+' *\)'
        self.assertFalse(re.search(print_float_re, self.getEditorText()), 'Checking for hardcoding')
        # LOOK FOR IF STATEMENTS
        outer_ifs = re.findall(r'^(if[ (].*: *)$', self.getEditorText(), re.M)
        outer_elifs = re.findall(r'^(elif[ (].*: *)$', self.getEditorText(), re.M)
        outer_elses = re.findall(r'^(else *: *)$', self.getEditorText(), re.M)
        inner_ifs = re.findall(r'^( +if[ (].*: *)$', self.getEditorText(), re.M)
        inner_elifs = re.findall(r'^( +elif[ (].*: *)$', self.getEditorText(), re.M)
        inner_elses = re.findall(r'^( +else *: *)$', self.getEditorText(), re.M)
        self.assertTrue(len(outer_ifs)==0 and len(outer_elifs)==0 and len(outer_elses)==0 and len(inner_ifs)>=1 and len(inner_elifs)>=0 and len(inner_elses)>=0, 'Checking if-statements')
        # LOOK FOR for STATEMENTS
        outer_loops = re.findall(r'^(for[ (].* in.*: *)$', self.getEditorText(), re.M)
        inner_loops = re.findall(r'^( +for[ (].* in.*: *)$', self.getEditorText(), re.M)
        self.assertTrue(len(outer_loops)==1 and len(inner_loops)==0, 'Checking for-statements')
myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q3' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac7_14_3 data-question_label="8.13.4">
<div id=ac7_14_3_question class="ac_question col-md-12">
<p>What do these expressions evaluate to?</p>
<ol class="arabic simple">
<li><p><code class="docutils literal notranslate"><span class="pre">3</span> <span class="pre">==</span> <span class="pre">3</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">3</span> <span class="pre">!=</span> <span class="pre">3</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">3</span> <span class="pre">&gt;=</span> <span class="pre">4</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">not</span> <span class="pre">(3</span> <span class="pre">&lt;</span> <span class="pre">4)</span></code></p></li>
</ol>

</div>
<textarea data-lang="python" id="ac7_14_3_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >
<ol class="arabic simple">
<li><p>True</p></li>
<li><p>False</p></li>
<li><p>False</p></li>
<li><p>False</p></li>
</ol>
</div>
    </div>

</li>
<li><div id='q4' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac7_14_4 data-question_label="8.13.5">
<div id=ac7_14_4_question class="ac_question col-md-12">
<p>Give the <strong>logical opposites</strong> of these conditions, meaning an expression that would
produce False whenever this expression produces True, and vice versa.  You are not
allowed to use the <code class="docutils literal notranslate"><span class="pre">not</span></code> operator.</p>
<ol class="arabic simple">
<li><p><code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">&gt;</span> <span class="pre">b</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">&gt;=</span> <span class="pre">b</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">&gt;=</span> <span class="pre">18</span>&#160; <span class="pre">and</span>&#160; <span class="pre">day</span> <span class="pre">==</span> <span class="pre">3</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">&gt;=</span> <span class="pre">18</span>&#160; <span class="pre">or</span>&#160; <span class="pre">day</span> <span class="pre">!=</span> <span class="pre">3</span></code></p></li>
</ol>

</div>
<textarea data-lang="python" id="ac7_14_4_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q5' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac7_14_5 data-question_label="8.13.6">
<div id=ac7_14_5_question class="ac_question col-md-12">
<p>Provided are the lengths of two sides of a right-angled triangle. Assign the length of the hypotenuse the the variable <code class="docutils literal notranslate"><span class="pre">hypo_len</span></code>.  (Hint:  <code class="docutils literal notranslate"><span class="pre">x</span> <span class="pre">**</span> <span class="pre">0.5</span></code> will return the square root, or use <code class="docutils literal notranslate"><span class="pre">sqrt</span></code> from the math module)</p>

</div>
<textarea data-lang="python" id="ac7_14_5_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

side1 = 3
side2 = 4

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(hypo_len,5,"Testing that hypo_len has been set correctly")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q6' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac7_14_6 data-question_label="8.13.7">
<div id=ac7_14_6_question class="ac_question col-md-12">
<p>Provided is a list of numbers. For each of the numbers in the list, determine whether they are even. If the number is even, add <code class="docutils literal notranslate"><span class="pre">True</span></code> to a new list called <code class="docutils literal notranslate"><span class="pre">is_even</span></code>. If the number is odd, then add <code class="docutils literal notranslate"><span class="pre">False</span></code>.</p>

</div>
<textarea data-lang="python" id="ac7_14_6_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
num_lst = [3, 20, -1, 9, 10]

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
     def testOne(self):
         self.assertEqual(is_even, [False, True, False, False, True],"Testing that is_even is set correctly.")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q7' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac7_14_7 data-question_label="8.13.8">
<div id=ac7_14_7_question class="ac_question col-md-12">
<p>Provided is a list of numbers. For each of the numbers in the list, determine whether they are odd. If the number is odd, add <code class="docutils literal notranslate"><span class="pre">True</span></code> to a new list called <code class="docutils literal notranslate"><span class="pre">is_odd</span></code>. If the number is even, then add <code class="docutils literal notranslate"><span class="pre">False</span></code>.</p>

</div>
<textarea data-lang="python" id="ac7_14_7_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
num_lst = [3, 20, -1, 9, 10]


====
from unittest.gui import TestCaseGui
import re

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(is_odd, [True, False, True, True, False],"Testing that is_odd is set correctly.")

        # LOOK FOR IF STATEMENTS
        outer_ifs = re.findall(r'^(if[ (].*: *)$', self.getEditorText(), re.M)
        outer_elifs = re.findall(r'^(elif[ (].*: *)$', self.getEditorText(), re.M)
        outer_elses = re.findall(r'^(else *: *)$', self.getEditorText(), re.M)
        inner_ifs = re.findall(r'^( +if[ (].*: *)$', self.getEditorText(), re.M)
        inner_elifs = re.findall(r'^( +elif[ (].*: *)$', self.getEditorText(), re.M)
        inner_elses = re.findall(r'^( +else *: *)$', self.getEditorText(), re.M)
        self.assertTrue(len(outer_ifs)==0 and len(outer_elifs)==0 and len(outer_elses)==0 and len(inner_ifs)==1 and len(inner_elifs)==0 and len(inner_elses)==1, 'Checking if-statements')

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q8' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac7_14_8 data-question_label="8.13.9">
<div id=ac7_14_8_question class="ac_question col-md-12">
<p>Given the lengths of three sides of a triange, determine whether the triangle is right angled. If it is, the assign <code class="docutils literal notranslate"><span class="pre">True</span></code> to the variable <code class="docutils literal notranslate"><span class="pre">is_rightangled</span></code>. If it’s not, then assign <code class="docutils literal notranslate"><span class="pre">False</span></code> to the variable <code class="docutils literal notranslate"><span class="pre">is_rightangled</span></code>.</p>
<p>Hint: floating point arithmetic is not always exactly accurate,
so it is not safe to test floating point numbers for equality.
If a good programmer wants to know whether
<code class="docutils literal notranslate"><span class="pre">x</span></code> is equal or close enough to <code class="docutils literal notranslate"><span class="pre">y</span></code>, they would probably code it up as</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">if</span>  <span class="nb">abs</span><span class="p">(</span><span class="n">x</span> <span class="o">-</span> <span class="n">y</span><span class="p">)</span> <span class="o">&lt;</span> <span class="mf">0.001</span><span class="p">:</span>      <span class="c1"># if x is approximately equal to y</span>
    <span class="o">...</span>
</pre></div>
</div>

</div>
<textarea data-lang="python" id="ac7_14_8_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
a = 5
b = 6
c = 8

====
from unittest.gui import TestCaseGui
import re
import math
class myTests(TestCaseGui):
    def testOne(self):
        output = self.getOutput().split('\n')
        editor = self.getEditorText().split('\n')
        float_re = r'[-+]?(\d+(\.\d*)?|\.\d+)([eE][-+]?\d+)?'

        _is_rightangled = False
        if abs(a**2 + b**2 - c**2) < 1e-7:
            _is_rightangled = True
        elif abs(a**2 - b**2 + c**2) < 1e-7:
            _is_rightangled = True
        elif abs(-a**2 + b**2 + c**2) < 1e-7:
            _is_rightangled = True

        self.assertEqual(is_rightangled, _is_rightangled, 'Checking answer')

        # hardcode check
        float_re = r'[-+]?(\d+(\.\d*)?|\.\d+)([eE][-+]?\d+)?'
        print_float_re = r'print\( *'+float_re+' *\)'
        self.assertFalse(re.search(print_float_re, self.getEditorText()), 'Checking for hardcoding')
            # LOOK FOR IF STATEMENTS
        outer_ifs = re.findall(r'^(if[ (].*: *)$', self.getEditorText(), re.M)
        outer_elifs = re.findall(r'^(elif[ (].*: *)$', self.getEditorText(), re.M)
        outer_elses = re.findall(r'^(else *: *)$', self.getEditorText(), re.M)
        inner_ifs = re.findall(r'^( +if[ (].*: *)$', self.getEditorText(), re.M)
        inner_elifs = re.findall(r'^( +elif[ (].*: *)$', self.getEditorText(), re.M)
        inner_elses = re.findall(r'^( +else *: *)$', self.getEditorText(), re.M)
        self.assertTrue(len(outer_ifs)>=1 and len(outer_elifs)>=0 and len(outer_elses)>=0 and
                        len(inner_ifs)>=0 and len(inner_elifs)>=0 and len(inner_elses)>=0,
                'Checking if-statements')
        # LOOK FOR for STATEMENTS
        outer_loops = re.findall(r'^(for[ (].* in.*: *)$', self.getEditorText(), re.M)
        inner_loops = re.findall(r'^( +for[ (].* in.*: *)$', self.getEditorText(), re.M)
        self.assertTrue(len(outer_loops)>=0 and len(inner_loops)>=0, 'Checking for-statements')
myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q9' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac7_14_9 data-question_label="8.13.10">
<div id=ac7_14_9_question class="ac_question col-md-12">
<p>Implement the calculator for the date of Easter.</p>
<p>The following algorithm computes the date for Easter Sunday for any year between 1900 to 2099.</p>
<p>Ask the user to enter a year.
Compute the following:</p>
<blockquote>
<div><ol class="arabic simple">
<li><p>a = year % 19</p></li>
<li><p>b = year % 4</p></li>
<li><p>c = year % 7</p></li>
<li><p>d = (19 * a + 24) % 30</p></li>
<li><p>e = (2 * b + 4 * c + 6 * d + 5) % 7</p></li>
<li><p>dateofeaster = 22 + d + e</p></li>
</ol>
</div></blockquote>
<p>Special note: The algorithm can give a date in April.  You will know that the date is in April if the calculation gives you an answer greater than 31.  (You’ll need to adjust)  Also, if the year is one of four special
years (1954, 1981, 2049, or 2076) then subtract 7 from the date.</p>
<p>Your program should print an error message if the user provides a date that is out of range.</p>

</div>
<textarea data-lang="python" id="ac7_14_9_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=answer_ex_6_13 data-question_label="8.13.11">
<div id=answer_ex_6_13_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="answer_ex_6_13_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
year = int(input("Please enter a year"))
if year >= 1900 and year <= 2099:
    a = year % 19
    b = year % 4
    c = year % 7
    d = (19*a + 24) % 30
    e = (2*b + 4*c + 6*d + 5) % 7
    dateofeaster = 22 + d + e

    if year == 1954 or year == 2981 or year == 2049 or year == 2076:
        dateofeaster = dateofeaster - 7

    if dateofeaster > 31:
        print("April", dateofeaster - 31)
    else:
        print("March", dateofeaster)
else:
    print("ERROR...year out of range")
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function disqus_2dfd6acf1ca849c2853dad606d1ba255_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_2dfd6acf1ca849c2853dad606d1ba255_' + eBookConfig.course;
                    this.page.url = 'https://disqus_2dfd6acf1ca849c2853dad606d1ba255_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_2dfd6acf1ca849c2853dad606d1ba255_' + eBookConfig.course;
            disqus_url = 'http://disqus_2dfd6acf1ca849c2853dad606d1ba255_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="disqus_2dfd6acf1ca849c2853dad606d1ba255" onclick="disqus_2dfd6acf1ca849c2853dad606d1ba255_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='disqus_2dfd6acf1ca849c2853dad606d1ba255']").attr('data-disqus-identifier', 'disqus_2dfd6acf1ca849c2853dad606d1ba255_' + eBookConfig.course);
  });
</script>
</div>
    </div>

</li>
<li><div id='q9' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac7_14_10 data-question_label="8.13.12">
<div id=ac7_14_10_question class="ac_question col-md-12">
<p>Get the user to enter some text and print out True if it’s a palindrome, False otherwise. (Hint: Start by reversing the input string, and then use the == operator to compare two values to see if they are the same)</p>

</div>
<textarea data-lang="python" id="ac7_14_10_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</div>
    </div>

</li>
<li>
        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="pp7_14_11" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write a program that will print out a greeting to each student in the list. This list should also keep track of how many students have been greeted and note that each time a new student has been greeted. When only one student has entered, the program should say “The first student has entered!”. Afterwards, the program should say “There are {number here} students in the classroom!”.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="8.13.13"         style="visibility: hidden;">
        students = ["Jay", "Stacy", "Iman", "Trisha", "Ahmed", "Daniel", "Shadae", "Tosin", "Charlotte"]
---
num_students = 0
---
for student in students:
---
    print("Welcome to class, " + student)
    num_students += 1
---
    if num_students == 1:
        print("The first student has entered!")
---
    elif num_students > 1:
        print("There are " + str(num_students) + " students in the classroom!")
        </pre>
        </div>
        </div>
    </li>
<li>
        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="pp7_14_12" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Piece together a program so that it can successfully print out one print statement, given the value of x.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="8.13.14"         style="visibility: hidden;">
        x = 16
---
if x > 10:
---
    if x > 20:
        print("This is a large number!")
---
    else:
        print("This is a pretty big number.")
        </pre>
        </div>
        </div>
    </li>
</ol>
<div class="section" id="contributed-exercises">
<h2><span class="section-number">8.13.1. </span>Contributed Exercises<a class="headerlink" href="#contributed-exercises" title="Permalink to this headline">¶</a></h2>
{% for q in questions: %}
    <div class='oneq full-width'>
        {{ q['htmlsrc']|safe }}
    </div>
{% endfor %}</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 8.12. Glossary' data-toggle="tooltip">
  <a href="Glossary.html" aria-label="Previous - 8.12. Glossary">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 8.14. Chapter Assessment' data-toggle="tooltip" >
    <a href="week3a1.html" aria-label="Next - 8.14. Chapter Assessment">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 6.0.9.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>