<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html  lang="en" class="runestone-sphinx">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>6.9. Chained conditionals &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=E2DC0A34" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=4182B24D" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-347.f9add1ca35d5ad93.css?v=306CBB19" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-runestone.f64dcc3632d7a0c3.css?v=F4F91F21" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script src="../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script src="../_static/sphinx_highlight.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/prefix-runtime.2438a8bdd6d43b5c.bundle.js?v=7D29D67B"></script>
    <script defer="" src="../_static/prefix-347.8328b60515045466.bundle.js?v=2CF09D3E"></script>
    <script defer="" src="../_static/prefix-runestone.aeac7cf9ad5c07ef.bundle.js?v=4FF65C30"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
<script src="https://media.ethicalads.io/media/client/ethicalads.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="6.10. The Accumulator Pattern with Conditionals" href="TheAccumulatorPatternwithConditionals.html" />
    <link rel="prev" title="6.8. Nested conditionals" href="Nestedconditionals.html" />

  
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'FOPP-PIE';
    eBookConfig.basecourse = 'FOPP-PIE';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '7.2.11';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" id="rs-book-toc-link" href="../index.html" aria-label="index-page">
          
          FOPP-PIE
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="ad_warning"></span></li>
        <li><span id="browsing_warning"></span></li>

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://guide.runestone.academy' target="_blank" rel="noopener noreferrer" >Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=FOPP-PIE&page=Chainedconditionals'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">6.8. </span>Nested conditionals" data-toggle="tooltip">
    <a href="Nestedconditionals.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">6.10. </span>The Accumulator Pattern with Conditionals' data-toggle="tooltip" >              
    <a href="TheAccumulatorPatternwithConditionals.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <section id="chained-conditionals">
<span id="index-0"></span><h1><span class="section-number">6.9. </span>Chained conditionals<a class="headerlink" href="#chained-conditionals" title="Permalink to this heading">¶</a></h1>
<p>Python provides an alternative way to write nested selection such as the one shown in the previous section.
This is sometimes referred to as a <strong>chained conditional</strong>.</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">if</span> <span class="n">x</span> <span class="o">&lt;</span> <span class="n">y</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;x is less than y&quot;</span><span class="p">)</span>
<span class="k">elif</span> <span class="n">x</span> <span class="o">&gt;</span> <span class="n">y</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;x is greater than y&quot;</span><span class="p">)</span>
<span class="k">else</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;x and y must be equal&quot;</span><span class="p">)</span>
</pre></div>
</div>
<p>The flow of control can be drawn in a different orientation but the resulting pattern is identical to the one shown above.</p>
<img alt="../_images/flowchart_chained_conditional.png" src="../_images/flowchart_chained_conditional.png" />
<p><code class="docutils literal notranslate"><span class="pre">elif</span></code> is an abbreviation of <code class="docutils literal notranslate"><span class="pre">else</span> <span class="pre">if</span></code>. Again, exactly one branch will be
executed. There is no limit of the number of <code class="docutils literal notranslate"><span class="pre">elif</span></code> statements but only a
single (and optional) final <code class="docutils literal notranslate"><span class="pre">else</span></code> statement is allowed and it must be the last
branch in the statement.</p>
<img alt="../_images/conditionals_overview.png" src="../_images/conditionals_overview.png" />
<p>Each condition is checked in order. If the first is false, the next is checked,
and so on. If one of them is true, the corresponding branch executes, and the
statement ends. Even if more than one condition is true, only the first true
branch executes.</p>
<p>Here is the same program using <code class="docutils literal notranslate"><span class="pre">elif</span></code>.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac6_9_1 data-question_label="6.9.1">
<div id=ac6_9_1_question class="ac_question">

</div>
<textarea data-lang="python" id="ac6_9_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
x = 10
y = 10

if x < y:
    print("x is less than y")
elif x > y:
    print("x is greater than y")
else:
    print("x and y must be equal")
</textarea>
</div>
</div>
<p>The following image highlights different kinds of valid conditionals that can be used. Though there are other
versions of conditionals that Python can understand (imagine an if statement with twenty elif statements), those
other versions must follow the same order as seen below.</p>
<img alt="shows a unary conditiona, a binary conditional, a conditional with if, elif, else, and a conditional with if, elif, and elif." src="../_images/valid_conditionals.png" />
<p>Here is a turtle example that uses chained conditionals.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac6_9_2 data-question_label="6.9.2">
<div id=ac6_9_2_question class="ac_question">

</div>
<textarea data-lang="python" id="ac6_9_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import turtle
wn = turtle.Screen()

amy = turtle.Turtle()
amy.right(170)

colors = ["Pink", "Purple", "Yellow", "Orange", "Pink", "Orange", "Yellow", "Purple", "Orange", "Pink", "Pink", "Orange", "Yellow", "Purple", "Orange", "Purple", "Yellow", "Orange", "Pink", "Orange", "Purle", "Purple", "Yellow", "Orange", "Pink", "Blue", "Orange", "Yellow", "Purple"]


for color in colors:
    amy.pencolor(color)

    if amy.pencolor() == "Purple":
        amy.forward(50)
        amy.right(59)
    elif amy.pencolor() == "Yellow":
        amy.forward(65)
        amy.left(98)
    elif amy.pencolor() == "Orange":
        amy.forward(30)
        amy.left(60)
    elif amy.pencolor() == "Pink":
        amy.forward(50)
        amy.right(57)
    else:
        print("Uh-oh! I don't know the color: ", color)

amy.hideturtle()
</textarea>
</div>
</div>
<p>The above example combines a for loop with a set of conditional statements. Here, we loop through a list of
colors and each iteration checks to see what amy’s pen color is. Depending on the pen color, the turtle will
move in a certain direction, for a certain distance. This example also demonstrates a common pattern for chained
conditionals in which all of the expected things are covered by the various <code class="docutils literal notranslate"><span class="pre">elif</span></code> branches, and the <code class="docutils literal notranslate"><span class="pre">else</span></code>
is used to notify the programmer with a console message that something unexpected has occurred. In this case, we
expected only the four colours purple, yellow, orange, and pink. If we get something else, we print a message out
(and typically print out the unexpected item), which in this case helps us notice two things: there is a typo and
“Purple” is mispelled in one spot, and there is a “Blue” in the list, and so maybe we want to add code to handle
that color. You may think that adding such an else clause is unnecessary because you are sure your code works correctly,
but it’s helpful more often than you think.</p>
<p><strong>Check your understanding</strong></p>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="6.9.3" data-multipleanswers="false"  id="question6_9_1"  style="visibility: hidden;">
    <p>Which of I, II, and III below gives the same result as the following nested if?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="c1"># nested if-else statement</span>
<span class="n">x</span> <span class="o">=</span> <span class="o">-</span><span class="mi">10</span>
<span class="k">if</span> <span class="n">x</span> <span class="o">&lt;</span> <span class="mi">0</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;The negative number &quot;</span><span class="p">,</span>  <span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is not valid here.&quot;</span><span class="p">)</span>
<span class="k">else</span><span class="p">:</span>
    <span class="k">if</span> <span class="n">x</span> <span class="o">&gt;</span> <span class="mi">0</span><span class="p">:</span>
        <span class="nb">print</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is a positive number&quot;</span><span class="p">)</span>
    <span class="k">else</span><span class="p">:</span>
        <span class="nb">print</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is 0&quot;</span><span class="p">)</span>
</pre></div>
</div>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">I</span><span class="o">.</span>

<span class="k">if</span> <span class="n">x</span> <span class="o">&lt;</span> <span class="mi">0</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;The negative number &quot;</span><span class="p">,</span>  <span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is not valid here.&quot;</span><span class="p">)</span>
<span class="k">else</span> <span class="p">(</span><span class="n">x</span> <span class="o">&gt;</span> <span class="mi">0</span><span class="p">):</span>
    <span class="nb">print</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is a positive number&quot;</span><span class="p">)</span>
<span class="k">else</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is 0&quot;</span><span class="p">)</span>
</pre></div>
</div>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">II</span><span class="o">.</span>

<span class="k">if</span> <span class="n">x</span> <span class="o">&lt;</span> <span class="mi">0</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;The negative number &quot;</span><span class="p">,</span>  <span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is not valid here.&quot;</span><span class="p">)</span>
<span class="k">elif</span> <span class="p">(</span><span class="n">x</span> <span class="o">&gt;</span> <span class="mi">0</span><span class="p">):</span>
    <span class="nb">print</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is a positive number&quot;</span><span class="p">)</span>
<span class="k">else</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is 0&quot;</span><span class="p">)</span>
</pre></div>
</div>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">III</span><span class="o">.</span>

<span class="k">if</span> <span class="n">x</span> <span class="o">&lt;</span> <span class="mi">0</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;The negative number &quot;</span><span class="p">,</span>  <span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is not valid here.&quot;</span><span class="p">)</span>
<span class="k">if</span> <span class="p">(</span><span class="n">x</span> <span class="o">&gt;</span> <span class="mi">0</span><span class="p">):</span>
    <span class="nb">print</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is a positive number&quot;</span><span class="p">)</span>
<span class="k">else</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is 0&quot;</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="question6_9_1_opt_a">I only</li><li data-component="feedback">You can not use a Boolean expression after an else.</li>
    
    <li data-component="answer" data-correct='yes' id="question6_9_1_opt_b">II only</li><li data-component="feedback">Yes, II will give the same result.</li>
    
    <li data-component="answer"  id="question6_9_1_opt_c">III only</li><li data-component="feedback">No, III will not give the same result.  The first if statement will be true, but the second will be false, so the else part will execute.</li>
    
    <li data-component="answer"  id="question6_9_1_opt_d">II and III</li><li data-component="feedback">No, Although II is correct III will not give the same result.  Try it.</li>
    
    <li data-component="answer"  id="question6_9_1_opt_e">I, II, and III</li><li data-component="feedback">No, in I you can not have a Boolean expression after an else.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="6.9.4" data-multipleanswers="false"  id="question6_9_2"  style="visibility: hidden;">
    <p>What will the following code print if x = 3, y = 5, and z = 2?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">if</span> <span class="n">x</span> <span class="o">&lt;</span> <span class="n">y</span> <span class="ow">and</span> <span class="n">x</span> <span class="o">&lt;</span> <span class="n">z</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;a&quot;</span><span class="p">)</span>
<span class="k">elif</span> <span class="n">y</span> <span class="o">&lt;</span> <span class="n">x</span> <span class="ow">and</span> <span class="n">y</span> <span class="o">&lt;</span> <span class="n">z</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;b&quot;</span><span class="p">)</span>
<span class="k">else</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;c&quot;</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="question6_9_2_opt_a">a</li><li data-component="feedback">While the value in x is less than the value in y (3 is less than 5) it is not less than the value in z (3 is not less than 2).</li>
    
    <li data-component="answer"  id="question6_9_2_opt_b">b</li><li data-component="feedback">The value in y is not less than the value in x (5 is not less than 3).</li>
    
    <li data-component="answer" data-correct='yes' id="question6_9_2_opt_c">c</li><li data-component="feedback">Since the first two Boolean expressions are false the else will be executed.</li>
    

    </ul>
    </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac6_9_3 data-question_label="6.9.5">
<div id=ac6_9_3_question class="ac_question">
<p>Create one conditional to find whether “false” is in string <code class="docutils literal notranslate"><span class="pre">str1</span></code>. If so, assign variable <code class="docutils literal notranslate"><span class="pre">output</span></code> the string “False. You aren’t you?”. Check to see if “true” is in string <code class="docutils literal notranslate"><span class="pre">str1</span></code> and if it is then assign “True! You are you!” to the variable <code class="docutils literal notranslate"><span class="pre">output</span></code>. If neither are in <code class="docutils literal notranslate"><span class="pre">str1</span></code>, assign “Neither true nor false!” to <code class="docutils literal notranslate"><span class="pre">output</span></code>.</p>

</div>
<textarea data-lang="python" id="ac6_9_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
str1 = "Today you are you! That is truer than true! There is no one alive who is you-er than you!"

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testThree(self):
      self.assertEqual(output, "True! You are you!", "Testing that output has the correct value, given the str1 provided.")
      self.assertIn("else", self.getEditorText(), "Testing output (Don't worry about actual and expected values).")
      self.assertIn("elif", self.getEditorText(), "Testing output (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac6_9_4 data-question_label="6.9.6">
<div id=ac6_9_4_question class="ac_question">
<p>Create an empty list called <code class="docutils literal notranslate"><span class="pre">resps</span></code>. Using the list <code class="docutils literal notranslate"><span class="pre">percent_rain</span></code>, for each percent, if it is above 90, add the string ‘Bring an umbrella.’ to <code class="docutils literal notranslate"><span class="pre">resps</span></code>, otherwise if it is above 80, add the string ‘Good for the flowers?’ to <code class="docutils literal notranslate"><span class="pre">resps</span></code>, otherwise if it is above 50, add the string ‘Watch out for clouds!’ to <code class="docutils literal notranslate"><span class="pre">resps</span></code>, otherwise, add the string ‘Nice day!’ to <code class="docutils literal notranslate"><span class="pre">resps</span></code>. Note: if you’re sure you’ve got the problem right but it doesn’t pass, then check that you’ve matched up the strings exactly.</p>

</div>
<textarea data-lang="python" id="ac6_9_4_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
percent_rain = [94.3, 45, 100, 78, 16, 5.3, 79, 86]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(resps, ['Bring an umbrella.','Nice day!','Bring an umbrella.','Watch out for clouds!',"Nice day!",'Nice day!','Watch out for clouds!',"Good for the flowers?"], "Testing the value of resps")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac6_9_5 data-question_label="6.9.7">
<div id=ac6_9_5_question class="ac_question">
<p>We have created conditionals for you to use. Do not change the provided conditional statements. Find an integer value for <code class="docutils literal notranslate"><span class="pre">x</span></code> that will cause <code class="docutils literal notranslate"><span class="pre">output</span></code> to hold the values <code class="docutils literal notranslate"><span class="pre">True</span></code> and <code class="docutils literal notranslate"><span class="pre">None</span></code>. (Drawing diagrams or flow charts for yourself may help!)</p>

</div>
<textarea data-lang="python" id="ac6_9_5_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
x =
output = []

if x > 63:
    output.append(True)
elif x > 55:
    output.append(False)
else:
    output.append("Neither")

if x > 67:
    output.append(True)
else:
    output.append(None)

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testSixA(self):
      self.assertEqual(output, [True, None], "Testing that value of output is correct.")

   def testSixB(self):
      self.assertEqual(x in [64, 65, 66, 67], True, "Testing that value of x is reasonable for this problem")

myTests().main()
</textarea>
</div>
</div>
</section>


  

  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 6.8. Nested conditionals' data-toggle="tooltip">
  <a href="Nestedconditionals.html" aria-label="Previous - 6.8. Nested conditionals">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 6.10. The Accumulator Pattern with Conditionals' data-toggle="tooltip" >
    <a href="TheAccumulatorPatternwithConditionals.html" aria-label="Next - 6.10. The Accumulator Pattern with Conditionals">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 7.2.11.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>