<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html  lang="en" class="runestone-sphinx">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>6.14. Chapter Assessment &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=E2DC0A34" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=4182B24D" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-347.f9add1ca35d5ad93.css?v=306CBB19" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-runestone.f64dcc3632d7a0c3.css?v=F4F91F21" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script src="../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script src="../_static/sphinx_highlight.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/prefix-runtime.2438a8bdd6d43b5c.bundle.js?v=7D29D67B"></script>
    <script defer="" src="../_static/prefix-347.8328b60515045466.bundle.js?v=2CF09D3E"></script>
    <script defer="" src="../_static/prefix-runestone.aeac7cf9ad5c07ef.bundle.js?v=4FF65C30"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
<script src="https://media.ethicalads.io/media/client/ethicalads.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="7. Integration Pause" href="../IntegrationPause/toctree.html" />
    <link rel="prev" title="6.13. Exercises" href="Exercises.html" />

  
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'FOPP-PIE';
    eBookConfig.basecourse = 'FOPP-PIE';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '7.2.11';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" id="rs-book-toc-link" href="../index.html" aria-label="index-page">
          
          FOPP-PIE
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="ad_warning"></span></li>
        <li><span id="browsing_warning"></span></li>

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://guide.runestone.academy' target="_blank" rel="noopener noreferrer" >Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=FOPP-PIE&page=ChapterAssessment'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">6.13. </span>Exercises" data-toggle="tooltip">
    <a href="Exercises.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">7. </span>Integration Pause' data-toggle="tooltip" >              
    <a href="../IntegrationPause/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <section id="chapter-assessment">
<h1><span class="section-number">6.14. </span>Chapter Assessment<a class="headerlink" href="#chapter-assessment" title="Permalink to this heading">¶</a></h1>
<p><strong>Check your understanding</strong></p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch6_1 data-question_label="6.14.1">
<div id=assess_ch6_1_question class="ac_question">
<p><code class="docutils literal notranslate"><span class="pre">rainfall_mi</span></code> is a string that contains the average number of inches of rainfall in Michigan for every month (in inches) with every month separated by a comma.
Write code to compute the number of months that have more than 3 inches of rainfall. Store the result in the variable <code class="docutils literal notranslate"><span class="pre">num_rainy_months</span></code>.
In other words, count the number of items with values <code class="docutils literal notranslate"><span class="pre">&gt;</span> <span class="pre">3.0</span></code>.</p>
<p>Hard-coded answers will receive no credit.</p>

</div>
<textarea data-lang="python" id="assess_ch6_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
rainfall_mi = "1.65, 1.46, 2.05, 3.03, 3.35, 3.46, 2.83, 3.23, 3.5, 2.52, 2.8, 1.85"
=====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertIn('for', self.getEditorText(), "Testing that your code has a for loop (Don't worry about actual and expected values).")
        self.assertEqual(num_rainy_months, 5, "Testing that num_rainy_months has the right value")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch6_2 data-question_label="6.14.2">
<div id=assess_ch6_2_question class="ac_question">
<p>The variable <code class="docutils literal notranslate"><span class="pre">sentence</span></code> stores a string. Write code to determine how many words in <code class="docutils literal notranslate"><span class="pre">sentence</span></code> start and end with the same letter, including one-letter words.
Store the result in the variable <code class="docutils literal notranslate"><span class="pre">same_letter_count</span></code>.</p>
<p>Hard-coded answers will receive no credit.</p>

</div>
<textarea data-lang="python" id="assess_ch6_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
sentence = "students flock to the arb for a variety of outdoor activities such as jogging and picnicking"

# Write your code here.


====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(same_letter_count, 2, "Checking that same_letter_count has the correct value")
        self.assertIn('for ', self.getEditorText(), "Testing that your code has a for loop")
myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch6_3 data-question_label="6.14.3">
<div id=assess_ch6_3_question class="ac_question">
<p>Write code to count the number of strings in list <code class="docutils literal notranslate"><span class="pre">items</span></code> that have the character <code class="docutils literal notranslate"><span class="pre">w</span></code> in it. Assign that number to the variable <code class="docutils literal notranslate"><span class="pre">acc_num</span></code>.</p>
<p>HINT 1: Use the accumulation pattern!</p>
<p>HINT 2: the <code class="docutils literal notranslate"><span class="pre">in</span></code> operator checks whether a substring is present in a string.</p>
<p>Hard-coded answers will receive no credit.</p>

</div>
<textarea data-lang="python" id="assess_ch6_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
items = ["whirring", "wow!", "calendar", "wry", "glass", "", "llama","tumultuous","owing"]


=====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertIn(' in ', self.getEditorText(), "Testing that you are using the in operator.")
        self.assertEqual(acc_num, 4, "Testing that acc_num has been set to the number of strings that have 'w' in them.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch6_4 data-question_label="6.14.4">
<div id=assess_ch6_4_question class="ac_question">
<p>Write code that counts the number of words in <code class="docutils literal notranslate"><span class="pre">sentence</span></code> that contain <em>either</em> an “a” or an “e”. Store the result in the variable <code class="docutils literal notranslate"><span class="pre">num_a_or_e</span></code>.</p>
<p>Note 1: be sure to not double-count words that contain both an a and an e.</p>
<p>HINT 1: Use the <code class="docutils literal notranslate"><span class="pre">in</span></code> operator.</p>
<p>HINT 2: You can either use <code class="docutils literal notranslate"><span class="pre">or</span></code> or <code class="docutils literal notranslate"><span class="pre">elif</span></code>.</p>
<p>Hard-coded answers will receive no credit.</p>

</div>
<textarea data-lang="python" id="assess_ch6_4_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
sentence = "python is a high level general purpose programming language that can be applied to many different classes of problems."


=====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertIn(' in ', self.getEditorText(), "Testing that you are using the in operator.")
        self.assertEqual(num_a_or_e, 14, "Testing that num_a_or_e has been set to the correct number.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch6_5 data-question_label="6.14.5">
<div id=assess_ch6_5_question class="ac_question">
<p>Write code that will count the number of vowels in the sentence <code class="docutils literal notranslate"><span class="pre">s</span></code> and assign the result to the variable <code class="docutils literal notranslate"><span class="pre">num_vowels</span></code>. For this problem, vowels are only a, e, i, o, and u. Hint: use the <code class="docutils literal notranslate"><span class="pre">in</span></code> operator with <code class="docutils literal notranslate"><span class="pre">vowels</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch6_5_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
s = "singing in the rain and playing in the rain are two entirely different situations but both can be fun"
vowels = ['a','e','i','o','u']

# Write your code here.


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertEqual(num_vowels, 32, "testing whether num_vowels is set correctly")

    def testOneA(self):
       self.assertIn('for', self.getEditorText(), "Testing that you are using a for loop.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch6_6 data-question_label="6.14.6">
<div id=assess_ch6_6_question class="ac_question">
<p>Create one conditional so that if “Friendly” is in <code class="docutils literal notranslate"><span class="pre">w</span></code>, then “Friendly is here!” should be assigned to the variable <code class="docutils literal notranslate"><span class="pre">wrd</span></code>. If it’s not, check if “Friend” is in <code class="docutils literal notranslate"><span class="pre">w</span></code>. If so, the string “Friend is here!” should be assigned to the variable <code class="docutils literal notranslate"><span class="pre">wrd</span></code>, otherwise “No variation of friend is in here.” should be assigned to the variable <code class="docutils literal notranslate"><span class="pre">wrd</span></code>. (Also consider: does the order of your conditional statements matter for this problem? Why?)</p>

</div>
<textarea data-lang="python" id="assess_ch6_6_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
w = "Friendship is a wonderful human experience!"

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(wrd, "Friend is here!", "Testing the value of wrd")
      self.assertIn("else", self.getEditorText(), "Checking that you used an else clause.")
      self.assertIn("elif", self.getEditorText(), "Checking that you used an elif clause.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch6_7 data-question_label="6.14.7">
<div id=assess_ch6_7_question class="ac_question">
<p>We have written conditionals for you to use. Create the variable x and assign it some integer so that at the end of the code, <code class="docutils literal notranslate"><span class="pre">output</span></code> will be assigned the string <code class="docutils literal notranslate"><span class="pre">&quot;Consistently</span> <span class="pre">working&quot;</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch6_7_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
if x >= 10:
    output = "working"
else:
    output = "Still working"
if x > 12:
    output = "Always working"
elif x < 7:
    output = "Forever working"
else:
    output = "Consistently working"

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(output, "Consistently working", "Testing the value of output")
   def testTwo(self):
      self.assertIn(x, [7,8,9,10,11,12], "Testing that x was assigned a correct number" )

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch6_8 data-question_label="6.14.8">
<div id=assess_ch6_8_question class="ac_question">
<p>Write code so that if <code class="docutils literal notranslate"><span class="pre">&quot;STATS</span> <span class="pre">250&quot;</span></code> is in the list <code class="docutils literal notranslate"><span class="pre">schedule</span></code>, then the string <code class="docutils literal notranslate"><span class="pre">&quot;You</span> <span class="pre">could</span> <span class="pre">be</span> <span class="pre">in</span> <span class="pre">Information</span> <span class="pre">Science!&quot;</span></code> is assigned to the variable <code class="docutils literal notranslate"><span class="pre">resp</span></code>. Otherwise, the string <code class="docutils literal notranslate"><span class="pre">&quot;That's</span> <span class="pre">too</span> <span class="pre">bad.&quot;</span></code> should be assigned to the variable <code class="docutils literal notranslate"><span class="pre">resp</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch6_8_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
schedule = ["SI 106", "STATS 250", "SI 110", "ENGLISH 124/125"]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(resp, "You could be in Information Science!", "Testing the value of resp given the schedule list provided.")
      self.assertIn("if", self.getEditorText(), "Testing that you used an if clause.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch6_9 data-question_label="6.14.9">
<div id=assess_ch6_9_question class="ac_question">
<p>Create the variable <code class="docutils literal notranslate"><span class="pre">z</span></code> whose value is <code class="docutils literal notranslate"><span class="pre">30</span></code>. Write code to see if <code class="docutils literal notranslate"><span class="pre">z</span></code> is greater than <code class="docutils literal notranslate"><span class="pre">y</span></code>. If so, add 5 to <code class="docutils literal notranslate"><span class="pre">y</span></code>’s value, otherwise do nothing. Then, multiply <code class="docutils literal notranslate"><span class="pre">z</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code>, and assign the resulting value to the variable <code class="docutils literal notranslate"><span class="pre">x</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch6_9_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
y = 22

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(x, 810, "Testing the value of x")
   def testTwo(self):
      self.assertEqual(z, 30, "Testing the value of z.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch6_10 data-question_label="6.14.10">
<div id=assess_ch6_10_question class="ac_question">
<p>For each string in <code class="docutils literal notranslate"><span class="pre">wrd_lst</span></code>, find the number of characters in the string. If the number of characters is less than 6, add 1 to <code class="docutils literal notranslate"><span class="pre">accum</span></code> so that in the end, <code class="docutils literal notranslate"><span class="pre">accum</span></code> will contain an integer representing the total number of words in the list that have fewer than 6 characters.</p>

</div>
<textarea data-lang="python" id="assess_ch6_10_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
wrd_lst = ["Hello", "activecode", "Java", "C#", "Python", "HTML and CSS", "Javascript", "Swift", "php"]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(accum, 5, "Testing the value of accum")
      self.assertIn('for ', self.getEditorText(), "Testing that your code has a for loop")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone ">
<div data-component="clickablearea"  id="assess_ch6_11_lh" data-question_label="6.14.11"  data-table data-cc="2,3;3,3;" data-ci="4,3;5,3" style="visibility: hidden;">
<span data-question>: The truth table below is supposed to represent “A and B”, but it has some incorrect values in the third column. Look at the code in CodeLens above to see how Python evaluates “and” operators. Click on the INCORRECT "A and B" values.</span>
<img alt="../_images/cdq6-1.png" src="../_images/cdq6-1.png" />
<table class="docutils align-default">
<tbody>
<tr class="row-odd"><td><p>A</p></td>
<td><p>B</p></td>
<td><p>A and B</p></td>
</tr>
<tr class="row-even"><td><p>True</p></td>
<td><p>False</p></td>
<td><p>True</p></td>
</tr>
<tr class="row-odd"><td><p>False</p></td>
<td><p>False</p></td>
<td><p>True</p></td>
</tr>
<tr class="row-even"><td><p>True</p></td>
<td><p>True</p></td>
<td><p>True</p></td>
</tr>
<tr class="row-odd"><td><p>False</p></td>
<td><p>True</p></td>
<td><p>False</p></td>
</tr>
</tbody>
</table>

</div>
</div>
<img alt="../_images/cdq6-2.png" src="../_images/cdq6-2.png" />

<div class="runestone ">
<div data-component="clickablearea"  id="assess_ch6_12_lh" data-question_label="6.14.12"     style="visibility: hidden;">
<span data-question>: This code has a semantic error. Notice in the CodeLens image above that the output states the user is underage. Without changing the values of the variables, select the portion of the code that needs to change.</span><span data-feedback>Consider the different boolean logic operators and what each means</span><pre>MANITOBA_LEGAL_AGE = <span data-incorrect>18</span>

my_age = <span data-incorrect>18</span>

<span data-incorrect>if my_age</span> <span data-correct>></span> <span data-incorrect>MANITOBA_LEGAL_AGE:</span>
  <span data-incorrect>print("You are of legal age in Manitoba")</span>
<span data-incorrect>else:</span>
    <span data-incorrect>print("You are underage in Manitoba")</span></pre>

</div>
</div>

        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="6.14.13" id="assess_ch6_13"  style="visibility: hidden;">
            <p>Given that same code/image from the previous question, what operator should be used so the code works as intended?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*>=\\s*$", "regexFlags": "i", "feedback": "<p>Correct! You are legal in Manitoba if your age is greater than OR equal to 18</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "i", "feedback": "<p>Hmmm try again, you may want to review the different types of boolean operators</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="6.14.14" id="assess_ch6_14_lh"  style="visibility: hidden;">
            <p>:</p>
<img alt="../_images/cdq6-3.png" src="../_images/cdq6-3.png" />
<p>Given num = 49 on line 3, what value of i will cause line 14 to run? Hint: Prime numbers are numbers greater than 1 that only have two factors, 1 and the number itself. Notice that in order to reach line 14, line 12 must evaluate to true. Walk through the code to determine what is happening here and what the value will be when line 12 becomes true.</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*seven|7\\s*$", "regexFlags": "i", "feedback": "<p>Correct! When i = 7, (num % i) will equal 0</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "i", "feedback": "<p>Hmmm try again, what causes the conditional on line 12 to be true?</p>\n"}]]
        </script>

        </div>
        </div>
            </section>


  

  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 6.13. Exercises' data-toggle="tooltip">
  <a href="Exercises.html" aria-label="Previous - 6.13. Exercises">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 7. Integration Pause' data-toggle="tooltip" >
    <a href="../IntegrationPause/toctree.html" aria-label="Next - 7. Integration Pause">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 7.2.11.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>