<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html  lang="en" class="runestone-sphinx">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>4.5. Index Operator: Working with the Characters of a String &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=E2DC0A34" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=4182B24D" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-347.f9add1ca35d5ad93.css?v=306CBB19" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-runestone.f64dcc3632d7a0c3.css?v=F4F91F21" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script src="../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script src="../_static/sphinx_highlight.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/prefix-runtime.2438a8bdd6d43b5c.bundle.js?v=7D29D67B"></script>
    <script defer="" src="../_static/prefix-347.8328b60515045466.bundle.js?v=2CF09D3E"></script>
    <script defer="" src="../_static/prefix-runestone.aeac7cf9ad5c07ef.bundle.js?v=4FF65C30"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
<script src="https://media.ethicalads.io/media/client/ethicalads.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="4.6. Strings and for loops" href="Stringsandforloops.html" />
    <link rel="prev" title="4.4. Flow of Execution of the for Loop" href="FlowofExecutionoftheforLoop.html" />

  
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'FOPP-PIE';
    eBookConfig.basecourse = 'FOPP-PIE';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '7.2.11';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" id="rs-book-toc-link" href="../index.html" aria-label="index-page">
          
          FOPP-PIE
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="ad_warning"></span></li>
        <li><span id="browsing_warning"></span></li>

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://guide.runestone.academy' target="_blank" rel="noopener noreferrer" >Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=FOPP-PIE&page=IndexOperatorForString'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">4.4. </span>Flow of Execution of the for Loop" data-toggle="tooltip">
    <a href="FlowofExecutionoftheforLoop.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">4.6. </span>Strings and <code class="docutils literal notranslate"><span class="pre">for</span></code> loops' data-toggle="tooltip" >              
    <a href="Stringsandforloops.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <section id="index-operator-working-with-the-characters-of-a-string">
<span id="index-0"></span><h1><span class="section-number">4.5. </span>Index Operator: Working with the Characters of a String<a class="headerlink" href="#index-operator-working-with-the-characters-of-a-string" title="Permalink to this heading">¶</a></h1>
<p>The <strong>indexing operator</strong> (Python uses square brackets to enclose the index)
selects a single character from a string.  The characters are accessed by their position or
index value.  For example, in the string shown below, the 14 characters are indexed left to right
from postion 0 to position 13.</p>
<img alt="index values" src="../_images/indexvalues.png" />
<p>It is also the case that the positions are named from right to left using negative numbers where -1 is
the rightmost index and so on. Note that the character at index 6 (or -8) is the blank character.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_5_1 data-question_label="4.5.1">
<div id=ac4_5_1_question class="ac_question">

</div>
<textarea data-lang="python" id="ac4_5_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
school = "Luther College"
m = school[2]
print(m)

lastchar = school[-1]
print(lastchar)
</textarea>
</div>
</div>
<p>The expression <code class="docutils literal notranslate"><span class="pre">school[2]</span></code> selects the character at index 2 from <code class="docutils literal notranslate"><span class="pre">school</span></code>, and creates a new
string containing just this one character. The variable <code class="docutils literal notranslate"><span class="pre">m</span></code> refers to the result.</p>
<p>The letter at index zero of <code class="docutils literal notranslate"><span class="pre">&quot;Luther</span> <span class="pre">College&quot;</span></code> is <code class="docutils literal notranslate"><span class="pre">L</span></code>.  So at
position <code class="docutils literal notranslate"><span class="pre">[2]</span></code> we have the letter <code class="docutils literal notranslate"><span class="pre">t</span></code>.</p>
<p>If you want the zero-eth letter of a string, you just put 0, or any expression
with the value 0, in the brackets.  Give it a try.</p>
<p>The expression in brackets is called an <strong>index</strong>. An index specifies a member
of an ordered collection.  In this case the collection of characters in the string. The index
<em>indicates</em> which character you want. It can be any integer
expression so long as it evaluates to a valid index value.</p>
<p>Note that indexing returns a <em>string</em> — Python has no special type for a single character.
It is just a string of length 1.</p>
<section id="index-operator-accessing-elements-of-a-list">
<h2><span class="section-number">4.5.1. </span>Index Operator: Accessing Elements of a List<a class="headerlink" href="#index-operator-accessing-elements-of-a-list" title="Permalink to this heading">¶</a></h2>
<p>The syntax for accessing the elements of a list is the same as the syntax for
accessing the characters of a string.  We use the index operator  <code class="docutils literal notranslate"><span class="pre">[]</span></code>
. The expression inside the brackets specifies
the index. Remember that the indices start at 0. Any integer expression can be used
as an index and as with strings, negative index values will locate items from the right instead
of from the left.</p>
<p>When we say the first, third or nth character of a sequence, we generally mean counting the usual way, starting with 1.  The nth character and the character AT INDEX n are different then:  The nth character is at index n-1.  The fourth character in a string is at index 3. The seventh element in a list is at index 6. Make sure you are clear on what you mean!</p>
<p>Try to predict what will be printed out by the following code, and then run it to check your
prediction. (Actually, it’s a good idea to always do that with the code examples. You
will learn much more if you force yourself to make a prediction before you see the output.)</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_5_2 data-question_label="4.5.1.1">
<div id=ac4_5_2_question class="ac_question">

</div>
<textarea data-lang="python" id="ac4_5_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
numbers = [17, 123, 87, 34, 66, 8398, 44]
print(numbers[2])
print(numbers[9-8])
print(numbers[-2])
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_5_3 data-question_label="4.5.1.2">
<div id=ac4_5_3_question class="ac_question">

</div>
<textarea data-lang="python" id="ac4_5_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
prices = [1.99, 2.00, 5.50, 20.95, 100.98]
print(prices[0])
print(prices[-1])
print(prices[3-5])
</textarea>
</div>
</div>
<p><strong>Check your understanding</strong></p>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="4.5.1.3" data-multipleanswers="false"  id="question4_5_1"  style="visibility: hidden;">
    <p>What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">s</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">s</span><span class="p">[</span><span class="mi">3</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="question4_5_1_opt_a">t</li><li data-component="feedback">Index locations do not start with 1, they start with 0.</li>
    
    <li data-component="answer" data-correct='yes' id="question4_5_1_opt_b">h</li><li data-component="feedback">Yes, index locations start with 0.</li>
    
    <li data-component="answer"  id="question4_5_1_opt_c">c</li><li data-component="feedback">s[-3] would return c, counting from right to left.</li>
    
    <li data-component="answer"  id="question4_5_1_opt_d">Error, you cannot use the [ ] operator with a string.</li><li data-component="feedback">[ ] is the index operator.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="4.5.1.4" data-multipleanswers="false"  id="question4_5_2"  style="visibility: hidden;">
    <p>What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">s</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">s</span><span class="p">[</span><span class="mi">2</span><span class="p">]</span> <span class="o">+</span> <span class="n">s</span><span class="p">[</span><span class="o">-</span><span class="mi">4</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="question4_5_2_opt_a">tr</li><li data-component="feedback">Almost, t is at postion 2, counting left to right starting from 0; but r is at -5, counting right to left starting from -1.</li>
    
    <li data-component="answer" data-correct='yes' id="question4_5_2_opt_b">to</li><li data-component="feedback">For -4 you count from right to left, starting with -1.</li>
    
    <li data-component="answer"  id="question4_5_2_opt_c">ps</li><li data-component="feedback">p is at location 0, not 2.</li>
    
    <li data-component="answer"  id="question4_5_2_opt_d">nn</li><li data-component="feedback">n is at location 5, not 2.</li>
    
    <li data-component="answer"  id="question4_5_2_opt_e">Error, you cannot use the [ ] operator with the + operator.</li><li data-component="feedback">[ ] operator returns a string that can be concatenated with another string.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="4.5.1.5" data-multipleanswers="false"  id="question4_5_3"  style="visibility: hidden;">
    <p>What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">alist</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="n">alist</span><span class="p">[</span><span class="mi">5</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="question4_5_3_opt_a">[ ]</li><li data-component="feedback">The empty list is at index 4.</li>
    
    <li data-component="answer" data-correct='yes' id="question4_5_3_opt_b">3.14</li><li data-component="feedback">Yes, 3.14 is at index 5 since we start counting at 0 and sublists count as one item.</li>
    
    <li data-component="answer"  id="question4_5_3_opt_c">False</li><li data-component="feedback">False is at index 6.</li>
    
    <li data-component="answer"  id="question4_5_3_opt_d">"dog"</li><li data-component="feedback">Look again, the element at index 3 is a list. This list only counts as one element.</li>
    

    </ul>
    </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_5_4 data-question_label="4.5.1.6">
<div id=ac4_5_4_question class="ac_question">
<p>Assign the value of the 34th element of <code class="docutils literal notranslate"><span class="pre">lst</span></code> to the variable <code class="docutils literal notranslate"><span class="pre">output</span></code>.</p>

</div>
<textarea data-lang="python" id="ac4_5_4_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
lst = ["hi", "morning", "dog", "506", "caterpillar", "balloons", 106, "yo-yo", "python", "moon", "water", "sleepy", "daffy", 45, "donald", "whiteboard", "glasses", "markers", "couches", "butterfly", "100", "magazine", "door", "picture", "window", ["Olympics", "handle"], "chair", "pages", "readings", "burger", "juggle", "craft", ["store", "poster", "board"], "laptop", "computer", "plates", "hotdog", "salad", "backpack", "zipper", "ring", "watch", "finger", "bags", "boxes", "pods", "peas", "apples", "horse", "guinea pig", "bowl", "EECS"]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(output, "laptop", "Testing that output value is assigned to correct value.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_5_5 data-question_label="4.5.1.7">
<div id=ac4_5_5_question class="ac_question">
<p>Assign the value of the 23rd element of <code class="docutils literal notranslate"><span class="pre">l</span></code> to the variable <code class="docutils literal notranslate"><span class="pre">checking</span></code>.</p>

</div>
<textarea data-lang="python" id="ac4_5_5_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
l = ["hi", "goodbye", "python", "106", "506", 91, ['all', 'Paul', 'Jackie', "UMSI", 1, "Stephen", 4.5], 109, "chair", "pizza", "wolverine", 2017, 3.92, 1817, "account", "readings", "papers", 12, "facebook", "twitter", 193.2, "snapchat", "leaders and the best", "social", "1986", 9, 29, "holiday", ["women", "olympics", "gold", "rio", 21, "2016", "men"], "26trombones"]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(checking, "leaders and the best", "Testing that checking has the correct element assigned.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_5_6 data-question_label="4.5.1.8">
<div id=ac4_5_6_question class="ac_question">
<p>Assign the value of the last chacter of <code class="docutils literal notranslate"><span class="pre">lst</span></code> to the variable <code class="docutils literal notranslate"><span class="pre">output</span></code>. Do this so that the length of lst doesn’t matter.</p>

</div>
<textarea data-lang="python" id="ac4_5_6_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
lst = "Every chess or checkers game begins from the same position and has a finite number of moves that can be played. While the number of possible scenarios and moves is quite large, it is still possible for computers to calculate that number and even be programmed to respond well against a human player..."

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testThree(self):
      self.assertEqual(output, ".", "Testing that output value is assigned to correct value.")

myTests().main()
</textarea>
</div>
</div>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>Why does counting start at 0 going from left to right, but at -1 going from right to left? Well, indexing starting at 0
has a long history in computer science having to do with some low-level implementation details that we won’t
go into. For indexing from right to left, it might seem natural to do the analgous thing
and start at -0. Unfortunately, -0 is the same as 0, so s[-0] can’t be the last item. Remember we
said that programming languages are formal languages where details matter and
everything is taken literally?</p>
</div>
</section>
</section>


  

  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 4.4. Flow of Execution of the for Loop' data-toggle="tooltip">
  <a href="FlowofExecutionoftheforLoop.html" aria-label="Previous - 4.4. Flow of Execution of the for Loop">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 4.6. Strings and  for  loops' data-toggle="tooltip" >
    <a href="Stringsandforloops.html" aria-label="Next - 4.6. Strings and  for  loops">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 7.2.11.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>