<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html  lang="en" class="runestone-sphinx">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>4.9. The Accumulator Pattern &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=E2DC0A34" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=4182B24D" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-347.f9add1ca35d5ad93.css?v=306CBB19" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-runestone.f64dcc3632d7a0c3.css?v=F4F91F21" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script src="../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script src="../_static/sphinx_highlight.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/prefix-runtime.2438a8bdd6d43b5c.bundle.js?v=7D29D67B"></script>
    <script defer="" src="../_static/prefix-347.8328b60515045466.bundle.js?v=2CF09D3E"></script>
    <script defer="" src="../_static/prefix-runestone.aeac7cf9ad5c07ef.bundle.js?v=4FF65C30"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
<script src="https://media.ethicalads.io/media/client/ethicalads.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="4.10. 👩‍💻 Printing Intermediate Results" href="WPPrintingIntermediateResults.html" />
    <link rel="prev" title="4.8. 👩‍💻 Naming Variables in For Loops" href="WPNamingVariablesinForLoops.html" />

  
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'FOPP-PIE';
    eBookConfig.basecourse = 'FOPP-PIE';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '7.2.11';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" id="rs-book-toc-link" href="../index.html" aria-label="index-page">
          
          FOPP-PIE
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="ad_warning"></span></li>
        <li><span id="browsing_warning"></span></li>

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://guide.runestone.academy' target="_blank" rel="noopener noreferrer" >Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=FOPP-PIE&page=TheAccumulatorPattern'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">4.8. </span>👩‍💻 Naming Variables in For Loops" data-toggle="tooltip">
    <a href="WPNamingVariablesinForLoops.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">4.10. </span>👩‍💻 Printing Intermediate Results' data-toggle="tooltip" >              
    <a href="WPPrintingIntermediateResults.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <section id="the-accumulator-pattern">
<h1><span class="section-number">4.9. </span>The Accumulator Pattern<a class="headerlink" href="#the-accumulator-pattern" title="Permalink to this heading">¶</a></h1>
<p>One common programming “pattern” is to traverse a sequence, <strong>accumulating</strong> a value as we go,
such as the sum-so-far or the maximum-so-far. That way, at the end of the traversal we have
accumulated a single value, such as the sum total of all the items or the largest item.</p>
<dl class="simple">
<dt>The anatomy of the accumulation pattern includes:</dt><dd><ul class="simple">
<li><p><strong>initializing</strong> an “accumulator” variable to an initial value (such as 0 if accumulating a sum)</p></li>
<li><p><strong>iterating</strong> (e.g., traversing the items in a sequence)</p></li>
<li><p><strong>updating</strong> the accumulator variable on each iteration (i.e., when processing each item in the sequence)</p></li>
</ul>
</dd>
</dl>
<p>For example, consider the following code, which computes the sum of the numbers in a list.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_9_1 data-question_label="4.9.1">
<div id=ac4_9_1_question class="ac_question">

</div>
<textarea data-lang="python" id="ac4_9_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
nums = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
accum = 0
for w in nums:
    accum = accum + w
print(accum)
</textarea>
</div>
</div>
<p>In the program above, notice that the variable <code class="docutils literal notranslate"><span class="pre">accum</span></code> starts out with a value of 0.
Next, the iteration is performed 10 times.  Inside the for loop, the update occurs.
<code class="docutils literal notranslate"><span class="pre">w</span></code> has the value of the current item (1 the first time, then 2, then 3, etc.).
<code class="docutils literal notranslate"><span class="pre">accum</span></code> is reassigned a new value: the old value plus the current value of <code class="docutils literal notranslate"><span class="pre">w</span></code>.</p>
<p>This pattern of iterating the updating of a variable is commonly referred to as the
<strong>accumulator pattern</strong>. We refer to the variable as the <strong>accumulator</strong>. This pattern will come up
over and over again. Remember that the key to making it work successfully is to be sure to
initialize the variable <strong>before</strong> you start the iteration. Once inside the iteration, it is required
that you update the accumulator.</p>
<p>Here is the same program in CodeLens.  Step through the function and watch the “running total”
accumulate the result.</p>

<div class="runestone codelens ">
    <div class="cd_section" data-component="codelens" data-question_label="4.9.2">
        <div id=clens4_9_1_question class="ac_question">
        </div>
        <div class="pytutorVisualizer" id="clens4_9_1"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 4.9.2 (clens4_9_1)</span> </p>
    </div>

<script>
var clens4_9_1_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["clens4_9_1"] = {"code": "nums = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]\naccum = 0\nfor w in nums:\n   accum = accum + w\nprint(accum)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1]}, "ordered_globals": ["nums"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 0}, "ordered_globals": ["nums", "accum"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 0, "w": 1}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 1, "w": 1}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 1, "w": 2}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 3, "w": 2}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 3, "w": 3}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 6, "w": 3}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 6, "w": 4}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 10, "w": 4}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 10, "w": 5}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 15, "w": 5}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 15, "w": 6}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 21, "w": 6}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 21, "w": 7}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 28, "w": 7}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 28, "w": 8}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 36, "w": 8}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 36, "w": 9}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 45, "w": 9}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 45, "w": 10}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 55, "w": 10}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 55, "w": 10}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 5, "event": "return", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 55, "w": 10}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": "55\n"}]};
</script>
</div>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>What would happen if we indented the print accum statement? Not sure? Make a prediction, then try it and find out.</p>
</div>
<p>We can utilize the range function in this situation as well. Previously, you’ve seen it used when we wanted to do repeated drawing instructions with a turtle. There we used it to iterate a certain number of times. We can do more than that though. The <code class="docutils literal notranslate"><span class="pre">range</span></code> function
takes at least one input - which should be an integer - and returns a list as long as your input. While you can provide
two inputs, we will focus on using range with just one input. With one input, range will start at zero and go up to - but
not include - the input. Here are some examples:</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_9_2 data-question_label="4.9.3">
<div id=ac4_9_2_question class="ac_question">

</div>
<textarea data-lang="python" id="ac4_9_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
print("range(5): ")
for i in range(5):
  print(i)

print("range(0,5): ")
for i in range(0, 5):
  print(i)

# We can also print this out by casting `range` to a `list`:
print(list(range(5)))
print(list(range(0,5)))
</textarea>
</div>
</div>
<p>One important thing to know about the range function in Python 3 is that if we want to use it outside of iteration, we
have to cast it to a list using <code class="docutils literal notranslate"><span class="pre">list()</span></code>.</p>
<p>Here’s how you could use the range function in the previous problem.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_9_3 data-question_label="4.9.4">
<div id=ac4_9_3_question class="ac_question">

</div>
<textarea data-lang="python" id="ac4_9_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
accum = 0
for w in range(11):
    accum = accum + w
print(accum)

# or, if you use two inputs for the range function

sec_accum = 0
for w in range(1,11):
    sec_accum = sec_accum + w
print(sec_accum)
</textarea>
</div>
</div>
<p>Because the range function is exclusive of the ending number, we have to use 11 as the function input.</p>
<p>We can use the accumulation pattern to count the number of something or to sum up a total. The
above examples only covered how to get the sum for a list, but we can also count how many items are
in the list if we wanted to.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_9_4 data-question_label="4.9.5">
<div id=ac4_9_4_question class="ac_question">

</div>
<textarea data-lang="python" id="ac4_9_4_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
nums = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
count = 0
for w in nums:
    count = count + 1
print(count)
</textarea>
</div>
</div>
<p>In this example we don’t make use of <code class="docutils literal notranslate"><span class="pre">w</span></code> even though the iterator variable (loop variable) is a necessary part of
constructing a for loop. Instead of adding the value of <code class="docutils literal notranslate"><span class="pre">w</span></code> to <code class="docutils literal notranslate"><span class="pre">count</span></code> we add a 1 to it,
because we’re incrementing the value of count when we iterate each time through the loop. Though in
this scenario we could have used the <code class="docutils literal notranslate"><span class="pre">len</span></code> function, there are other cases later on where len
won’t be useful but we will still need to count.</p>
<p><strong>Check your understanding</strong></p>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="4.9.6" data-multipleanswers="false"  id="question4_9_1"  style="visibility: hidden;">
    <p>Consider the following code:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">nums</span> <span class="o">=</span> <span class="p">[</span><span class="mi">1</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">3</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">6</span><span class="p">,</span> <span class="mi">7</span><span class="p">,</span> <span class="mi">8</span><span class="p">,</span> <span class="mi">9</span><span class="p">,</span> <span class="mi">10</span><span class="p">]</span>
<span class="k">for</span> <span class="n">w</span> <span class="ow">in</span> <span class="n">nums</span><span class="p">:</span>
   <span class="n">accum</span> <span class="o">=</span> <span class="mi">0</span>
   <span class="n">accum</span> <span class="o">=</span> <span class="n">accum</span> <span class="o">+</span> <span class="n">w</span>
<span class="nb">print</span><span class="p">(</span><span class="n">accum</span><span class="p">)</span>
</pre></div>
</div>
<p>What happens if you put the initialization of accum inside the for loop as the first
instruction in the loop?</p>

    <li data-component="answer" data-correct='yes' id="question4_9_1_opt_a">It will print out 10 instead of 55</li><li data-component="feedback">The variable accum will be reset to 0 each time through the loop. Then it will add the current item. Only the last item will count.</li>
    
    <li data-component="answer"  id="question4_9_1_opt_b">It will cause a run-time error</li><li data-component="feedback">Assignment statements are perfectly legal inside loops and will not cause an error.</li>
    
    <li data-component="answer"  id="question4_9_1_opt_c">It will print out 0 instead of 55</li><li data-component="feedback">Good thought: the variable accum will be reset to 0 each time through the loop. But then it adds the current item.</li>
    

    </ul>
    </div>
    
        <div class="runestone parsons-container ">
        <div data-component="parsons" id="pp4_9_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Rearrange the code statements so that the program will add up the first n odd numbers where n is provided by the user. Note that indentation matters, so drag code blocks to the right if they should be indented.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="4.9.7"          style="visibility: hidden;">
        n_str = input('How many odd numbers would you like to add together?')
n = int(n_str)
thesum = 0
oddnumber = 1
---
for counter in range(n):
---
   thesum = thesum + oddnumber
   oddnumber = oddnumber + 2
---
print(thesum)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_9_5 data-question_label="4.9.8">
<div id=ac4_9_5_question class="ac_question">
<p>Write code to create a list of integers from 0 through 52 and assign that list to the variable <code class="docutils literal notranslate"><span class="pre">numbers</span></code>. You should use a special Python function – do not type out the whole list yourself. HINT: You can do this in one line of code!</p>

</div>
<textarea data-lang="python" id="ac4_9_5_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(numbers, range(53), "Testing that numbers is a list that contains the correct elements.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_9_6 data-question_label="4.9.9">
<div id=ac4_9_6_question class="ac_question">
<p>Count the number of characters in string <code class="docutils literal notranslate"><span class="pre">str1</span></code>. Do not use <code class="docutils literal notranslate"><span class="pre">len()</span></code>. Save the number in variable <code class="docutils literal notranslate"><span class="pre">numbs</span></code>.</p>

</div>
<textarea data-lang="python" id="ac4_9_6_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
str1 = "I like nonsense, it wakes up the brain cells. Fantasy is a necessary ingredient in living."

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testEight(self):
      self.assertEqual(numbs, 90, "Testing that numbs is assigned to correct values.")
      self.assertNotIn("len(", self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac4_9_7 data-question_label="4.9.10">
<div id=ac4_9_7_question class="ac_question">
<p>Create a list of numbers 0 through 40 and assign this list to the variable <code class="docutils literal notranslate"><span class="pre">numbers</span></code>. Then, accumulate the total of the list’s values and assign that sum to the variable <code class="docutils literal notranslate"><span class="pre">sum1</span></code>.</p>

</div>
<textarea data-lang="python" id="ac4_9_7_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testNineA(self):
      self.assertEqual(numbers, [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40], "Testing that numbers is assigned to correct values.")

   def testNineB(self):
      self.assertEqual(sum1, 820, "Testing that sum1 has the correct value.")

myTests().main()
</textarea>
</div>
</div>
</section>


  

  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 4.8. 👩‍💻 Naming Variables in For Loops' data-toggle="tooltip">
  <a href="WPNamingVariablesinForLoops.html" aria-label="Previous - 4.8. 👩‍💻 Naming Variables in For Loops">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 4.10. 👩‍💻 Printing Intermediate Results' data-toggle="tooltip" >
    <a href="WPPrintingIntermediateResults.html" aria-label="Next - 4.10. 👩‍💻 Printing Intermediate Results">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 7.2.11.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>