<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html  lang="en" class="runestone-sphinx">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>4.17. Chapter Assessment &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=E2DC0A34" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=4182B24D" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-347.f9add1ca35d5ad93.css?v=306CBB19" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-runestone.f64dcc3632d7a0c3.css?v=F4F91F21" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script src="../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script src="../_static/sphinx_highlight.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/prefix-runtime.2438a8bdd6d43b5c.bundle.js?v=7D29D67B"></script>
    <script defer="" src="../_static/prefix-347.8328b60515045466.bundle.js?v=2CF09D3E"></script>
    <script defer="" src="../_static/prefix-runestone.aeac7cf9ad5c07ef.bundle.js?v=4FF65C30"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
<script src="https://media.ethicalads.io/media/client/ethicalads.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="5. Functions and Files" href="../FunctionsAndFiles/toctree.html" />
    <link rel="prev" title="4.16. Exercises" href="Exercises.html" />

  
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'FOPP-PIE';
    eBookConfig.basecourse = 'FOPP-PIE';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '7.2.11';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" id="rs-book-toc-link" href="../index.html" aria-label="index-page">
          
          FOPP-PIE
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="ad_warning"></span></li>
        <li><span id="browsing_warning"></span></li>

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://guide.runestone.academy' target="_blank" rel="noopener noreferrer" >Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=FOPP-PIE&page=week2a2'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">4.16. </span>Exercises" data-toggle="tooltip">
    <a href="Exercises.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">5. </span>Functions and Files' data-toggle="tooltip" >              
    <a href="../FunctionsAndFiles/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <section id="chapter-assessment">
<h1><span class="section-number">4.17. </span>Chapter Assessment<a class="headerlink" href="#chapter-assessment" title="Permalink to this heading">¶</a></h1>
<p><strong>Check your understanding</strong></p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch4_1 data-question_label="4.17.1">
<div id=assess_ch4_1_question class="ac_question">
<p>Write one for loop to print out each character of the string <code class="docutils literal notranslate"><span class="pre">my_str</span></code> on a separate line.</p>

</div>
<textarea data-lang="python" id="assess_ch4_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
my_str = "MANITOBA"

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def test_output(self):
    self.assertIn('for', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
    self.assertIn("M\nA\nN\nI\nT\nO\nB\nA", self.getOutput(), "Testing output (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch4_2 data-question_label="4.17.2">
<div id=assess_ch4_2_question class="ac_question">
<p>Write one for loop to print out each element of the list <code class="docutils literal notranslate"><span class="pre">several_things</span></code>. Then, write <em>another</em> for loop to print out the TYPE of each element of the list <code class="docutils literal notranslate"><span class="pre">several_things</span></code>. To complete this problem you should have written two different for loops, each of which iterates over the list <code class="docutils literal notranslate"><span class="pre">several_things</span></code>, but each of those 2 for loops should have a different result.</p>

</div>
<textarea data-lang="python" id="assess_ch4_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
several_things = ["hello", 2, 4, 6.0, 7.5, 234352354, "the end", "", 99]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def test_output(self):
      self.assertIn('for', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
      str1 = "hello\n2\n4\n6.0\n7.5\n234352354\nthe end\n\n99\n<class 'str'>\n<class 'int'>\n<class 'int'>\n<class 'float'>\n<class 'float'>\n<class 'int'>\n<class 'str'>\n<class 'str'>\n<class 'int'>"
      str2 = "hello\n2\n4\n6.0\n7.5\n234352354\nthe end\n\n99\n<type 'str'>\n<type 'int'>\n<type 'int'>\n<type 'float'>\n<type 'float'>\n<type 'int'>\n<type 'str'>\n<type 'str'>\n<type 'int'>"
      self.assertTrue(str1 in self.getOutput() or str2 in self.getOutput(), "Testing output (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch4_3 data-question_label="4.17.3">
<div id=assess_ch4_3_question class="ac_question">
<p>Write code that uses iteration to print out <strong>the length</strong> of each element of the list stored in <code class="docutils literal notranslate"><span class="pre">str_list</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch4_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
str_list = ["hello", "", "goodbye", "wonderful", "I love Python"]

# Write your code here.
=====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def test_output(self):
        self.assertIn("for", self.getEditorText(), "Testing whether you used a for loop (Don't worry about actual and expected values).")
        self.assertIn("5\n0\n7\n9\n13", self.getOutput(), "Testing output (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch4_4 data-question_label="4.17.4">
<div id=assess_ch4_4_question class="ac_question">
<p>Write a program that uses the turtle module <strong>and</strong> a for loop to draw something. It doesn’t have to be complicated, but draw something different than we have done in the past. (Hint: if you are drawing something complicated, it could get tedious to watch it draw over and over. Try setting <code class="docutils literal notranslate"><span class="pre">.speed(10)</span></code> for the turtle to draw fast, or <code class="docutils literal notranslate"><span class="pre">.speed(0)</span></code> for it to draw super fast with no animation.)</p>

</div>
<textarea data-lang="python" id="assess_ch4_4_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import turtle
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch4_5 data-question_label="4.17.5">
<div id=assess_ch4_5_question class="ac_question">
<p>Write code to count the number of characters in <code class="docutils literal notranslate"><span class="pre">original_str</span></code> using the accumulation pattern and assign the answer to a variable <code class="docutils literal notranslate"><span class="pre">num_chars</span></code>. Do NOT use the <code class="docutils literal notranslate"><span class="pre">len</span></code> function to solve the problem (if you use it while you are working on this problem, comment it out afterward!)</p>

</div>
<textarea data-lang="python" id="assess_ch4_5_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
original_str = "The quick brown rhino jumped over the extremely lazy fox."


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertEqual(num_chars, len(original_str), "Testing whether num_chars_sent has the correct value")
       self.assertNotIn('len', self.getEditorText(), "Testing that you are not including the len function in your code. (Don't worry about Actual and Expected Values.)")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch4_7 data-question_label="4.17.6">
<div id=assess_ch4_7_question class="ac_question">
<p>Write code to create a list of numbers from 0 to 67 and assign that list to the variable <code class="docutils literal notranslate"><span class="pre">nums</span></code>. Do not hard code the list.</p>

</div>
<textarea data-lang="python" id="assess_ch4_7_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(nums, list(range(68)), "Testing that nums is a list that contains the correct elements.")

myTests().main()
</textarea>
</div>
</div>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="4.17.7" data-multipleanswers="false"  id="assess_ch4_8_lh"  style="visibility: hidden;">
    <blockquote>
<div><a class="reference internal image-reference" href="../_images/cdq4-1.png"><img alt="image of numerous variables being declared and assigned values" class="align-center" src="../_images/cdq4-1.png" style="width: 900px;" /></a>
<p>Given the following output, frame, and arrow indicators, what happens in the next step?</p>
</div></blockquote>

    <li data-component="answer"  id="assess_ch4_8_lh_opt_a">0 is printed</li><li data-component="feedback">Not quite, as shown, 0 has already printed</li>
    
    <li data-component="answer"  id="assess_ch4_8_lh_opt_b">1 is printed</li><li data-component="feedback">Not quite, The next line to execute is at the top of the for loop, not the print</li>
    
    <li data-component="answer" data-correct='yes' id="assess_ch4_8_lh_opt_c">The variable i checks for the end of the loop</li><li data-component="feedback">Correct! The code needs to check the list and see if there is anything left to do (line 1). Once it realizes it still has numbers to process, the loop continues.</li>
    
    <li data-component="answer"  id="assess_ch4_8_lh_opt_d">An error is thrown</li><li data-component="feedback">Not quite, this code runs properly, try it yourself.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="4.17.8" data-multipleanswers="false"  id="assess_ch4_9_lh"  style="visibility: hidden;">
    <p>:</p>
<a class="reference internal image-reference" href="../_images/cdq4-2.png"><img alt="image of numerous variables being declared and assigned values" class="align-center" src="../_images/cdq4-2.png" style="width: 900px;" /></a>
<p>The red arrow is still pointing to line 6 but we can see in the global frame that j is ‘o’. Why is the next line to execute line 6?</p>

    <li data-component="answer"  id="assess_ch4_9_lh_opt_a">Because 'o' gets printed twice</li><li data-component="feedback">Not quite, there is no reason for 'o' to print again</li>
    
    <li data-component="answer"  id="assess_ch4_9_lh_opt_b">Because the loop starts again</li><li data-component="feedback">Not quite, once a loop finishes running in its entirety, it can't hop back up to run again</li>
    
    <li data-component="answer" data-correct='yes' id="assess_ch4_9_lh_opt_c">Because the code is checking for another iteration</li><li data-component="feedback">Correct! The code needs to check in the list and see if there is anything left to do (line 6). Once it realizes there isn't, the code block completes</li>
    
    <li data-component="answer"  id="assess_ch4_9_lh_opt_d">Because there is an error in the code</li><li data-component="feedback">Not quite, this code runs properly, try it yourself.</li>
    

    </ul>
    </div>
    
        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="4.17.9" id="assess_ch4_10_lh"  style="visibility: hidden;">
            <p>:</p>
<a class="reference internal image-reference" href="../_images/cdq4-3.png"><img alt="image of numerous variables being declared and assigned values" class="align-center" src="../_images/cdq4-3.png" style="width: 900px;" /></a>
<p>The next line to execute is shown with the red arrow (line 4). What line of code will execute next?
<input type="text"></p>

        <script type="application/json">
            [[{"number": [6, 6], "feedback": "<p>Correct. The code goes to line 4 and realizes it\u2019s done looping over j so it then jumps down to the next line (6).</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Hmmm not quite, this one is tricky. The next line to execute will be the top of the inner for loop, but then what? What letter just printed?</p>\n"}]]
        </script>

        </div>
        </div>
            </section>


  

  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 4.16. Exercises' data-toggle="tooltip">
  <a href="Exercises.html" aria-label="Previous - 4.16. Exercises">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 5. Functions and Files' data-toggle="tooltip" >
    <a href="../FunctionsAndFiles/toctree.html" aria-label="Next - 5. Functions and Files">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 7.2.11.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>