<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html  lang="en" class="runestone-sphinx">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>2.17. Exercises &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=E2DC0A34" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=4182B24D" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-347.f9add1ca35d5ad93.css?v=306CBB19" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-runestone.f64dcc3632d7a0c3.css?v=F4F91F21" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script src="../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script src="../_static/sphinx_highlight.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/prefix-runtime.2438a8bdd6d43b5c.bundle.js?v=7D29D67B"></script>
    <script defer="" src="../_static/prefix-347.8328b60515045466.bundle.js?v=2CF09D3E"></script>
    <script defer="" src="../_static/prefix-runestone.aeac7cf9ad5c07ef.bundle.js?v=4FF65C30"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
<script src="https://media.ethicalads.io/media/client/ethicalads.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="2.18. Chapter 2 Assessment" href="week1a2.html" />
    <link rel="prev" title="2.16. Glossary" href="Glossary.html" />

  
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'FOPP-PIE';
    eBookConfig.basecourse = 'FOPP-PIE';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '7.2.11';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" id="rs-book-toc-link" href="../index.html" aria-label="index-page">
          
          FOPP-PIE
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="ad_warning"></span></li>
        <li><span id="browsing_warning"></span></li>

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://guide.runestone.academy' target="_blank" rel="noopener noreferrer" >Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=FOPP-PIE&page=Exercises'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">2.16. </span>Glossary" data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">2.18. </span>Chapter 2 Assessment' data-toggle="tooltip" >              
    <a href="week1a2.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <p></p>
<section id="exercises">
<h1><span class="section-number">2.17. </span>Exercises<a class="headerlink" href="#exercises" title="Permalink to this heading">¶</a></h1>

    <div data-component="question" class="full-width container question " data-question_label="2.17.1" id="spd_ex_1" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    <div id='q1' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ch02_ex1 data-question_label="2.17.2">
<div id=ch02_ex1_question class="ac_question">
<p>Evaluate the following numerical expressions in your head, then use
the active code window to check your results:</p>
<ol class="arabic simple">
<li><p><code class="docutils literal notranslate"><span class="pre">5</span> <span class="pre">**</span> <span class="pre">2</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">9</span> <span class="pre">*</span> <span class="pre">5</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">15</span> <span class="pre">/</span> <span class="pre">12</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">12</span> <span class="pre">/</span> <span class="pre">15</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">15</span> <span class="pre">//</span> <span class="pre">12</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">12</span> <span class="pre">//</span> <span class="pre">15</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">5</span> <span class="pre">%</span> <span class="pre">2</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">9</span> <span class="pre">%</span> <span class="pre">5</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">15</span> <span class="pre">%</span> <span class="pre">12</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">12</span> <span class="pre">%</span> <span class="pre">15</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">6</span> <span class="pre">%</span> <span class="pre">6</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">0</span> <span class="pre">%</span> <span class="pre">7</span></code></p></li>
</ol>

</div>
<textarea data-lang="python" id="ch02_ex1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

print(5**2)
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >
<ol class="arabic simple">
<li><p><code class="docutils literal notranslate"><span class="pre">5</span> <span class="pre">**</span> <span class="pre">2</span>&#160; <span class="pre">=</span> <span class="pre">25</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">9</span> <span class="pre">*</span> <span class="pre">5</span> <span class="pre">=</span> <span class="pre">45</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">15</span> <span class="pre">/</span> <span class="pre">12</span> <span class="pre">=</span> <span class="pre">1.25</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">12</span> <span class="pre">/</span> <span class="pre">15</span> <span class="pre">=</span> <span class="pre">0.8</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">15</span> <span class="pre">//</span> <span class="pre">12</span> <span class="pre">=</span> <span class="pre">1</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">12</span> <span class="pre">//</span> <span class="pre">15</span> <span class="pre">=</span> <span class="pre">0</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">5</span> <span class="pre">%</span> <span class="pre">2</span> <span class="pre">=</span> <span class="pre">1</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">9</span> <span class="pre">%</span> <span class="pre">5</span> <span class="pre">=</span> <span class="pre">4</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">15</span> <span class="pre">%</span> <span class="pre">12</span> <span class="pre">=</span> <span class="pre">3</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">12</span> <span class="pre">%</span> <span class="pre">15</span> <span class="pre">=</span> <span class="pre">12</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">6</span> <span class="pre">%</span> <span class="pre">6</span> <span class="pre">=</span> <span class="pre">0</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">0</span> <span class="pre">%</span> <span class="pre">7</span> <span class="pre">=</span> <span class="pre">0</span></code></p></li>
</ol>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function c0a62044cac248859ce3695b46697ecc_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'c0a62044cac248859ce3695b46697ecc_' + eBookConfig.course;
                    this.page.url = 'https://c0a62044cac248859ce3695b46697ecc_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'c0a62044cac248859ce3695b46697ecc_' + eBookConfig.course;
            disqus_url = 'http://c0a62044cac248859ce3695b46697ecc_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="c0a62044cac248859ce3695b46697ecc" onclick="c0a62044cac248859ce3695b46697ecc_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='c0a62044cac248859ce3695b46697ecc']").attr('data-disqus-identifier', 'c0a62044cac248859ce3695b46697ecc_' + eBookConfig.course);
  });
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.3" id="spd_ex_2" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_2_2 data-question_label="2.17.4">
<div id=ex_2_2_question class="ac_question">
<p>What is the order of the arithmetic operations in the following expression.  Evaluate the expression by hand and then check your work.</p>

</div>
<textarea data-lang="python" id="ex_2_2_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
2 + (3 - 1) * 10 / 5 * (2 + 3)
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.5" id="spd_ex_5" >
    <ol start=3 class=arabic><li class="alert alert-warning">

    <div id='q5' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_2_5 data-question_label="2.17.6">
<div id=ex_2_5_question class="ac_question">
<p>Challenge: Many people keep time using a 24 hour clock (11 is 11am and 23 is 11pm, 0 is midnight).
If it is currently 13 and you set your alarm to go off in 50 hours, it will be 15 (3pm).
Write a Python program to solve the general version of the above problem.
Ask the user for the time now (in hours), and then ask for the number of hours to wait for the alarm.
Your program should output what the time will be on the clock when the alarm goes off.</p>

</div>
<textarea data-lang="python" id="ex_2_5_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=spd_q5_answer data-question_label="2.17.7">
<div id=spd_q5_answer_question class="ac_question">

</div>
<textarea data-lang="python" id="spd_q5_answer_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
## question 5 solution ##

current_time_string = input("What is the current time (in hours)? ")
waiting_time_string = input("How many hours do you have to wait? ")

current_time_int = int(current_time_string)
waiting_time_int = int(waiting_time_string)

hours = current_time_int + waiting_time_int

timeofday = hours % 24

print(timeofday)
</textarea>
</div>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.8" id="spd_ex_6" >
    <ol start=4 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_2_6 data-question_label="2.17.9">
<div id=ex_2_6_question class="ac_question">
<p>It is possible to name the days 0 thru 6 where day 0 is Sunday and day 6 is Saturday.
If you go on a wonderful holiday leaving on day number 3 (a Wednesday) and you return
home after 10 nights you would return home on a Saturday (day 6). Write a general
version of the program which asks for the starting day number, and the length of your
stay, and it will tell you the number of day of the week you will return on.</p>

</div>
<textarea data-lang="python" id="ex_2_6_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.10" id="spd_ex_7" >
    <ol start=5 class=arabic><li class="alert alert-warning">

    <div id='q7' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >
<p>Challenge: Take the sentence: <em>All work and no play makes Jack a dull boy.</em>
Store each word in a separate variable, then print out the sentence on
one line using <code class="docutils literal notranslate"><span class="pre">print</span></code>.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_2_7 data-question_label="2.17.11">
<div id=ex_2_7_question class="ac_question">

</div>
<textarea data-lang="python" id="ex_2_7_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">


</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=spd_q7_answer data-question_label="2.17.12">
<div id=spd_q7_answer_question class="ac_question">

</div>
<textarea data-lang="python" id="spd_q7_answer_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
## question 7 solution ##

word1 = "All"
word2 = "work"
word3 = "and"
word4 = "no"
word5 = "play"
word6 = "makes"
word7 = "Jack"
word8 = "a"
word9 = "dull"
word10 = "boy."

print(word1 + ' ' +  word2 + ' ' + word3 + ' ' +  word4 + ' ' +  word5 + ' ' +  word6 + ' ' +  word7 + ' ' +  word8 + ' ' +  word9 + ' ' +  word10)
</textarea>
</div>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.13" id="spd_ex_8" >
    <ol start=6 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_2_8 data-question_label="2.17.14">
<div id=ex_2_8_question class="ac_question">
<p>Add parentheses to the expression <code class="docutils literal notranslate"><span class="pre">6</span> <span class="pre">*</span> <span class="pre">1</span> <span class="pre">-</span> <span class="pre">2</span></code> to change its value from 4 to -6.</p>

</div>
<textarea data-lang="python" id="ex_2_8_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

print(6 * 1 -2)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual('-6\n', self.getOutput(), 'Checking output.')
myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.15" id="spd_ex_9" >
    <ol start=7 class=arabic><li class="alert alert-warning">

    <div id='q9' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_2_9 data-question_label="2.17.16">
<div id=ex_2_9_question class="ac_question">
<p>Challenge: The formula for computing the final amount if one is earning
compound interest is given on Wikipedia as</p>
<img alt="formula for compound interest" src="../_images/compoundInterest.png" />
<p>Write a Python program that assigns the principal amount of 10000 to
variable <code class="docutils literal notranslate"><span class="pre">P</span></code>, assign to <code class="docutils literal notranslate"><span class="pre">n</span></code> the value 12, and assign to <code class="docutils literal notranslate"><span class="pre">r</span></code> the interest
rate of 8% (0.08).  Then have the program prompt the user for the number of years,
<code class="docutils literal notranslate"><span class="pre">t</span></code>, that the money will be compounded for.  Calculate and print the final
amount after <code class="docutils literal notranslate"><span class="pre">t</span></code> years.</p>

</div>
<textarea data-lang="python" id="ex_2_9_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

P = 10000
n = 12
r = 0.08

t = ??

====

from unittest.gui import TestCaseGui
import re
class myTests(TestCaseGui):
    def testOne(self):
        self.assertTrue(re.search(str(P * (1 + (r/n))**(n*t))[:5],
            self.getOutput()),
            'Checking Answer.')

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=spd_q9_answer data-question_label="2.17.17">
<div id=spd_q9_answer_question class="ac_question">

</div>
<textarea data-lang="python" id="spd_q9_answer_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
## question 9 solution ##

P = 10000
n = 12
r = 0.08

t = int(input("Compound for how many years? "))

final = P * ( ((1 + (r/n)) ** (n * t)) )

print("The final amount after", t, "years is", final)
</textarea>
</div>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.18" id="spd_ex_10" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_2_10 data-question_label="2.17.19">
<div id=ex_2_10_question class="ac_question">
<p>Write a program that will compute the area of a circle. Prompt the user to enter the radius and save it to a variable called <code class="docutils literal notranslate"><span class="pre">radius</span></code>. Print a nice message back to the user with the answer.</p>

</div>
<textarea data-lang="python" id="ex_2_10_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
pi = 3.14

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def test_radius(self):
        self.assertIn("radius", self.getEditorText(), 'radius variable')
        self.assertIn(str(3.14*radius**2)[:4], self.getOutput(), 'Checking answer.')
myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.20" id="spd_ex_11" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    <div id='q11' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_2_11 data-question_label="2.17.21">
<div id=ex_2_11_question class="ac_question">
<p>Challenge: Write a program that will compute the area of a rectangle.  Prompt the user to enter the width and height of the rectangle  and store the values in variables called <code class="docutils literal notranslate"><span class="pre">width</span></code> and
<code class="docutils literal notranslate"><span class="pre">height</span></code>. Print a nice message with the answer..</p>

</div>
<textarea data-lang="python" id="ex_2_11_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertIn('width', self.getEditorText())
        self.assertIn('height', self.getEditorText())
        self.assertIn(str(width*height), self.getOutput(), 'Checking answer.')
myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=spd_q11_answer data-question_label="2.17.22">
<div id=spd_q11_answer_question class="ac_question">

</div>
<textarea data-lang="python" id="spd_q11_answer_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
## question 11 solution

width = int(input("Width? "))
height = int(input("Height? "))

area = width * height

print("The area of the rectangle is", area)
</textarea>
</div>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.23" id="spd_ex_12" >
    <ol start=10 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_2_12 data-question_label="2.17.24">
<div id=ex_2_12_question class="ac_question">
<p>Write a program that will compute MPG for a car. Prompt the user to enter the number of
miles driven and store it in a variable called <cite>miles</cite> and the number of gallons used stored in a variable <cite>gallons</cite>.
Print a nice message with the answer.</p>

</div>
<textarea data-lang="python" id="ex_2_12_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
import re
class myTests(TestCaseGui):
    def testOne(self):
       self.assertTrue(re.search(str(miles/gallons)[:5], self.getOutput()), 'Checking answer.')
myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.25" id="spd_ex_13" >
    <ol start=11 class=arabic><li class="alert alert-warning">

    <div id='q13' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_2_13 data-question_label="2.17.26">
<div id=ex_2_13_question class="ac_question">
<p>Challenge: Write a program that will convert degrees celsius to degrees fahrenheit.</p>

</div>
<textarea data-lang="python" id="ex_2_13_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=spd_q13_answer data-question_label="2.17.27">
<div id=spd_q13_answer_question class="ac_question">

</div>
<textarea data-lang="python" id="spd_q13_answer_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
## question 13 solution ##

deg_c = int(input("What is the temperature in Celsius? "))

# formula to convert C to F is: (degrees Celcius) times (9/5) plus (32)
deg_f = deg_c * (9 / 5) + 32

print(deg_c, " degrees Celsius is", deg_f, " degrees Farenheit.")
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function c4a929d598ab4c46b484f6abbcec2655_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'c4a929d598ab4c46b484f6abbcec2655_' + eBookConfig.course;
                    this.page.url = 'https://c4a929d598ab4c46b484f6abbcec2655_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'c4a929d598ab4c46b484f6abbcec2655_' + eBookConfig.course;
            disqus_url = 'http://c4a929d598ab4c46b484f6abbcec2655_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="c4a929d598ab4c46b484f6abbcec2655" onclick="c4a929d598ab4c46b484f6abbcec2655_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='c4a929d598ab4c46b484f6abbcec2655']").attr('data-disqus-identifier', 'c4a929d598ab4c46b484f6abbcec2655_' + eBookConfig.course);
  });
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.28" id="spd_ex_14" >
    <ol start=12 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_2_14 data-question_label="2.17.29">
<div id=ex_2_14_question class="ac_question">
<p>Ask the user for the temperature in Fahrenheit and store it in a variable call <code class="docutils literal notranslate"><span class="pre">deg_f</span></code>. Calculate the equivalent temperature in degrees Celsius and store it in <code class="docutils literal notranslate"><span class="pre">deg_c</span></code>. Output a message to the user giving the temperature in Celsius.</p>

</div>
<textarea data-lang="python" id="ex_2_14_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertIn('deg_f', self.getEditorText())
        self.assertIn('deg_c', self.getEditorText())
        ans = (deg_f - 32)*5/9
        self.assertAlmostEqual(ans, deg_c)
        self.assertIn(str(ans)[:min(len(str(ans)), 5)],    self.getOutput(), 'Checking answer.')
myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.30" id="spd_ex_15" >
    <ol start=13 class=arabic><li class="alert alert-warning">

    
        <div class="runestone parsons-container ">
        <div data-component="parsons" id="pp_2_15" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Piece together the code so that a user is asked for two numbers, and then the sum of those two numbers is printed out.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2.17.31"      data-noindent="true"     style="visibility: hidden;">
        num_one = input("Please enter your first number: ")
---
num_two = input("Please enter your second number: ")
---
sum_of_input = int(num_one) + int(num_two)
---
print(sum_of_input)
        </pre>
        </div>
        </div>
    
    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.32" id="spd_ex_16" >
    <ol start=14 class=arabic><li class="alert alert-warning">

    
        <div class="runestone parsons-container ">
        <div data-component="parsons" id="pp_2_16" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write a program that will convert gallons to liters. This program will also need to get input from a user to see how many gallons should be converted and the result should be printed to the user.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2.17.33"      data-noindent="true"     style="visibility: hidden;">
        user_gallons = input("How many gallons should be converted?: ")
---
num_gallons = int(user_gallons)
---
liters = num_gallons * 3.785
---
print("Number of liters: " + str(liters))
        </pre>
        </div>
        </div>
    
    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="2.17.34" id="spd_ex_17" >
    <ol start=15 class=arabic><li class="alert alert-warning">

    
        <div class="runestone parsons-container ">
        <div data-component="parsons" id="pp_2_17" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write a program that will convert tablespoons to teaspoons. This program will also need to get input from a user to see how many tablespoons should be converted and the result should be printed to the user.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2.17.35"      data-noindent="true"     style="visibility: hidden;">
        user_tablespoons = float(input("How many tablespoons should be converted?: "))
---
teaspoons = user_tablespoons * 3
---
print("Number of teaspoons: " + str(teaspoons))
        </pre>
        </div>
        </div>
    
    </li></ol>
    </div>
    </section>


  

  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 2.16. Glossary' data-toggle="tooltip">
  <a href="Glossary.html" aria-label="Previous - 2.16. Glossary">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 2.18. Chapter 2 Assessment' data-toggle="tooltip" >
    <a href="week1a2.html" aria-label="Next - 2.18. Chapter 2 Assessment">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 7.2.11.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>