<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html  lang="en" class="runestone-sphinx">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>2.18. Chapter 2 Assessment &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=E2DC0A34" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=4182B24D" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-347.f9add1ca35d5ad93.css?v=306CBB19" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-runestone.f64dcc3632d7a0c3.css?v=F4F91F21" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script src="../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script src="../_static/sphinx_highlight.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/prefix-runtime.2438a8bdd6d43b5c.bundle.js?v=7D29D67B"></script>
    <script defer="" src="../_static/prefix-347.8328b60515045466.bundle.js?v=2CF09D3E"></script>
    <script defer="" src="../_static/prefix-runestone.aeac7cf9ad5c07ef.bundle.js?v=4FF65C30"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
<script src="https://media.ethicalads.io/media/client/ethicalads.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="3. Debugging and Modules" href="../DebuggingAndModules/toctree.html" />
    <link rel="prev" title="2.17. Exercises" href="Exercises.html" />

  
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'FOPP-PIE';
    eBookConfig.basecourse = 'FOPP-PIE';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '7.2.11';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" id="rs-book-toc-link" href="../index.html" aria-label="index-page">
          
          FOPP-PIE
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="ad_warning"></span></li>
        <li><span id="browsing_warning"></span></li>

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://guide.runestone.academy' target="_blank" rel="noopener noreferrer" >Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=FOPP-PIE&page=week1a2'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">2.17. </span>Exercises" data-toggle="tooltip">
    <a href="Exercises.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">3. </span>Debugging and Modules' data-toggle="tooltip" >              
    <a href="../DebuggingAndModules/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <section id="chapter-2-assessment">
<h1><span class="section-number">2.18. </span>Chapter 2 Assessment<a class="headerlink" href="#chapter-2-assessment" title="Permalink to this heading">¶</a></h1>
<p><strong>Check your understanding</strong></p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch2_1 data-question_label="2.18.1">
<div id=assess_ch2_1_question class="ac_question">
<p>There is a function we are providing in for you in this problem called <code class="docutils literal notranslate"><span class="pre">square</span></code>. It takes one integer and returns the square of that integer value. Write code to assign a variable called <code class="docutils literal notranslate"><span class="pre">xyz</span></code> the value <code class="docutils literal notranslate"><span class="pre">5*5</span></code> (five squared). Use the square function, rather than just multiplying with <code class="docutils literal notranslate"><span class="pre">*</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch2_1_editor" 
     data-include="assess_Ch2_addl_functions" data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
xyz = ""

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(type(xyz), type(3), "Checking type of xyz")
        self.assertEqual(xyz, 25, "Checking if xyz is 25")
        self.assertIn('square', self.getEditorText(), "Testing that 'square' is in your code. (Don't worry about Actual and Expected Values.)")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch2_2 data-question_label="2.18.2">
<div id=assess_ch2_2_question class="ac_question">
<p>Write code to assign the number of <em>characters</em> in the string <code class="docutils literal notranslate"><span class="pre">rv</span></code> to a variable <code class="docutils literal notranslate"><span class="pre">num_chars</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch2_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
rv = """Once upon a midnight dreary, while I pondered, weak and weary,
    Over many a quaint and curious volume of forgotten lore,
    While I nodded, nearly napping, suddenly there came a tapping,
    As of some one gently rapping, rapping at my chamber door.
    'Tis some visitor, I muttered, tapping at my chamber door;
    Only this and nothing more."""

# Write your code here!

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertEqual(num_chars, len(rv), "Testing that num_chars has been set to the length of rv")

myTests().main()
</textarea>
</div>
</div>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.18.3" data-multipleanswers="true"  id="assess_ch2_3"  style="visibility: hidden;">
    <p>The code below initializes two variables, <code class="docutils literal notranslate"><span class="pre">z</span></code> and <code class="docutils literal notranslate"><span class="pre">y</span></code>. We want to assign the total number of characters in <code class="docutils literal notranslate"><span class="pre">z</span></code> and in <code class="docutils literal notranslate"><span class="pre">y</span></code> to the variable <code class="docutils literal notranslate"><span class="pre">a</span></code>. Which of the following solutions, if any, would be considered hard coding?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">z</span> <span class="o">=</span> <span class="s2">&quot;hello world&quot;</span>
<span class="n">y</span> <span class="o">=</span> <span class="s2">&quot;welcome!&quot;</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="assess_ch2_3_opt_a">a = len("hello worldwelcome!")</li><li data-component="feedback">Though we are using the len function here, we are hardcoding what len should return the length of. We are not referencing z or y.</li>
    
    <li data-component="answer" data-correct='yes' id="assess_ch2_3_opt_b">a = 11 + 8</li><li data-component="feedback">This is hardcoding, we are writing in the value without referencing z or y.</li>
    
    <li data-component="answer"  id="assess_ch2_3_opt_c">a = len(z) + len(y)</li><li data-component="feedback">This is not considered hard coding. We are using the function len to determine the length of what is stored in z and y, which is a correct way to approach this problem.</li>
    
    <li data-component="answer" data-correct='yes' id="assess_ch2_3_opt_d">a = len("hello world") + len("welcome!")</li><li data-component="feedback">Though we are using the len function here, we are hardcoding what len should return the length of each time we call len. We are not referencing z or y.</li>
    
    <li data-component="answer"  id="assess_ch2_3_opt_e">none of the above are hardcoding.</li><li data-component="feedback">At least one of these solutions is considered hardcoding. Take another look.</li>
    

    </ul>
    </div>
    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch2_4 data-question_label="2.18.4">
<div id=assess_ch2_4_question class="ac_question">
<p>(This is not an assessment question) The code below defines functions used by one of the questions above. Do not modify the code, but feel free to take a look.</p>

</div>
<textarea data-lang="python" id="assess_ch2_4_editor" 
    data-hidecode="true"  data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

def square(num):
    return num**2
</textarea>
</div>
</div>

        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="2.18.5" id="assess_ch2_5_lh"  style="visibility: hidden;">
            <blockquote>
<div><a class="reference internal image-reference" href="../_images/cdq2-1.png"><img alt="image of numerous variables being declared and assigned values" class="align-center" src="../_images/cdq2-1.png" style="width: 900px;" /></a>
</div></blockquote>
<p>Which variables in the image above are assigned string values? List each <em>variable</em> in the box below, separated by a white space.</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*a\\s+e\\s+f\\s+g\\s+h\\s+i\\s*$", "regexFlags": "", "feedback": "<p>Correct. You typically use whole numbers for ages after age 1.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>There are 6 in total, can you find them all? Remember that variables are case sensitive.</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="2.18.6" id="assess_ch2_6_lh"  style="visibility: hidden;">
            <blockquote>
<div><a class="reference internal image-reference" href="../_images/cdq2-2.png"><img alt="image of numerous variables being declared and assigned values" class="align-center" src="../_images/cdq2-2.png" style="width: 900px;" /></a>
</div></blockquote>
<p>Which line number(s) will print an integer as the result of the given operation? List each line number in the box below, separated by a COMMA.</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*3,6\\s*$", "regexFlags": "", "feedback": "<p>Correct! All the others either use floats or integers with the \u2018/\u2019 operation which always returns a float.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Not quite. Try it out yourself to check your understanding of \u2018/\u2019 and \u2018//\u2019. Make sure you format your answer as numbers separated by spaces as well.</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="2.18.7" id="assess_ch2_7_lh"  style="visibility: hidden;">
            <blockquote>
<div><a class="reference internal image-reference" href="../_images/cdq2-3.png"><img alt="image of numerous variables being declared and assigned values" class="align-center" src="../_images/cdq2-3.png" style="width: 900px;" /></a>
</div></blockquote>
<p>The image above shows a program that just completed executing (we know this because the green arrow is on the last line and there is no red arrow in the image). At the end of the program, which variables changed from their original value during the program’s execution? List each <em>variable</em> in the box below, separated by a white space.</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^\\s*a\\s+c\\s*$", "regexFlags": "", "feedback": "<p>Correct! Even though c finished with the same value it started with, it still was updated during the code execution.</p>\n"}, {"regex": "^\\s*a\\s+c\\s+e\\s*$", "regexFlags": "", "feedback": "<p>Not quite. Variable e is made up of values from a and b but it never gets reassigned.</p>\n"}, {"regex": "^\\s*a\\s+e\\s*$", "regexFlags": "", "feedback": "<p>Not quite. Variable e is made up of values from a and b but it never gets reassigned. Read through the program again to double check which variables get new values during code execution.</p>\n"}, {"regex": "^\\s*a\\s*$", "regexFlags": "", "feedback": "<p>Almost. Variable a is definitely reassigned during the program but is that the only one? Go back and read the code again to see if there is any others</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Not quite. Read through the program again to see which variables are updated during the code execution.</p>\n"}]]
        </script>

        </div>
        </div>
            
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.18.8" data-multipleanswers="false"  id="assess_ch2_8_lh"  style="visibility: hidden;">
    <p>:</p>
<img alt="../_images/cdq2-4.png" src="../_images/cdq2-4.png" />
<p>What does the following code shown in the codelens image print when it finishes executing?</p>

    <li data-component="answer" data-correct='yes' id="assess_ch2_8_lh_opt_a">an error</li><li data-component="feedback">This is a correct. When using + to stick strings together, all variables need to be a string. This throws a type error.</li>
    
    <li data-component="answer"  id="assess_ch2_8_lh_opt_b">sara 11</li><li data-component="feedback">Hmm, this would be true if 11 was a string but it isn't</li>
    
    <li data-component="answer"  id="assess_ch2_8_lh_opt_c">sara11</li><li data-component="feedback">Hmm, not quite, we stick a space on our string in line 3</li>
    
    <li data-component="answer"  id="assess_ch2_8_lh_opt_d">"sara11"</li><li data-component="feedback">Hmm, not quite, we stick a space on our string in line 3 and we don't have quotes</li>
    
    <li data-component="answer"  id="assess_ch2_8_lh_opt_e">"sara 11"</li><li data-component="feedback">Hmm, we don't have quotes in our string</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.18.9" data-multipleanswers="false"  id="assess_ch2_9_lh"  style="visibility: hidden;">
    <p>:</p>
<img alt="../_images/cdq2-5.png" src="../_images/cdq2-5.png" />
<p>Here is that same image of code again. Which line number shows a line of code we would define as an expression?</p>

    <li data-component="answer"  id="assess_ch2_9_lh_opt_a">1</li><li data-component="feedback">This is a variable being asssigned a string value</li>
    
    <li data-component="answer"  id="assess_ch2_9_lh_opt_b">2</li><li data-component="feedback">This is a variable being assigned an integer value</li>
    
    <li data-component="answer" data-correct='yes' id="assess_ch2_9_lh_opt_c">3</li><li data-component="feedback">Correct, "sara" + " " + age is an expression that we want to evaluate to "sara 11" and store the string in the info variable</li>
    
    <li data-component="answer"  id="assess_ch2_9_lh_opt_d">4</li><li data-component="feedback">This prints the string value stored the info variable</li>
    

    </ul>
    </div>
    </section>


  

  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 2.17. Exercises' data-toggle="tooltip">
  <a href="Exercises.html" aria-label="Previous - 2.17. Exercises">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 3. Debugging and Modules' data-toggle="tooltip" >
    <a href="../DebuggingAndModules/toctree.html" aria-label="Next - 3. Debugging and Modules">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 7.2.11.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>