<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>11.2. The while Statement &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.fafafbd97df8a0d1.css?v=B5361BC8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.d062db89aaf1dddd.css?v=B02DFDA" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/runtime.b6a5f441f397ace5.bundle.js?v=3F299E7"></script>
    <script defer="" src="../_static/637.d54be67956c5c660.bundle.js?v=915B1670"></script>
    <script defer="" src="../_static/runestone.f2a5bb1b0f5cce0b.bundle.js?v=871C0DAB"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="11.3. The Listener Loop" href="listenerLoop.html" />
    <link rel="prev" title="11.1. Introduction" href="intro-indefiniteiteration.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'COMP1000UofM';
    eBookConfig.basecourse = 'COMP1000UofM';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = false;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '6.0.10';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          COMP1000UofM
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=COMP1000UofM&page=ThewhileStatement'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">11.1. </span>Introduction" data-toggle="tooltip">
    <a href="intro-indefiniteiteration.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">11.3. </span>The Listener Loop' data-toggle="tooltip" >              
    <a href="listenerLoop.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="the-while-statement">
<h1><span class="section-number">11.2. </span>The <code class="docutils literal notranslate"><span class="pre">while</span></code> Statement<a class="headerlink" href="#the-while-statement" title="Permalink to this headline">¶</a></h1>

    <div class="runestone" style="margin-left: auto; margin-right:auto">
        <div id="whileloop" data-component="youtube" class="align-left youtube-video"
            data-video-height="315" data-question_label="11.2.1"
            data-video-width="560" data-video-videoid="iUbToOoAykE"
            data-video-divid="whileloop" data-video-start="0"
            data-video-end="-1" ></div>
    </div>
    <p>There is another Python statement that can also be used to build an iteration. It is called the <code class="docutils literal notranslate"><span class="pre">while</span></code> statement.
The <code class="docutils literal notranslate"><span class="pre">while</span></code> statement provides a much more general mechanism for iterating. Similar to the <code class="docutils literal notranslate"><span class="pre">if</span></code> statement, it uses
a boolean expression to control the flow of execution. The body of while will be repeated as long as the controlling
boolean expression evaluates to <code class="docutils literal notranslate"><span class="pre">True</span></code>.</p>
<p>The following two figures show the flow of control. The first focuses on the flow inside the while loop and the second
shows the while loop in context.</p>
<img alt="a diamond at the top has the phrase &quot;Is the condition True?&quot;. Two arrows come out it with either the phrase yes or no on the arrows. The yes arrow points to a box that says &quot;evaluate the statemenets in the body of the loop&quot;. It then has an arrow that unconditionally points back to &quot;Is the condition True?&quot; diamond. The no arrow escapes the loop and points down past the &quot;evaluate&quot; square." src="../_images/while_flow.png" />
<img alt="image showing a rectangle with &quot;code block&quot; written on it on top. Then, text that read &quot;while {condition}&quot;: followed by an indented block with &quot;run if {condition} is True&quot; written on it. An arrow points from the bottom of the indented block to the top of the while loop and says &quot;run loop again&quot;. At the bottom of the image is an unindented block that says the phrase &quot;code block.&quot;" src="../_images/while_loop.png" />
<p>We can use the <code class="docutils literal notranslate"><span class="pre">while</span></code> loop to create any type of iteration we wish, including anything that we have previously
done with a <code class="docutils literal notranslate"><span class="pre">for</span></code> loop. For example, the program in the previous section could be rewritten using <code class="docutils literal notranslate"><span class="pre">while</span></code>.
Instead of relying on the <code class="docutils literal notranslate"><span class="pre">range</span></code> function to produce the numbers for our summation, we will need to produce them
ourselves. To do this, we will create a variable called <code class="docutils literal notranslate"><span class="pre">aNumber</span></code> and initialize it to 1, the first number in the
summation. Every iteration will add <code class="docutils literal notranslate"><span class="pre">aNumber</span></code> to the running total until all the values have been used. In order to
control the iteration, we must create a boolean expression that evaluates to <code class="docutils literal notranslate"><span class="pre">True</span></code> as long as we want to keep
adding values to our running total. In this case, as long as <code class="docutils literal notranslate"><span class="pre">aNumber</span></code> is less than or equal to the bound, we
should keep going.</p>
<p>Here is a new version of the summation program that uses a while statement.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac14_2_1 data-question_label="11.2.2">
<div id=ac14_2_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac14_2_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def sumTo(aBound):
    """ Return the sum of 1+2+3 ... n """

    theSum  = 0
    aNumber = 1
    while aNumber <= aBound:
        theSum = theSum + aNumber
        aNumber = aNumber + 1
    return theSum

print(sumTo(4))

print(sumTo(1000))
</textarea>
</div>
</div>
<p>You can almost read the <code class="docutils literal notranslate"><span class="pre">while</span></code> statement as if it were in natural language. It means, while <code class="docutils literal notranslate"><span class="pre">aNumber</span></code> is less
than or equal to <code class="docutils literal notranslate"><span class="pre">aBound</span></code>, continue executing the body of the loop. Within the body, each time, update <code class="docutils literal notranslate"><span class="pre">theSum</span></code>
using the accumulator pattern and increment <code class="docutils literal notranslate"><span class="pre">aNumber</span></code>. After the body of the loop, we go back up to the condition
of the <code class="docutils literal notranslate"><span class="pre">while</span></code> and reevaluate it. When <code class="docutils literal notranslate"><span class="pre">aNumber</span></code> becomes greater than <code class="docutils literal notranslate"><span class="pre">aBound</span></code>, the condition fails and flow
of control continues to the <code class="docutils literal notranslate"><span class="pre">return</span></code> statement.</p>
<p>The same program in codelens will allow you to observe the flow of execution.</p>

<div class="runestone" style="max-width: none;">
    <div class="alert alert-warning cd_section" data-component="codelens" data-question_label="11.2.3">
        <div class="pytutorVisualizer" id="clens14_2_1"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 11.2.3 (clens14_2_1)</span> </p>
    </div>

<script>
var clens14_2_1_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["clens14_2_1"] = {"code": "def sumTo(aBound):\n    \"\"\" Return the sum of 1+2+3 ... n \"\"\"\n\n    theSum  = 0\n    aNumber = 1\n    while aNumber <= aBound:\n        theSum = theSum + aNumber\n        aNumber = aNumber + 1\n    return theSum\n\nprint(sumTo(4))", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 11, "event": "step_line", "func_name": "<module>", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4}, "ordered_varnames": ["aBound"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4}, "ordered_varnames": ["aBound"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 0}, "ordered_varnames": ["aBound", "theSum"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 0, "aNumber": 1}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 7, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 0, "aNumber": 1}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 8, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 1, "aNumber": 1}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 1, "aNumber": 2}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 7, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 1, "aNumber": 2}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 8, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 3, "aNumber": 2}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 3, "aNumber": 3}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 7, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 3, "aNumber": 3}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 8, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 6, "aNumber": 3}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 6, "aNumber": 4}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 7, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 6, "aNumber": 4}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 8, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 10, "aNumber": 4}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 10, "aNumber": 5}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 9, "event": "step_line", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 10, "aNumber": 5}, "ordered_varnames": ["aBound", "theSum", "aNumber"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 9, "event": "return", "func_name": "sumTo", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [{"func_name": "sumTo", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"aBound": 4, "theSum": 10, "aNumber": 5, "__return__": 10}, "ordered_varnames": ["aBound", "theSum", "aNumber", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "sumTo_f1"}], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": ""}, {"line": 11, "event": "return", "func_name": "<module>", "globals": {"sumTo": ["REF", 1]}, "ordered_globals": ["sumTo"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "sumTo(aBound)", null]}, "stdout": "10\n"}]};
</script>
</div>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The names of the variables have been chosen to help readability.</p>
</div>
<p>More formally, here is the flow of execution for a <code class="docutils literal notranslate"><span class="pre">while</span></code> statement:</p>
<ol class="arabic simple">
<li><p>Evaluate the condition, yielding <code class="docutils literal notranslate"><span class="pre">False</span></code> or <code class="docutils literal notranslate"><span class="pre">True</span></code>.</p></li>
<li><p>If the condition is <code class="docutils literal notranslate"><span class="pre">False</span></code>, exit the <code class="docutils literal notranslate"><span class="pre">while</span></code> statement and continue
execution at the next statement.</p></li>
<li><p>If the condition is <code class="docutils literal notranslate"><span class="pre">True</span></code>, execute each of the statements in the body and
then go back to step 1.</p></li>
</ol>
<p>The body consists of all of the statements below the header with the same indentation.</p>
<p>This type of flow is called a <strong>loop</strong> because the third step loops back around to the top. Notice that if the
condition is <code class="docutils literal notranslate"><span class="pre">False</span></code> the first time through the loop, the statements inside the loop are never executed.</p>
<p>The body of the loop should change the value of one or more variables so that eventually the condition becomes
<code class="docutils literal notranslate"><span class="pre">False</span></code> and the loop terminates. Otherwise the loop will repeat forever. This is called an <strong>infinite loop</strong>.
An endless source of amusement for computer scientists is the observation that the directions written on the back of
the shampoo bottle (lather, rinse, repeat) create an infinite loop.</p>
<p>In the case shown above, we can prove that the loop terminates because we know that the value of <code class="docutils literal notranslate"><span class="pre">aBound</span></code> is
finite, and we can see that the value of <code class="docutils literal notranslate"><span class="pre">aNumber</span></code> increments each time through the loop, so eventually it will
have to exceed <code class="docutils literal notranslate"><span class="pre">aBound</span></code>. In other cases, it is not so easy to tell.</p>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>Introduction of the while statement causes us to think about the types of iteration we have seen. The <code class="docutils literal notranslate"><span class="pre">for</span></code>
statement will always iterate through a sequence of values like the list of names for the party or the list of
numbers created by <code class="docutils literal notranslate"><span class="pre">range</span></code>. Since we know that it will iterate once for each value in the collection, it is
often said that a <code class="docutils literal notranslate"><span class="pre">for</span></code> loop creates a <strong>definite iteration</strong> because we definitely know how many times we are
going to iterate.  On the other hand, the <code class="docutils literal notranslate"><span class="pre">while</span></code> statement is dependent on a condition that needs to evaluate
to <code class="docutils literal notranslate"><span class="pre">False</span></code> in order for the loop to terminate. Since we do not necessarily know when this will happen, it
creates what we call <strong>indefinite iteration</strong>. Indefinite iteration simply means that we don’t know how many
times we will repeat but eventually the condition controlling the iteration will fail and the iteration will
stop. (Unless we have an infinite loop which is of course a problem)</p>
</div>
<p>What you will notice here is that the <code class="docutils literal notranslate"><span class="pre">while</span></code> loop is more work for you — the programmer — than the equivalent
<code class="docutils literal notranslate"><span class="pre">for</span></code> loop.  When using a <code class="docutils literal notranslate"><span class="pre">while</span></code> loop you have to control the loop variable yourself. You give it an initial
value, test for completion, and then make sure you change something in the body so that the loop terminates. That
also makes a while loop harder to read and understand than the equivalent for loop. So, while you <em>can</em> implement
definite iteration with a while loop, it’s not a good idea to do that. Use a for loop whenever it will be known at
the beginning of the iteration process how many times the block of code needs to be executed.</p>
<p><strong>Check your understanding</strong></p>

            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="11.2.4" data-multipleanswers="false"  id="question14_2_1"  style="visibility: hidden;">
            <p>moreiter-2-4: True or False: You can rewrite any for-loop as a while-loop.</p>

            <li data-component="answer" data-correct id="question14_2_1_opt_a">True</li><li data-component="feedback">Although the while loop uses a different syntax, it is just as powerful as a for-loop and often more flexible.</li>
            
            <li data-component="answer"  id="question14_2_1_opt_b">False</li><li data-component="feedback">Often a for-loop is more natural and convenient for a task, but that same task can always be expressed using a while loop.</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="11.2.5" data-multipleanswers="false"  id="question14_2_2"  style="visibility: hidden;">
            <p>moreiter-2-5: The following code contains an infinite loop. Which is the best explanation for why the loop does not terminate?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">n</span> <span class="o">=</span> <span class="mi">10</span>
<span class="n">answer</span> <span class="o">=</span> <span class="mi">1</span>
<span class="k">while</span> <span class="p">(</span> <span class="n">n</span> <span class="o">&gt;</span> <span class="mi">0</span> <span class="p">):</span>
  <span class="n">answer</span> <span class="o">=</span> <span class="n">answer</span> <span class="o">+</span> <span class="n">n</span>
  <span class="n">n</span> <span class="o">=</span> <span class="n">n</span> <span class="o">+</span> <span class="mi">1</span>
<span class="nb">print</span><span class="p">(</span><span class="n">answer</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="question14_2_2_opt_a">n starts at 10 and is incremented by 1 each time through the loop, so it will always be positive</li><li data-component="feedback">The loop will run as long as n is positive.  In this case, we can see that n will never become non-positive.</li>
            
            <li data-component="answer"  id="question14_2_2_opt_b">answer starts at 1 and is incremented by n each time, so it will always be positive</li><li data-component="feedback">While it is true that answer will always be positive, answer is not considered in the loop condition.</li>
            
            <li data-component="answer"  id="question14_2_2_opt_c">You cannot compare n to 0 in while loop.  You must compare it to another variable.</li><li data-component="feedback">It is perfectly valid to compare n to 0.  Though indirectly, this is what causes the infinite loop.</li>
            
            <li data-component="answer"  id="question14_2_2_opt_d">In the while loop body, we must set n to False, and this code does not do that.</li><li data-component="feedback">The loop condition must become False for the loop to terminate, but n by itself is not the condition in this case.</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="11.2.6" data-multipleanswers="false"  id="question14_2_3"  style="visibility: hidden;">
            <p>moreiter-2-6: Which type of loop can be used to perform the following iteration: You choose a positive integer at random and then print the numbers from 1 up to and including the selected integer.</p>

            <li data-component="answer" data-correct id="question14_2_3_opt_a">a for-loop or a while-loop</li><li data-component="feedback">Although you do not know how many iterations you loop will run before the program starts running, once you have chosen your random integer, Python knows exactly how many iterations the loop will run, so either a for-loop or a while-loop will work.</li>
            
            <li data-component="answer"  id="question14_2_3_opt_b">only a for-loop</li><li data-component="feedback">As you learned in section 7.2, a while-loop can always be used for anything a for-loop can be used for.</li>
            
            <li data-component="answer"  id="question14_2_3_opt_c">only a while-loop</li><li data-component="feedback">Although you do not know how many iterations you loop will run before the program starts running, once you have chosen your random integer, Python knows exactly how many iterations the loop will run, so this is an example of definite iteration.</li>
            

            </ul>
            </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac14_2_2 data-question_label="11.2.7">
<div id=ac14_2_2_question class="ac_question col-md-12">
<p>Write a while loop that is initialized at 0 and stops at 15. If the counter is an even number, append the counter to a list called <code class="docutils literal notranslate"><span class="pre">eve_nums</span></code>.</p>

</div>
<textarea data-lang="python" id="ac14_2_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
     self.assertEqual(eve_nums, [0,2,4,6,8,10,12,14], "Testing that eve_nums has been assigned the correct elements")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac14_2_3 data-question_label="11.2.8">
<div id=ac14_2_3_question class="ac_question col-md-12">
<p>Below, we’ve provided a for loop that sums all the elements of <code class="docutils literal notranslate"><span class="pre">list1</span></code>. Write code that accomplishes the same task, but instead uses a while loop. Assign the accumulator variable to the name <code class="docutils literal notranslate"><span class="pre">accum</span></code>.</p>

</div>
<textarea data-lang="python" id="ac14_2_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

list1 = [8, 3, 4, 5, 6, 7, 9]

tot = 0
for elem in list1:
    tot = tot + elem

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(accum, 42, "Testing that accum has the correct value.")
      self.assertIn('while', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac14_2_4 data-question_label="11.2.9">
<div id=ac14_2_4_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">stop_at_four</span></code> that iterates through a list of numbers. Using a while loop, append each number to a new list until the number 4 appears. The function should return the new list.</p>

</div>
<textarea data-lang="python" id="ac14_2_4_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

def stop_at_four():



=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(stop_at_four([0, 9, 4.5, 1, 7, 4, 8, 9, 3]), [0, 9, 4.5, 1, 7], "Testing the function stop_at_four on the input [0, 9, 4.5, 1, 7, 4, 8, 9, 3].")
      self.assertEqual(stop_at_four([4, 1, 2, 8]), [], "Testing the function stop_at_four on the input [4, 1, 2, 8].")
      self.assertEqual(stop_at_four([4]), [], "Testing the function stop_at_four on the input [4].")
      self.assertEqual(stop_at_four([1, 6, 2, 3, 9]), [1, 6, 2, 3, 9], "Testing that stop_at_four([1, 6, 2, 3, 9]) returns ([1, 6, 2, 3, 9])")

myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 11.1. Introduction' data-toggle="tooltip">
  <a href="intro-indefiniteiteration.html" aria-label="Previous - 11.1. Introduction">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 11.3. The Listener Loop' data-toggle="tooltip" >
    <a href="listenerLoop.html" aria-label="Next - 11.3. The Listener Loop">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 6.0.10.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>