<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html  lang="en" class="runestone-sphinx">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>Introduction to Problem Types &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/matrixeq.css?v=E2DC0A34" />
    <link rel="stylesheet" type="text/css" href="../../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../../_static/webglinteractive.css?v=4182B24D" />
    <link rel="stylesheet" type="text/css" href="../../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/prefix-347.f9add1ca35d5ad93.css?v=306CBB19" />
    <link rel="stylesheet" type="text/css" href="../../_static/prefix-runestone.f64dcc3632d7a0c3.css?v=F4F91F21" />



    <script data-url_root="../../" id="documentation_options" src="../../_static/documentation_options.js"></script>
    <script defer="" src="../../_static/jquery.js"></script>
    <script defer="" src="../../_static/underscore.js"></script>
    <script src="../../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script defer="" src="../../_static/doctools.js"></script>
    <script src="../../_static/sphinx_highlight.js"></script>
    <script defer="" src="../../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../../_static/prefix-runtime.2438a8bdd6d43b5c.bundle.js?v=7D29D67B"></script>
    <script defer="" src="../../_static/prefix-347.8328b60515045466.bundle.js?v=2CF09D3E"></script>
    <script defer="" src="../../_static/prefix-runestone.aeac7cf9ad5c07ef.bundle.js?v=4FF65C30"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
<script src="https://media.ethicalads.io/media/client/ethicalads.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />

  
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'FOPP-PIE';
    eBookConfig.basecourse = 'FOPP-PIE';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '7.2.11';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" id="rs-book-toc-link" href="../../index.html" aria-label="index-page">
          
          FOPP-PIE
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="ad_warning"></span></li>
        <li><span id="browsing_warning"></span></li>

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://guide.runestone.academy' target="_blank" rel="noopener noreferrer" >Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=FOPP-PIE&page=ps-intro'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <section id="introduction-to-problem-types">
<h1>Introduction to Problem Types<a class="headerlink" href="#introduction-to-problem-types" title="Permalink to this heading">¶</a></h1>
<p>Please read the following, watch the videos, and try to solve the problems.</p>
<section id="solving-mixed-up-code-problems">
<h2>Solving Mixed-up Code Problems<a class="headerlink" href="#solving-mixed-up-code-problems" title="Permalink to this heading">¶</a></h2>
<p>If you see a problem like the one below you will need to put the mixed-up
code in the correct order on the right side. You
may need to indent the blocks as well.  There may also be extra blocks that are not
needed in a correct solution that you can leave on the left side. Click the “Check” button
to check your solution.</p>
<p>See the video below for an example.</p>

    <div class="runestone " style="margin-left: auto; margin-right:auto">
        <div id="iwgex1-parsons1-ps" data-component="youtube" class="align-center youtube-video"
            data-video-height="415" data-question_label="1"
            data-video-width="650" data-video-videoid="Rf7oWHlo-e0"
            data-video-divid="iwgex1-parsons1-ps" data-video-start="0"
            data-video-end="-1" ></div>
    </div>
    <p>Try to solve the following mixed-up code problem.  This problem doesn’t require any indentation.</p>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="intro-simple-parsons-no-indent-ps" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Drag the blocks from the left and put them in the correct order on the right. The text in each block
defines the order.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2"   data-adaptive="true"   data-order="3, 1, 2, 0"     data-numbered="left"  style="visibility: hidden;">
        First block
---
Second block
---
Third block
        </pre>
        </div>
        </div>
    <p>Try to solve the following mixed-up code problem. This problem requires indentation.</p>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="intro-simple-parsons-indent-ps" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Drag the blocks from the left and put them in the correct order on the right with the correct indentation.
The text in each block defines the order and indentation.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="3"   data-adaptive="true"   data-order="3, 1, 2, 0"     data-numbered="left"  style="visibility: hidden;">
        First block
---
Second block
---
    Third block that needs to be indented
        </pre>
        </div>
        </div>
    <p>Try to solve the following mixed-up code problem. This problem requires indentation and has extra blocks that are not needed in a correct solution.</p>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="intro-simple-parsons-indent-with-dist-ps" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Drag the blocks from the left and put them in the correct order on the right with the correct indentation.
There is an extra block that is not needed in the correct solution.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="4"   data-adaptive="true"   data-order="3, 1, 2, 0"     data-numbered="left"  style="visibility: hidden;">
        First block
---
Second block
---
Extra block that is not needed #paired: This block is not needed
---
    Third block that needs to be indented
        </pre>
        </div>
        </div>
    <p>The mixed-up code problems have a “Help me” button at the bottom of the
problem. Once you have checked at least three incorrect solutions you can
click the button for help.  It will remove an incorrect code block, if you used
one in your solution, or combine two blocks into one if there are more
than three blocks left.</p>
<p>See the video below for an example.</p>

    <div class="runestone " style="margin-left: auto; margin-right:auto">
        <div id="iwgex1-parsons2-ps" data-component="youtube" class="align-center youtube-video"
            data-video-height="415" data-question_label="5"
            data-video-width="650" data-video-videoid="QejZ7u642IU"
            data-video-divid="iwgex1-parsons2-ps" data-video-start="0"
            data-video-end="-1" ></div>
    </div>
    </section>
<section id="solving-write-code-problems">
<h2>Solving Write Code Problems<a class="headerlink" href="#solving-write-code-problems" title="Permalink to this heading">¶</a></h2>
<p>If you see a problem like the one below, you will need to write code.  The problem
will have unit tests that you can run to check that your code is working
correctly.  Click on the “Run” button to compile and run your code.  Look after
the code area for compiler errors and/or unit test results.</p>
<p>See the video below for an example.</p>

    <div class="runestone " style="margin-left: auto; margin-right:auto">
        <div id="js-write-code-video-ex" data-component="youtube" class="align-center youtube-video"
            data-video-height="826" data-question_label="6"
            data-video-width="1020" data-video-videoid="w9hTOJ7iJpE"
            data-video-divid="js-write-code-video-ex" data-video-start="0"
            data-video-end="-1" ></div>
    </div>
    <p>Finish writing the code for the following problem.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=intro-sample-write-code-triple-ps data-question_label="7">
<div id=intro-sample-write-code-triple-ps_question class="ac_question">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">triple(num)</span></code> that takes a number <code class="docutils literal notranslate"><span class="pre">num</span></code> and
returns the number times 3. For example, <code class="docutils literal notranslate"><span class="pre">triple(2)</span></code> should return 6 and
<code class="docutils literal notranslate"><span class="pre">triple(-1)</span></code> should return -3.  Look below the code to check for any
compiler errors or the results
from the test cases.  Be sure to <code class="docutils literal notranslate"><span class="pre">return</span></code> the result.</p>

</div>
<textarea data-lang="python" id="intro-sample-write-code-triple-ps_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def triple(num):
    # write code here

print(triple(2))
print(triple(-1))

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(triple(2),6,"triple(2)")
        self.assertEqual(triple(3),9,"triple(3)")
        self.assertEqual(triple(-1),-3,"triple(-1)")
        self.assertEqual(triple(0),0,"triple(0)")
        self.assertEqual(triple(11),33,"triple(11)")

myTests().main()
</textarea>
</div>
</div>
</section>
<section id="feedback">
<h2>Feedback<a class="headerlink" href="#feedback" title="Permalink to this heading">¶</a></h2>

<div class="runestone">
<div data-component="shortanswer" data-question_label="8" class="journal" id=ps-intro-sa    >
<p>Please provide feedback here. Please share any comments, problems, or suggestions.</p>

</div>
</div> <!-- end of runestone div -->
</section>
<section id="what-to-do-next">
<h2>What to do next:<a class="headerlink" href="#what-to-do-next" title="Permalink to this heading">¶</a></h2>
<p>Click on the following link to go the practice problems: <a id="ps-practice"><font size="+2">Practice Problems</font></a></p><script type="text/javascript" >

  window.onload = function() {

    a = document.getElementById("ps-practice")
    a.href = "ps-parsons.html"


    };
</script></section>
</section>


  

  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 7.2.11.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>