<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>6.11. Chapter Assessment &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.fafafbd97df8a0d1.css?v=B5361BC8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.d062db89aaf1dddd.css?v=B02DFDA" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/runtime.b6a5f441f397ace5.bundle.js?v=3F299E7"></script>
    <script defer="" src="../_static/637.d54be67956c5c660.bundle.js?v=915B1670"></script>
    <script defer="" src="../_static/runestone.f2a5bb1b0f5cce0b.bundle.js?v=871C0DAB"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="7. Iteration" href="../Iteration/toctree.html" />
    <link rel="prev" title="6.10. Exercises" href="Exercises.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'COMP1000UofM';
    eBookConfig.basecourse = 'COMP1000UofM';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = false;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '6.0.10';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          COMP1000UofM
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=COMP1000UofM&page=week2a1'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">6.10. </span>Exercises" data-toggle="tooltip">
    <a href="Exercises.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">7. </span>Iteration' data-toggle="tooltip" >              
    <a href="../Iteration/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="chapter-assessment">
<h1><span class="section-number">6.11. </span>Chapter Assessment<a class="headerlink" href="#chapter-assessment" title="Permalink to this headline">¶</a></h1>
<p><strong>Check your understanding</strong></p>

            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="6.11.1" data-multipleanswers="false"  id="assess_question2_1_1_1"  style="visibility: hidden;">
            <p>sequences-10-1: What will the output be for the following code?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">let</span> <span class="o">=</span> <span class="s2">&quot;z&quot;</span>
<span class="n">let_two</span> <span class="o">=</span> <span class="s2">&quot;p&quot;</span>
<span class="n">c</span> <span class="o">=</span> <span class="n">let_two</span> <span class="o">+</span> <span class="n">let</span>
<span class="n">m</span> <span class="o">=</span> <span class="n">c</span><span class="o">*</span><span class="mi">5</span>
<span class="nb">print</span><span class="p">(</span><span class="n">m</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer"  id="assess_question2_1_1_1_opt_a">zpzpzpzpzp</li><li data-component="feedback">The order of concatenation matters.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_1_opt_b">zzzzzppppp</li><li data-component="feedback">Think about the order that the program is executed in, what occurs first?</li>
            
            <li data-component="answer" data-correct id="assess_question2_1_1_1_opt_c">pzpzpzpzpz</li><li data-component="feedback">Yes, because let_two was put before let, c has "pz" and then that is repeated five times.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_1_opt_d">pppppzzzzz</li><li data-component="feedback">Think about the order that the program is executed in, what occurs first?</li>
            
            <li data-component="answer"  id="assess_question2_1_1_1_opt_e">None of the above, an error will occur.</li><li data-component="feedback">This is correct syntax and no errors will occur.</li>
            

            </ul>
            </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=assess_ac_2_1_1_2 data-question_label="6.11.2">
<div id=assess_ac_2_1_1_2_question class="ac_question col-md-12">
<p>Write a program that extracts the last three items in the list <code class="docutils literal notranslate"><span class="pre">sports</span></code> and assigns it to the variable <code class="docutils literal notranslate"><span class="pre">last</span></code>. Make sure to write your code so that it works no matter how many items are in the list.</p>

</div>
<textarea data-lang="python" id="assess_ac_2_1_1_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
sports = ['cricket', 'football', 'volleyball', 'baseball', 'softball', 'track and field', 'curling', 'ping pong', 'hockey']


=====

from unittest.gui import TestCaseGui
import re

class myTests(TestCaseGui):

  def test_output(self):
     self.assertEqual(last, ['curling', 'ping pong', 'hockey'], "Testing that the value of last is the last three items in sports.")
     self.assertTrue(re.search(r'last\s*=\s*\S*3:', self.getEditorText()), "Hardcode check")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=assess_ac_2_1_1_3 data-question_label="6.11.3">
<div id=assess_ac_2_1_1_3_question class="ac_question col-md-12">
<p>Write code that combines the following variables so that the sentence “You are doing a great job, keep it up!” is assigned to the variable <code class="docutils literal notranslate"><span class="pre">message</span></code>. Do not edit the values assigned to <code class="docutils literal notranslate"><span class="pre">by</span></code>, <code class="docutils literal notranslate"><span class="pre">az</span></code>, <code class="docutils literal notranslate"><span class="pre">io</span></code>, or <code class="docutils literal notranslate"><span class="pre">qy</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ac_2_1_1_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
by = "You are"
az = "doing a great "
io = "job"
qy = "keep it up!"


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def test_output(self):
    self.assertEqual(by, 'You are', "Testing original variables.")
    self.assertEqual(az, 'doing a great ', "Testing original variables.")
    self.assertEqual(io, 'job', "Testing original variables.")
    self.assertEqual(qy, 'keep it up!', "Testing original variables.")
    self.assertEqual(message, 'You are doing a great job, keep it up!', "Testing that the value of message is what was expected.")
    self.assertNotIn("You are doing a great job, keep it up!", self.getEditorText(), "Testing for hardcoding (Don't worry about actual and expected values).")


myTests().main()
</textarea>
</div>
</div>

            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="6.11.4" data-multipleanswers="false"  id="assess_question2_1_1_4"  style="visibility: hidden;">
            <p>sequences-10-4: What will the output be for the following code?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">ls</span> <span class="o">=</span> <span class="p">[</span><span class="s1">&#39;run&#39;</span><span class="p">,</span> <span class="s1">&#39;world&#39;</span><span class="p">,</span> <span class="s1">&#39;travel&#39;</span><span class="p">,</span> <span class="s1">&#39;lights&#39;</span><span class="p">,</span> <span class="s1">&#39;moon&#39;</span><span class="p">,</span> <span class="s1">&#39;baseball&#39;</span><span class="p">,</span> <span class="s1">&#39;sea&#39;</span><span class="p">]</span>
<span class="n">new</span> <span class="o">=</span> <span class="n">ls</span><span class="p">[</span><span class="mi">2</span><span class="p">:</span><span class="mi">4</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="n">new</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer"  id="assess_question2_1_1_4_opt_a">['travel', 'lights', 'moon']</li><li data-component="feedback">When we take a slice of something, it includes the item at the first index and excludes the item at the second index.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_4_opt_b">['world', 'travel', 'lights']</li><li data-component="feedback">When we take a slice of something, it includes the item at the first index and excludes the item at the second index. Additionally, Python is a zero-index based language.</li>
            
            <li data-component="answer" data-correct id="assess_question2_1_1_4_opt_c">['travel', 'lights']</li><li data-component="feedback">Yes, python is a zero-index based language and slices are inclusive of the first index and exclusive of the second.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_4_opt_d">['world', 'travel']</li><li data-component="feedback">Python is a zero-index based language.</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="6.11.5" data-multipleanswers="false"  id="assess_question2_1_1_5"  style="visibility: hidden;">
            <p>sequences-10-5: What is the type of <code class="docutils literal notranslate"><span class="pre">m</span></code>?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">l</span> <span class="o">=</span> <span class="p">[</span><span class="s1">&#39;w&#39;</span><span class="p">,</span> <span class="s1">&#39;7&#39;</span><span class="p">,</span> <span class="mi">0</span><span class="p">,</span> <span class="mi">9</span><span class="p">]</span>
<span class="n">m</span> <span class="o">=</span> <span class="n">l</span><span class="p">[</span><span class="mi">1</span><span class="p">:</span><span class="mi">2</span><span class="p">]</span>
</pre></div>
</div>

            <li data-component="answer"  id="assess_question2_1_1_5_opt_a">string</li><li data-component="feedback">Not quite, is it slicing or accessing an element?</li>
            
            <li data-component="answer"  id="assess_question2_1_1_5_opt_b">integer</li><li data-component="feedback">What is happening in the assignment statement for m?</li>
            
            <li data-component="answer"  id="assess_question2_1_1_5_opt_c">float</li><li data-component="feedback">What is happening in the assignment statement for m?</li>
            
            <li data-component="answer" data-correct id="assess_question2_1_1_5_opt_d">list</li><li data-component="feedback">Yes, a slice returns a list no matter how large the slice.</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="6.11.6" data-multipleanswers="false"  id="assess_question2_1_1_6"  style="visibility: hidden;">
            <p>sequences-10-6: What is the type of <code class="docutils literal notranslate"><span class="pre">m</span></code>?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">l</span> <span class="o">=</span> <span class="p">[</span><span class="s1">&#39;w&#39;</span><span class="p">,</span> <span class="s1">&#39;7&#39;</span><span class="p">,</span> <span class="mi">0</span><span class="p">,</span> <span class="mi">9</span><span class="p">]</span>
<span class="n">m</span> <span class="o">=</span> <span class="n">l</span><span class="p">[</span><span class="mi">1</span><span class="p">]</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="assess_question2_1_1_6_opt_a">string</li><li data-component="feedback">Yes, the quotes around the number mean that this is a string.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_6_opt_b">integer</li><li data-component="feedback">Not quite, look again at what is being extracted.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_6_opt_c">float</li><li data-component="feedback">Not quite, look again at what is being extracted.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_6_opt_d">list</li><li data-component="feedback">Not quite, is it slicing or accessing an element?</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="6.11.7" data-multipleanswers="false"  id="assess_question2_1_1_7"  style="visibility: hidden;">
            <p>sequences-10-7: What is the type of <code class="docutils literal notranslate"><span class="pre">x</span></code>?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">b</span> <span class="o">=</span> <span class="s2">&quot;My, what a lovely day&quot;</span>
<span class="n">x</span> <span class="o">=</span> <span class="n">b</span><span class="o">.</span><span class="n">split</span><span class="p">(</span><span class="s1">&#39;,&#39;</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer"  id="assess_question2_1_1_7_opt_a">string</li><li data-component="feedback">Not quite; .split() returns a list, each of whose elements is a string.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_7_opt_b">integer</li><li data-component="feedback">Not quite, look again at what types are present and what the result of .split() is.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_7_opt_c">float</li><li data-component="feedback">Not quite, look again at what types are present and what the result of .split() is.</li>
            
            <li data-component="answer" data-correct id="assess_question2_1_1_7_opt_d">list</li><li data-component="feedback">Yes, the .split() method returns a list.</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="6.11.8" data-multipleanswers="false"  id="assess_question2_1_1_8"  style="visibility: hidden;">
            <p>sequences-10-8: What is the type of <code class="docutils literal notranslate"><span class="pre">a</span></code>?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">b</span> <span class="o">=</span> <span class="s2">&quot;My, what a lovely day&quot;</span>
<span class="n">x</span> <span class="o">=</span> <span class="n">b</span><span class="o">.</span><span class="n">split</span><span class="p">(</span><span class="s1">&#39;,&#39;</span><span class="p">)</span>
<span class="n">z</span> <span class="o">=</span> <span class="s2">&quot;&quot;</span><span class="o">.</span><span class="n">join</span><span class="p">(</span><span class="n">x</span><span class="p">)</span>
<span class="n">y</span> <span class="o">=</span> <span class="n">z</span><span class="o">.</span><span class="n">split</span><span class="p">()</span>
<span class="n">a</span> <span class="o">=</span> <span class="s2">&quot;&quot;</span><span class="o">.</span><span class="n">join</span><span class="p">(</span><span class="n">y</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="assess_question2_1_1_8_opt_a">string</li><li data-component="feedback">Yes, the string is split into a list, then joined back into a string, then split again, and finally joined back into a string.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_8_opt_b">integer</li><li data-component="feedback">Not quite, look again at what types are present and what the result of .split() is.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_8_opt_c">float</li><li data-component="feedback">Not quite, look again at what types are present and what the result of .split() is.</li>
            
            <li data-component="answer"  id="assess_question2_1_1_8_opt_d">list</li><li data-component="feedback">Not quite, think about what .split() and .join() return.</li>
            

            </ul>
            </div>
            
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=assess_ac2_1_1_9 data-question_label="6.11.9">
<div id=assess_ac2_1_1_9_question class="ac_question col-md-12">
<p>Write code to determine how many 9’s are in the list <code class="docutils literal notranslate"><span class="pre">nums</span></code> and assign that value to the variable <code class="docutils literal notranslate"><span class="pre">how_many</span></code>. Do not use a for loop to do this.</p>

</div>
<textarea data-lang="python" id="assess_ac2_1_1_9_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
nums = [4, 2, 23.4, 9, 545, 9, 1, 234.001, 5, 49, 8, 9 , 34, 52, 1, -2, 9.1, 4]


=====

from unittest.gui import TestCaseGui
import re

class myTests(TestCaseGui):

  def test_output(self):
    self.assertEqual(how_many, 3, "Testing that how_many is set correctly.")
    self.assertNotIn('for', self.getEditorText(), "Testing that you didn't use a for loop.")
    self.assertFalse(re.search(r'how_many\s*=\s*3', self.getEditorText()), "Hardcode check")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=assess_ac2_1_1_10 data-question_label="6.11.10">
<div id=assess_ac2_1_1_10_question class="ac_question col-md-12">
<p>Write code that uses slicing to get rid of the the second 8 so that here are only two 8’s in the list bound to the variable <cite>nums</cite>.</p>

</div>
<textarea data-lang="python" id="assess_ac2_1_1_10_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
nums = [4, 2, 8, 23.4, 8, 9, 545, 9, 1, 234.001, 5, 49, 8, 9 , 34, 52, 1, -2, 9.1, 4]


=====

from unittest.gui import TestCaseGui
import re

class myTests(TestCaseGui):

  def test_output(self):
    self.assertEqual(nums, [4, 2, 8, 23.4, 9, 545, 9, 1, 234.001, 5, 49, 8, 9 , 34, 52, 1, -2, 9.1, 4], "Testing that nums is set correctly.")
    self.assertTrue(re.search(r'\s*:', self.getEditorText()), "Testing that you are using slices to remove the second 8 (Don't worry about actual and expected values)")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=access_ac_2_1_1_11 data-question_label="6.11.11">
<div id=access_ac_2_1_1_11_question class="ac_question col-md-12">
<p>Assign the last element of <code class="docutils literal notranslate"><span class="pre">lst</span></code> to the variable <code class="docutils literal notranslate"><span class="pre">end_elem</span></code>. Do this so that it works no matter how long lst is.</p>

</div>
<textarea data-lang="python" id="access_ac_2_1_1_11_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
lst = ["hi", "goodbye", "python", "106", "506", 91, ['all', 'Paul', 'Jackie', "UMSI", 1, "Stephen", 4.5], 109, "chair", "pizza", "wolverine", 2017, 3.92, 1817, "account", "readings", "papers", 12, "facebook", "twitter", 193.2, "snapchat", "leaders and the best", "social", "1986", 9, 29, "holiday", ["women", "olympics", "gold", "rio", 21, "2016", "men"], "26trombones"]

=====

from unittest.gui import TestCaseGui
import re

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(end_elem, lst[-1], "Testing that end_elem has the correct element assigned.")
      self.assertFalse(re.search(r'end_elem\s*=\s*\S26trombones\S', self.getEditorText()), "Hardcoding Check (Don't worry about actual and expected values)")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=access_ac_2_1_1_12 data-question_label="6.11.12">
<div id=access_ac_2_1_1_12_question class="ac_question col-md-12">
<p>Assign the number of elements in <code class="docutils literal notranslate"><span class="pre">lst</span></code> to the variable <code class="docutils literal notranslate"><span class="pre">num_lst</span></code>.</p>

</div>
<textarea data-lang="python" id="access_ac_2_1_1_12_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
lst = ["hi", "goodbye", "python", "106", "506", 91, ['all', 'Paul', 'Jackie', "UMSI", 1, "Stephen", 4.5], 109, "chair", "pizza", "wolverine", 2017, 3.92, 1817, "account", "readings", "papers", 12, "facebook", "twitter", 193.2, "snapchat", "leaders and the best", "social", "1986", 9, 29, "holiday", ["women", "olympics", "gold", "rio", 21, "2016", "men"], "26trombones"]

=====

from unittest.gui import TestCaseGui
import re

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(num_lst, 30, "Testing that num_lst has the correct length assigned.")
      self.assertFalse(re.search(r'num_lst\s*=\s*30', self.getEditorText()), "Hardcoding Check (Don't worry about actual and expected values)")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=assess_ac_2_1_1_13 data-question_label="6.11.13">
<div id=assess_ac_2_1_1_13_question class="ac_question col-md-12">
<p>Create a variable called <code class="docutils literal notranslate"><span class="pre">wrds</span></code> and assign to it a list whose elements are the words in the string <code class="docutils literal notranslate"><span class="pre">sent</span></code>. Do not worry about punctuation.</p>

</div>
<textarea data-lang="python" id="assess_ac_2_1_1_13_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
sent = "The bicentennial for our university was in 2017"

=====

from unittest.gui import TestCaseGui
import re

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(wrds, sent.split(), "Testing that wrds has been correctly assigned.")
      self.assertFalse(re.search(r'wrds\s*=\s*\S*The\S', self.getEditorText()), "Hardcoding Check (Don't worry about actual and expected values)")


myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 6.10. Exercises' data-toggle="tooltip">
  <a href="Exercises.html" aria-label="Previous - 6.10. Exercises">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 7. Iteration' data-toggle="tooltip" >
    <a href="../Iteration/toctree.html" aria-label="Next - 7. Iteration">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 6.0.10.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>