<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html  lang="en" class="runestone-sphinx">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>9.12. Exercises &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=E2DC0A34" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=4182B24D" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-347.f9add1ca35d5ad93.css?v=306CBB19" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-runestone.f64dcc3632d7a0c3.css?v=F4F91F21" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script src="../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script src="../_static/sphinx_highlight.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/prefix-runtime.2438a8bdd6d43b5c.bundle.js?v=7D29D67B"></script>
    <script defer="" src="../_static/prefix-347.8328b60515045466.bundle.js?v=2CF09D3E"></script>
    <script defer="" src="../_static/prefix-runestone.aeac7cf9ad5c07ef.bundle.js?v=4FF65C30"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
<script src="https://media.ethicalads.io/media/client/ethicalads.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="9.13. Chapter Assessment" href="week2a1.html" />
    <link rel="prev" title="9.11. Mutable Object Side Effects" href="MutableObjectSideEffects.html" />

  
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'FOPP-PIE';
    eBookConfig.basecourse = 'FOPP-PIE';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '7.2.11';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" id="rs-book-toc-link" href="../index.html" aria-label="index-page">
          
          FOPP-PIE
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="ad_warning"></span></li>
        <li><span id="browsing_warning"></span></li>

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://guide.runestone.academy' target="_blank" rel="noopener noreferrer" >Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=FOPP-PIE&page=Exercises'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">9.11. </span>Mutable Object Side Effects" data-toggle="tooltip">
    <a href="MutableObjectSideEffects.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">9.13. </span>Chapter Assessment' data-toggle="tooltip" >              
    <a href="week2a1.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <p></p>
<section id="exercises">
<h1><span class="section-number">9.12. </span>Exercises<a class="headerlink" href="#exercises" title="Permalink to this heading">¶</a></h1>
<ol class="arabic">
<li>
        <div class="runestone parsons-container ">
        <div data-component="parsons" id="pp9_11_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write a program that will print out the length of each item in the list as well as the first and last characters of the item.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9.12.1"          style="visibility: hidden;">
        weather = ["sunny", "cloudy", "partially sunny",
           "rainy", "storming", "windy", "foggy",
           "snowy", "hailing"]
---
for condition in weather:
---
    print("The word is", len(condition), "characters")
---
    first_char = condition[0]
    last_char = condition[-1]
---
    print("The first character is: " + first_char)
    print("The last character is: " + last_char)
        </pre>
        </div>
        </div>
    </li>
<li>
        <div class="runestone parsons-container ">
        <div data-component="parsons" id="pp9_11_2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write code to determine how many t’s are in the following sentences.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="9.12.2"          style="visibility: hidden;">
        phrases = ["My, what a lovely day today is!",
"Have you mastered cooking yet? A tasty treat could be in your future.",
"Have you ever seen the leaves change color?"]
---
for sentence in phrases:
---
    print(sentence.count("t"))
        </pre>
        </div>
        </div>
    </li>
<li><div id='q6' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac9_11_6 data-question_label="9.12.3">
<div id=ac9_11_6_question class="ac_question">
<p>Although Python provides us with many list methods, it is good practice and very instructive to think about how they are implemented.  Implement a Python function that works like the following:</p>
<ol class="loweralpha simple">
<li><p>count</p></li>
<li><p>in</p></li>
<li><p>reverse</p></li>
<li><p>index</p></li>
<li><p>insert</p></li>
</ol>

</div>
<textarea data-lang="python" id="ac9_11_6_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=answer9_11_6 data-question_label="9.12.4">
<div id=answer9_11_6_question class="ac_question">

</div>
<textarea data-lang="python" id="answer9_11_6_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def count(obj, lst):
    count = 0
    for e in lst:
        if e == obj:
            count = count + 1
    return count

def is_in(obj, lst):  # cannot be called in() because in is a reserved keyword
    for e in lst:
        if e == obj:
            return True
    return False

def reverse(lst):
    reversed = []
    for i in range(len(lst)-1, -1, -1): # step through the original list backwards
        reversed.append(lst[i])
    return reversed

def index(obj, lst):
    for i in range(len(lst)):
        if lst[i] == obj:
            return i
    return -1

def insert(obj, index, lst):
    newlst = []
    for i in range(len(lst)):
        if i == index:
            newlst.append(obj)
        newlst.append(lst[i])
    return newlst

lst = [0, 1, 1, 2, 2, 3, 4, 5, 6, 7, 8, 9]
print(count(1, lst))
print(is_in(4, lst))
print(reverse(lst))
print(index(2, lst))
print(insert('cat', 4, lst))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function disqus_39ee0274e51d4c888cc20b6fefa4069c_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_39ee0274e51d4c888cc20b6fefa4069c_' + eBookConfig.course;
                    this.page.url = 'https://disqus_39ee0274e51d4c888cc20b6fefa4069c_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_39ee0274e51d4c888cc20b6fefa4069c_' + eBookConfig.course;
            disqus_url = 'http://disqus_39ee0274e51d4c888cc20b6fefa4069c_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="disqus_39ee0274e51d4c888cc20b6fefa4069c" onclick="disqus_39ee0274e51d4c888cc20b6fefa4069c_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='disqus_39ee0274e51d4c888cc20b6fefa4069c']").attr('data-disqus-identifier', 'disqus_39ee0274e51d4c888cc20b6fefa4069c_' + eBookConfig.course);
  });
</script>
</div>
    </div>

</li>
<li><div id='q8' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac9_11_8 data-question_label="9.12.5">
<div id=ac9_11_8_question class="ac_question">
<p>Write a Python function that will take a list of 100 random integers between 0 and 1000 and return the maximum value.  (Note: there is a builtin function named <code class="docutils literal notranslate"><span class="pre">max</span></code> but pretend you cannot use it.)</p>

</div>
<textarea data-lang="python" id="ac9_11_8_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
import random as r
lst = []

for i in range(100):
    num = r.randint(1, 1000)
    lst.append(num)

def largest(lst):
    #your code here

====
from unittest.gui import TestCaseGui
import re
class myTests(TestCaseGui):
    def testOne(self):
        output = self.getOutput().split('\n')
        editor = self.getEditorText().split('\n')
        float_re = r'[-+]?(\d+(\.\d*)?|\.\d+)([eE][-+]?\d+)?'

        self.assertEqual(largest(lst), max(lst), 'Checking for the list'+ str(lst))

        # hardcode check
        self.assertFalse(re.search(r'max', self.getEditorText()), 'Checking for max')
myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=answer9_11_8 data-question_label="9.12.6">
<div id=answer9_11_8_question class="ac_question">

</div>
<textarea data-lang="python" id="answer9_11_8_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import random

def max(lst):
    max = 0
    for e in lst:
        if e > max:
            max = e
    return max

lst = []
for i in range(100):
    lst.append(random.randint(0, 1000))

print(max(lst))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function disqus_714fd5537ebf41189ce5fb6fb16d1d26_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_714fd5537ebf41189ce5fb6fb16d1d26_' + eBookConfig.course;
                    this.page.url = 'https://disqus_714fd5537ebf41189ce5fb6fb16d1d26_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_714fd5537ebf41189ce5fb6fb16d1d26_' + eBookConfig.course;
            disqus_url = 'http://disqus_714fd5537ebf41189ce5fb6fb16d1d26_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="disqus_714fd5537ebf41189ce5fb6fb16d1d26" onclick="disqus_714fd5537ebf41189ce5fb6fb16d1d26_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='disqus_714fd5537ebf41189ce5fb6fb16d1d26']").attr('data-disqus-identifier', 'disqus_714fd5537ebf41189ce5fb6fb16d1d26_' + eBookConfig.course);
  });
</script>
</div>
    </div>

</li>
<li><div id='q9' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac9_11_9 data-question_label="9.12.7">
<div id=ac9_11_9_question class="ac_question">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">sum_of_squares(xs)</span></code> that computes the sum
of the squares of the numbers in the list <code class="docutils literal notranslate"><span class="pre">xs</span></code>.  For example,
<code class="docutils literal notranslate"><span class="pre">sum_of_squares([2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> should return 4+9+16 which is 29:</p>

</div>
<textarea data-lang="python" id="ac9_11_9_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def sum_of_squares(xs):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum_of_squares([2,3,4]),29,"Tested sum_of_squares on input [2,3,4]")
        self.assertEqual(sum_of_squares([0,1,-1]),2,"Tested sum_of_squares on input [0,1,-1]")
        self.assertEqual(sum_of_squares([5,12,14]),365,"Tested sum_of_squares on input [5,12,14]")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q10' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac9_11_10 data-question_label="9.12.8">
<div id=ac9_11_10_question class="ac_question">
<p>Write a function to count how many odd numbers are in a list.</p>

</div>
<textarea data-lang="python" id="ac9_11_10_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def countOdd(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(countOdd([1,3,5,7,9]),5,"Tested countOdd on input [1,3,5,7,9]")
        self.assertEqual(countOdd([1,2,3,4,5]),3,"Tested countOdd on input [-1,-2,-3,-4,-5]")
        self.assertEqual(countOdd([2,4,6,8,10]),0,"Tested countOdd on input [2,4,6,8,10]")
        self.assertEqual(countOdd([0,-1,12,-33]),2,"Tested countOdd on input [0,-1,12,-33]")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=answer9_11_10 data-question_label="9.12.9">
<div id=answer9_11_10_question class="ac_question">

</div>
<textarea data-lang="python" id="answer9_11_10_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import random

def countOdd(lst):
    odd = 0
    for e in lst:
        if e % 2 != 0:
            odd = odd + 1
    return odd

# make a random list to test the function
lst = []
for i in range(100):
    lst.append(random.randint(0, 1000))

print(countOdd(lst))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function disqus_fdd366b1b4c8494082a385e1e1197844_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_fdd366b1b4c8494082a385e1e1197844_' + eBookConfig.course;
                    this.page.url = 'https://disqus_fdd366b1b4c8494082a385e1e1197844_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_fdd366b1b4c8494082a385e1e1197844_' + eBookConfig.course;
            disqus_url = 'http://disqus_fdd366b1b4c8494082a385e1e1197844_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="disqus_fdd366b1b4c8494082a385e1e1197844" onclick="disqus_fdd366b1b4c8494082a385e1e1197844_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='disqus_fdd366b1b4c8494082a385e1e1197844']").attr('data-disqus-identifier', 'disqus_fdd366b1b4c8494082a385e1e1197844_' + eBookConfig.course);
  });
</script>
</div>
    </div>

</li>
<li><div id='q11' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac9_11_11 data-question_label="9.12.10">
<div id=ac9_11_11_question class="ac_question">
<p>Sum up all the even numbers in a list.</p>

</div>
<textarea data-lang="python" id="ac9_11_11_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def sumEven(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(sumEven([1,3,5,7,9]),0,"Tested sumEven on input [1,3,5,7,9]")
      self.assertEqual(sumEven([-1,-2,-3,-4,-5]),-6,"Tested sumEven on input [-1,-2,-3,-4,-5]")
      self.assertEqual(sumEven([2,4,6,7,9]),12,"Tested sumEven on input [2,4,6,7,9]")
      self.assertEqual(sumEven([0,1,12,33]),12,"Tested sumEven on input [0,1,12,33]")

myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</li>
<li><div id='q12' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ac9_11_12 data-question_label="9.12.11">
<div id=ac9_11_12_question class="ac_question">
<p>Sum up all the negative numbers in a list.</p>

</div>
<textarea data-lang="python" id="ac9_11_12_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">
def sumNegatives(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumNegatives([-1,-2,-3,-4,-5]),-15,"Tested sumNegatives on input [-1,-2,-3,-4,-5]")
        self.assertEqual(sumNegatives([1,-3,5,-7,9]),-10,"Tested sumNegatives on input [1,-3,5,-7,9]")
        self.assertEqual(sumNegatives([-2,-4,6,-7,9]),-13,"Tested sumNegatives on input [-2,-4,6,-7,9]")
        self.assertEqual(sumNegatives([0,1,2,3,4]),0,"Tested sumNegatives on input [0,1,2,3,4]")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=answer9_11_12 data-question_label="9.12.12">
<div id=answer9_11_12_question class="ac_question">

</div>
<textarea data-lang="python" id="answer9_11_12_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import random

def sumNegative(lst):
    sum = 0
    for e in lst:
        if e < 0:
            sum = sum + e
    return sum

lst = []
for i in range(100):
    lst.append(random.randrange(-1000, 1000))

print(sumNegative(lst))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function disqus_bfe671ac1e0942f2be4de7179921f83f_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_bfe671ac1e0942f2be4de7179921f83f_' + eBookConfig.course;
                    this.page.url = 'https://disqus_bfe671ac1e0942f2be4de7179921f83f_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_bfe671ac1e0942f2be4de7179921f83f_' + eBookConfig.course;
            disqus_url = 'http://disqus_bfe671ac1e0942f2be4de7179921f83f_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="disqus_bfe671ac1e0942f2be4de7179921f83f" onclick="disqus_bfe671ac1e0942f2be4de7179921f83f_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='disqus_bfe671ac1e0942f2be4de7179921f83f']").attr('data-disqus-identifier', 'disqus_bfe671ac1e0942f2be4de7179921f83f_' + eBookConfig.course);
  });
</script>
</div>
    </div>

</li>
</ol>
</section>


  

  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 9.11. Mutable Object Side Effects' data-toggle="tooltip">
  <a href="MutableObjectSideEffects.html" aria-label="Previous - 9.11. Mutable Object Side Effects">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 9.13. Chapter Assessment' data-toggle="tooltip" >
    <a href="week2a1.html" aria-label="Next - 9.13. Chapter Assessment">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 7.2.11.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>