<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>4.16. Exercises &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.cd688d1623c9932e.css?v=630908D8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.675cfab3e021b8f9.css?v=AC6C1FF6" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/runtime.15a489335c5c6109.bundle.js?v=8DB56D48"></script>
    <script defer="" src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-72cd89.51d9dde4e5162190.bundle.js?v=9CAA2D33"></script>
    <script defer="" src="../_static/runestone.a1906624d1effde5.bundle.js?v=919C0428"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="4.17. Chapter Assessment" href="week2a2.html" />
    <link rel="prev" title="4.15. Glossary" href="Glossary.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'COMP1000UofM';
    eBookConfig.basecourse = 'COMP1000UofM';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = false;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '6.0.9';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          COMP1000UofM
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="browsing_warning"></span></li>

        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=COMP1000UofM&page=Exercises'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">4.15. </span>Glossary" data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">4.17. </span>Chapter Assessment' data-toggle="tooltip" >              
    <a href="week2a2.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <p></p>
<div class="section" id="exercises">
<h1><span class="section-number">4.16. </span>Exercises<a class="headerlink" href="#exercises" title="Permalink to this headline">¶</a></h1>

    <div data-component="question" class="full-width container question " data-question_label="4.16.1" id="q4_16_1" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex4_16_1 data-question_label="4.16.2">
<div id=ex4_16_1_question class="ac_question col-md-12">
<p>In Robert McCloskey’s
book <em>Make Way for Ducklings</em>, the names of the ducklings are Jack, Kack, Lack,
Mack, Nack, Ouack, Pack, and Quack.  This loop tries to output these names in order.</p>
<p>Of course, that’s not quite right because Ouack and Quack are misspelled.
Can you fix it?</p>

</div>
<textarea data-lang="python" id="ex4_16_1_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

prefixes = "JKLMNOPQ"
suffix = "ack"

for p in prefixes:
    print(p + suffix)
====
from unittest.gui import TestCaseGui
import re
class myTests(TestCaseGui):
    def testOne(self):
        # check the answer
        self.assertTrue(re.search('Jack', self.getOutput()), 'Checking answer.')
        self.assertTrue(re.search('Kack', self.getOutput()), 'Checking answer.')
        self.assertTrue(re.search('Lack', self.getOutput()), 'Checking answer.')
        self.assertTrue(re.search('Mack', self.getOutput()), 'Checking answer.')
        self.assertTrue(re.search('Nack', self.getOutput()), 'Checking answer.')
        self.assertTrue(re.search('Ouack', self.getOutput()), 'Checking answer.')
        self.assertTrue(re.search('Pack', self.getOutput()), 'Checking answer.')
        self.assertTrue(re.search('Quack', self.getOutput()), 'Checking answer.')

        # hardcode check
        prefixes = ['J','K','L','M','N','Ou','P','Qu']
        self.assertFalse(re.search(r'('+'|'.join([p+suffix for p in prefixes])+r')', self.getEditorText()), 'Checking for hardcoding')
myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="4.16.3" id="q4_16_2" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex4_16_2 data-question_label="4.16.4">
<div id=ex4_16_2_question class="ac_question col-md-12">
<p>Get the user to enter some text and print it out in reverse order.</p>

</div>
<textarea data-lang="python" id="ex4_16_2_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
import re
class myTests(TestCaseGui):
    def testOne(self):
        self.assertTrue(re.search(user_input[::-1], self.getOutput()), 'Checking answer.')

myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="4.16.5" id="iter_ex_3" >
    <ol start=10 class=arabic><li class="alert alert-warning">

    <div id='q3' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_3_3 data-question_label="4.16.6">
<div id=ex_3_3_question class="ac_question col-md-12">
<dl>
<dt>Write a program that uses a for loop to print</dt><dd><div class="line-block">
<div class="line"><code class="docutils literal notranslate"><span class="pre">One</span> <span class="pre">of</span> <span class="pre">the</span> <span class="pre">months</span> <span class="pre">of</span> <span class="pre">the</span> <span class="pre">year</span> <span class="pre">is</span> <span class="pre">January</span></code></div>
<div class="line"><code class="docutils literal notranslate"><span class="pre">One</span> <span class="pre">of</span> <span class="pre">the</span> <span class="pre">months</span> <span class="pre">of</span> <span class="pre">the</span> <span class="pre">year</span> <span class="pre">is</span> <span class="pre">February</span></code></div>
<div class="line"><code class="docutils literal notranslate"><span class="pre">One</span> <span class="pre">of</span> <span class="pre">the</span> <span class="pre">months</span> <span class="pre">of</span> <span class="pre">the</span> <span class="pre">year</span> <span class="pre">is</span> <span class="pre">March</span></code></div>
<div class="line">etc …</div>
</div>
</dd>
</dl>

</div>
<textarea data-lang="python" id="ex_3_3_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=q3_answer data-question_label="4.16.7">
<div id=q3_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="q3_answer_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
for amonth in ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'November', 'December']:
    print("One of the months of the year is", amonth)
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function b271442ee0864973a023c19f27aeb401_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'b271442ee0864973a023c19f27aeb401_' + eBookConfig.course;
                    this.page.url = 'https://b271442ee0864973a023c19f27aeb401_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'b271442ee0864973a023c19f27aeb401_' + eBookConfig.course;
            disqus_url = 'http://b271442ee0864973a023c19f27aeb401_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="b271442ee0864973a023c19f27aeb401" onclick="b271442ee0864973a023c19f27aeb401_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='b271442ee0864973a023c19f27aeb401']").attr('data-disqus-identifier', 'b271442ee0864973a023c19f27aeb401_' + eBookConfig.course);
  });
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="4.16.8" id="iter_ex_4" >
    <ol start=11 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_3_4 data-question_label="4.16.9">
<div id=ex_3_4_question class="ac_question col-md-12">
<p>Assume you have a list of numbers <code class="docutils literal notranslate"><span class="pre">12,</span> <span class="pre">10,</span> <span class="pre">32,</span> <span class="pre">3,</span> <span class="pre">66,</span> <span class="pre">17,</span> <span class="pre">42,</span> <span class="pre">99,</span> <span class="pre">20</span></code></p>
<ol class="loweralpha simple">
<li><p>Write a loop that prints each of the numbers on a new line.</p></li>
<li><p>Write a loop that prints each number and its square on a new line.</p></li>
</ol>

</div>
<textarea data-lang="python" id="ex_3_4_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

====
from unittest.gui import TestCaseGui
import re
class myTests(TestCaseGui):
    def testOne(self):
        _seq = [12, 10, 32, 3, 66, 17, 42, 99, 20]
        _seq.extend(["{:d} +{:d}".format(number, number**2) for number in _seq])
        _seq = [str(s) for s in _seq]

        output = self.getOutput().split('\n')
        for expected, line in zip(_seq,output):
            self.assertTrue(re.search(expected, line), 'Checking line')
        outer_loops = re.findall(r'^(for[ (].* in.*: *)$', self.getEditorText(), re.M)
        inner_loops = re.findall(r'^( +for[ (].* in.*: *)$', self.getEditorText(), re.M)
        self.assertTrue(len(outer_loops)==2 and len(inner_loops)>=0, 'Checking for-statements')
myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="4.16.10" id="iter_ex_8" >
    <ol start=12 class=arabic><li class="alert alert-warning">

    
        <div class="runestone parsons-container ">
        <div data-component="parsons" id="pp_4_16_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write a program that will go through a list of temperatures and print them out to the user.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="4.16.11"          style="visibility: hidden;">
        temperatures = [-3, 78, 95, 28, 56, 42, 56, 81, -10, -]
---
for temp in temperatures:
---
    print("The weather outside is: " + str(temp))
        </pre>
        </div>
        </div>
    
    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="4.16.12" id="iter_ex_9" >
    <ol start=13 class=arabic><li class="alert alert-warning">

    
        <div class="runestone parsons-container ">
        <div data-component="parsons" id="pp_4_16_2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Write a program that will print out a greeting to each student in the list. This list should also keep track of how many students have been greeted and note that each time a new student has been greeted.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="4.16.13"          style="visibility: hidden;">
        students = ["Jay", "Stacy", "Iman", "Trisha", "Ahmed", "Daniel", "Shadae", "Tosin", "Charlotte"]
---
num_students = 0
---
for student in students:
---
    print("Welcome to class, " + student)
---
    num_students += 1
    print(str(num_students) + "student(s) have entered the classroom")
        </pre>
        </div>
        </div>
    
    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question " data-question_label="4.16.14" id="iter_ex_7" >
    <ol start=14 class=arabic><li class="alert alert-warning">

    <div id='q7' data-component="tabbedStuff"  class='alert alert-warning '><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ex_3_7 data-question_label="4.16.15">
<div id=ex_3_7_question class="ac_question col-md-12">
<p>A drunk pirate makes a random turn and then takes 100 steps forward, makes another random turn, takes another 100 steps, turns another random amount, etc.  A social science student records the angle of each turn before the next 100 steps are taken.  Her experimental data is <code class="docutils literal notranslate"><span class="pre">160,</span> <span class="pre">-43,</span> <span class="pre">270,</span> <span class="pre">-97,</span> <span class="pre">-43,</span> <span class="pre">200,</span> <span class="pre">-940,</span> <span class="pre">17,</span> <span class="pre">-86</span></code>. (Positive angles are counter-clockwise.)  Use a turtle to draw the path taken by our drunk friend.  After the pirate is done walking, print the current heading.  Assume that the turtle originally has a heading of 0 and accumulate the changes in heading to print out the final.  Your solution should work for any sequence of experimental data.</p>

</div>
<textarea data-lang="python" id="ex_3_7_editor" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
     style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=iter_q7_answer data-question_label="4.16.16">
<div id=iter_q7_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="iter_q7_answer_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
import turtle

wn = turtle.Screen()
lovelace = turtle.Turtle()

# move the turtle forward a little so that the whole path fits on the screen
lovelace.penup()
lovelace.forward(60)

# now draw the drunk pirate's path
lovelace.pendown()
current_heading = 0
for angle in [160, -43, 270, -97, -43, 200, -940, 17, -86]:

    # we use .left() so that positive angles are counter-clockwise
    # and negative angles are clockwise
    current_heading = (current_heading + angle) % 360
    lovelace.left(angle)
    lovelace.forward(100)

# the .heading() method gives us the turtle's current heading in degrees
print("The pirate's final heading was", current_heading)

wn.exitonclick()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script>
    function a7e34946f59f348f2bfeb3f918eb57b7a_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'a7e34946f59f348f2bfeb3f918eb57b7a_' + eBookConfig.course;
                    this.page.url = 'https://a7e34946f59f348f2bfeb3f918eb57b7a_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'a7e34946f59f348f2bfeb3f918eb57b7a_' + eBookConfig.course;
            disqus_url = 'http://a7e34946f59f348f2bfeb3f918eb57b7a_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class="disqus_thread_link" data-disqus-identifier="a7e34946f59f348f2bfeb3f918eb57b7a" onclick="a7e34946f59f348f2bfeb3f918eb57b7a_disqus(this);">Show Comments</a>
<script>
  window.addEventListener('load', (event) => {
    $("a[data-disqus-identifier='a7e34946f59f348f2bfeb3f918eb57b7a']").attr('data-disqus-identifier', 'a7e34946f59f348f2bfeb3f918eb57b7a_' + eBookConfig.course);
  });
</script>
</div>
    </div>


    </li></ol>
    </div>
    </div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 4.15. Glossary' data-toggle="tooltip">
  <a href="Glossary.html" aria-label="Previous - 4.15. Glossary">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 4.17. Chapter Assessment' data-toggle="tooltip" >
    <a href="week2a2.html" aria-label="Next - 4.17. Chapter Assessment">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 6.0.9.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>