<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>7.6. The Accumulator Pattern &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.fafafbd97df8a0d1.css?v=B5361BC8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.d062db89aaf1dddd.css?v=B02DFDA" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/runtime.b6a5f441f397ace5.bundle.js?v=3F299E7"></script>
    <script defer="" src="../_static/637.d54be67956c5c660.bundle.js?v=915B1670"></script>
    <script defer="" src="../_static/runestone.f2a5bb1b0f5cce0b.bundle.js?v=871C0DAB"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="7.7. Traversal and the for Loop: By Index" href="TraversalandtheforLoopByIndex.html" />
    <link rel="prev" title="7.5. Lists and for loops" href="Listsandforloops.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'COMP1000UofM';
    eBookConfig.basecourse = 'COMP1000UofM';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = false;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '6.0.10';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          COMP1000UofM
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/ns/books/published/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=COMP1000UofM&page=TheAccumulatorPattern'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">7.5. </span>Lists and <code class="docutils literal notranslate"><span class="pre">for</span></code> loops" data-toggle="tooltip">
    <a href="Listsandforloops.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">7.7. </span>Traversal and the <code class="docutils literal notranslate"><span class="pre">for</span></code> Loop: By Index' data-toggle="tooltip" >              
    <a href="TraversalandtheforLoopByIndex.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="the-accumulator-pattern">
<span id="accum-pattern"></span><h1><span class="section-number">7.6. </span>The Accumulator Pattern<a class="headerlink" href="#the-accumulator-pattern" title="Permalink to this headline">¶</a></h1>
<p>One common programming “pattern” is to traverse a sequence, <strong>accumulating</strong> a value as we go,
such as the sum-so-far or the maximum-so-far. That way, at the end of the traversal we have
accumulated a single value, such as the sum total of all the items or the largest item.</p>
<dl class="simple">
<dt>The anatomy of the accumulation pattern includes:</dt><dd><ul class="simple">
<li><p><strong>initializing</strong> an “accumulator” variable to an initial value (such as 0 if accumulating a sum)</p></li>
<li><p><strong>iterating</strong> (e.g., traversing the items in a sequence)</p></li>
<li><p><strong>updating</strong> the accumulator variable on each iteration (i.e., when processing each item in the sequence)</p></li>
</ul>
</dd>
</dl>
<p>For example, consider the following code, which computes the sum of the numbers in a list.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac6_6_1 data-question_label="7.6.1">
<div id=ac6_6_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac6_6_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
nums = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
accum = 0
for w in nums:
    accum = accum + w
print(accum)
</textarea>
</div>
</div>
<p>In the program above, notice that the variable <code class="docutils literal notranslate"><span class="pre">accum</span></code> starts out with a value of 0.
Next, the iteration is performed 10 times.  Inside the for loop, the update occurs.
<code class="docutils literal notranslate"><span class="pre">w</span></code> has the value of current item (1 the first time, then 2, then 3, etc.).
<code class="docutils literal notranslate"><span class="pre">accum</span></code> is reassigned a new value which is the old value plus the current value of <code class="docutils literal notranslate"><span class="pre">w</span></code>.</p>
<p>This pattern of iterating the updating of a variable is commonly referred to as the
<strong>accumulator pattern</strong>. We refer to the variable as the <strong>accumulator</strong>. This pattern will come up
over and over again. Remember that the key to making it work successfully is to be sure to
initialize the variable before you start the iteration. Once inside the iteration, it is required
that you update the accumulator.</p>
<p>Here is the same program in codelens.  Step through the function and watch the “running total”
accumulate the result.</p>

<div class="runestone" style="max-width: none;">
    <div class="alert alert-warning cd_section" data-component="codelens" data-question_label="7.6.2">
        <div class="pytutorVisualizer" id="clens6_6_1"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 7.6.2 (clens6_6_1)</span> </p>
    </div>

<script>
var clens6_6_1_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["clens6_6_1"] = {"code": "nums = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]\naccum = 0\nfor w in nums:\n   accum = accum + w\nprint(accum)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1]}, "ordered_globals": ["nums"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 0}, "ordered_globals": ["nums", "accum"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 0, "w": 1}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 1, "w": 1}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 1, "w": 2}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 3, "w": 2}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 3, "w": 3}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 6, "w": 3}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 6, "w": 4}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 10, "w": 4}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 10, "w": 5}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 15, "w": 5}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 15, "w": 6}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 21, "w": 6}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 21, "w": 7}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 28, "w": 7}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 28, "w": 8}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 36, "w": 8}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 36, "w": 9}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 45, "w": 9}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 45, "w": 10}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 55, "w": 10}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 55, "w": 10}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": ""}, {"line": 5, "event": "return", "func_name": "<module>", "globals": {"nums": ["REF", 1], "accum": 55, "w": 10}, "ordered_globals": ["nums", "accum", "w"], "stack_to_render": [], "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "stdout": "55\n"}]};
</script>
</div>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>What would happen if we indented the print accum statement? Not sure? Make a prediction, then try it and find out.</p>
</div>
<p>We can utilize the range function in this situation as well. Previously, you’ve seen it used when we wanted to draw in
turtle. There we used it to iterate a certain number of times. We can do more than that though. The <code class="docutils literal notranslate"><span class="pre">range</span></code> function
takes at least one input - which should be an integer - and returns a list as long as your input. While you can provide
two inputs, we will focus on using range with just one input. With one input, range will start at zero and go up to - but
not include - the input. Here are the examples:</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac6_8_10 data-question_label="7.6.3">
<div id=ac6_8_10_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac6_8_10_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
print("range(5): ")
for i in range(5):
    print(i)

print("range(0,5): ")
for i in range(0, 5):
    print(i)

# Notice the casting of `range` to the `list`
print(list(range(5)))
print(list(range(0,5)))

# Note: `range` function is already casted as `list` in the textbook
print(range(5))
</textarea>
</div>
</div>
<p>One important thing to know about the range function in python3 is that if we want to use it outside of iteration, we
have to cast it as a list using <code class="docutils literal notranslate"><span class="pre">list()</span></code>. Inside the textbook you’ll notice that <code class="docutils literal notranslate"><span class="pre">range</span></code> works with or without
casting it as a list but it is best for you to try and get into the habit of casting it as a list. Here’s how you could use the range function in the previous problem.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac6_6_2 data-question_label="7.6.4">
<div id=ac6_6_2_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac6_6_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
accum = 0
for w in range(11):
    accum = accum + w
print(accum)

# or, if you use two inputs for the range function

sec_accum = 0
for w in range(1,11):
    sec_accum = sec_accum + w
print(sec_accum)
</textarea>
</div>
</div>
<p>Because the range function is exclusive of the ending number, we have to use 11 as the function input.</p>
<p>We can use the accumulation pattern is count the number of something or to sum up a total. The
above examples only covered how to get the sum for a list, but we can also count how many items are
in the list if we wanted to.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac6_6_3 data-question_label="7.6.5">
<div id=ac6_6_3_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac6_6_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
nums = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
count = 0
for w in nums:
    count = count + 1
print(count)
</textarea>
</div>
</div>
<p>In this example we don’t make use of <code class="docutils literal notranslate"><span class="pre">w</span></code> even though the iterator variable (loop variable) is a necessary part of
constructing a for loop. Instead of adding the value of <code class="docutils literal notranslate"><span class="pre">w</span></code> to <code class="docutils literal notranslate"><span class="pre">count</span></code> we add a 1 to it,
because we’re incrementing the value of count when we iterate each time through the loop. Though in
this scenario we could have used the <code class="docutils literal notranslate"><span class="pre">len</span></code> function, there are other cases later on where len
won’t be useful but we will still need to count.</p>
<p><strong>Check your understanding</strong></p>

            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-question_label="7.6.6" data-multipleanswers="false"  id="question6_6_1"  style="visibility: hidden;">
            <p>iter-6-6: Consider the following code:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">nums</span> <span class="o">=</span> <span class="p">[</span><span class="mi">1</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">3</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">6</span><span class="p">,</span> <span class="mi">7</span><span class="p">,</span> <span class="mi">8</span><span class="p">,</span> <span class="mi">9</span><span class="p">,</span> <span class="mi">10</span><span class="p">]</span>
<span class="k">for</span> <span class="n">w</span> <span class="ow">in</span> <span class="n">nums</span><span class="p">:</span>
   <span class="n">accum</span> <span class="o">=</span> <span class="mi">0</span>
   <span class="n">accum</span> <span class="o">=</span> <span class="n">accum</span> <span class="o">+</span> <span class="n">w</span>
<span class="nb">print</span><span class="p">(</span><span class="n">accum</span><span class="p">)</span>
</pre></div>
</div>
<p>What happens if you put the initialization of accum inside the for loop as the first
instruction in the loop?</p>

            <li data-component="answer" data-correct id="question6_6_1_opt_a">It will print out 10 instead of 55</li><li data-component="feedback">The variable accum will be reset to 0 each time through the loop. Then it will add the current item. Only the last item will count.</li>
            
            <li data-component="answer"  id="question6_6_1_opt_b">It will cause a run-time error</li><li data-component="feedback">Assignment statements are perfectly legal inside loops and will not cause an error.</li>
            
            <li data-component="answer"  id="question6_6_1_opt_c">It will print out 0 instead of 55</li><li data-component="feedback">Good thought: the variable accum will be reset to 0 each time through the loop. But then it adds the current item.</li>
            

            </ul>
            </div>
            
        <div class="runestone" style="max-width: none;">
        <div data-component="parsons" id="pp6_6_1" class="alert alert-warning parsons" >
        <div class="parsons_question parsons-text" >
    <p>Rearrange the code statements so that the program will add up the first n odd numbers where n is provided by the user.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="7.6.7"         style="visibility: hidden;">
        n = int(input('How many odd numbers would you like to add together?'))
thesum = 0
oddnumber = 1
---
for counter in range(n):
---
   thesum = thesum + oddnumber
   oddnumber = oddnumber + 2
---
print(thesum)
        </pre>
        </div>
        </div>
    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac6_6_4 data-question_label="7.6.8">
<div id=ac6_6_4_question class="ac_question col-md-12">
<p>Write code to create a list of integers from 0 through 52 and assign that list to the variable <code class="docutils literal notranslate"><span class="pre">numbers</span></code>. You should use a special Python function – do not type out the whole list yourself. HINT: You can do this in one line of code!</p>

</div>
<textarea data-lang="python" id="ac6_6_4_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(numbers, range(53), "Testing that numbers is a list that contains the correct elements.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac6_6_10 data-question_label="7.6.9">
<div id=ac6_6_10_question class="ac_question col-md-12">
<p>Count the number of characters in string <code class="docutils literal notranslate"><span class="pre">str1</span></code>. Do not use <code class="docutils literal notranslate"><span class="pre">len()</span></code>. Save the number in variable <code class="docutils literal notranslate"><span class="pre">numbs</span></code>.</p>

</div>
<textarea data-lang="python" id="ac6_6_10_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
str1 = "I like nonsense, it wakes up the brain cells. Fantasy is a necessary ingredient in living."

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testEight(self):
      self.assertEqual(numbs, 90, "Testing that numbs is assigned to correct values.")
      self.assertNotIn("len(", self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac6_8_9 data-question_label="7.6.10">
<div id=ac6_8_9_question class="ac_question col-md-12">
<p>Create a list of numbers 0 through 40 and assign this list to the variable <code class="docutils literal notranslate"><span class="pre">numbers</span></code>. Then, accumulate the total of the list’s values and assign that sum to the variable <code class="docutils literal notranslate"><span class="pre">sum1</span></code>.</p>

</div>
<textarea data-lang="python" id="ac6_8_9_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testNineA(self):
      self.assertEqual(numbers, [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40], "Testing that numbers is assigned to correct values.")

   def testNineB(self):
      self.assertEqual(sum1, 820, "Testing that sum1 has the correct value.")

myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 7.5. Lists and  for  loops' data-toggle="tooltip">
  <a href="Listsandforloops.html" aria-label="Previous - 7.5. Lists and  for  loops">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 7.7. Traversal and the  for  Loop: By Index' data-toggle="tooltip" >
    <a href="TraversalandtheforLoopByIndex.html" aria-label="Next - 7.7. Traversal and the  for  Loop: By Index">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 6.0.10.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>