<!-- ****************************************
     |docname| - main template for this theme
     **************************************** -->
<!DOCTYPE html>

<html  lang="en" class="runestone-sphinx">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>5.15. Chapter Assessment &#8212; COMP 1000: Think Like a Computer</title>
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=E2DC0A34" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=4182B24D" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-347.f9add1ca35d5ad93.css?v=306CBB19" />
    <link rel="stylesheet" type="text/css" href="../_static/prefix-runestone.f64dcc3632d7a0c3.css?v=F4F91F21" />



    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script defer="" src="../_static/jquery.js"></script>
    <script defer="" src="../_static/underscore.js"></script>
    <script src="../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script defer="" src="../_static/doctools.js"></script>
    <script src="../_static/sphinx_highlight.js"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/prefix-runtime.2438a8bdd6d43b5c.bundle.js?v=7D29D67B"></script>
    <script defer="" src="../_static/prefix-347.8328b60515045466.bundle.js?v=2CF09D3E"></script>
    <script defer="" src="../_static/prefix-runestone.aeac7cf9ad5c07ef.bundle.js?v=4FF65C30"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
<script src="https://media.ethicalads.io/media/client/ethicalads.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="6. Conditionals" href="../Conditionals/toctree.html" />
    <link rel="prev" title="5.14. Exercises" href="Exercises.html" />

  
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />



<link rel="canonical" href="https://{{canonical_host}}{{new_server_prefix}}/books/published/{{base_course}}/{{pagepath}}" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'FOPP-PIE';
    eBookConfig.basecourse = 'FOPP-PIE';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
    eBookConfig.new_server_prefix = '';
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "good";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '7.2.11';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.enableScratchAC = true;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/RAIcon.png" alt=""></a>
        
        <a class="navbar-brand" id="rs-book-toc-link" href="../index.html" aria-label="index-page">
          
          FOPP-PIE
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>
        <li><span id="ad_warning"></span></li>
        <li><span id="browsing_warning"></span></li>

        <li class="divider-vertical"></li>
        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://guide.runestone.academy' target="_blank" rel="noopener noreferrer" >Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=FOPP-PIE&page=ChapterAssessment'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">5.14. </span>Exercises" data-toggle="tooltip">
    <a href="Exercises.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">6. </span>Conditionals' data-toggle="tooltip" >              
    <a href="../Conditionals/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  window.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

<!-- Ad Serving for Runestone Campaign -->


  
  <section id="chapter-assessment">
<h1><span class="section-number">5.15. </span>Chapter Assessment<a class="headerlink" href="#chapter-assessment" title="Permalink to this heading">¶</a></h1>

<div class="runestone datafile ">
<div class="datafile_caption">Data file: <code>travel_plans.txt</code></div>
<pre data-component="datafile" id=travel_plans.txt data-hidden data-edit="false" data-rows="20" data-cols="65">
This summer I will be travelling.
I will go to...
Italy: Rome
Greece: Athens
England: London, Manchester
France: Paris, Nice, Lyon
Spain: Madrid, Barcelona, Granada
Austria: Vienna
I will probably not even want to come back!
However, I wonder how I will get by with all the different languages.
I only know English!
</pre></div>

<div class="runestone datafile ">
<div class="datafile_caption">Data file: <code>school_prompt.txt</code></div>
<pre data-component="datafile" id=school_prompt.txt data-hidden data-edit="false" data-rows="20" data-cols="65">
Writing essays for school can be difficult but
many students find that by researching their topic that they
have more to say and are better informed. Here are the university
we require many undergraduate students to take a first year writing requirement
so that they can
have a solid foundation for their writing skills. This comes
in handy for many students.
Different schools have different requirements, but everyone uses
writing at some point in their academic career, be it essays, research papers,
technical write ups, or scripts.
</pre></div>

<div class="runestone datafile ">
<div class="datafile_caption">Data file: <code>emotion_words.txt</code></div>
<pre data-component="datafile" id=emotion_words.txt data-hidden data-edit="false" data-rows="20" data-cols="62">
Sad upset blue down melancholy somber bitter troubled
Angry mad enraged irate irritable wrathful outraged infuriated
Happy cheerful content elated joyous delighted lively glad
Confused disoriented puzzled perplexed dazed befuddled
Excited eager thrilled delighted
Scared afraid fearful panicked terrified petrified startled
Nervous anxious jittery jumpy tense uneasy apprehensive
</pre></div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch5_1 data-question_label="5.15.4">
<div id=assess_ch5_1_question class="ac_question">
<p>The textfile, <code class="docutils literal notranslate"><span class="pre">travel_plans.txt</span></code>, contains the summer travel plans for someone with some commentary. Find the total number of characters in the file and save to the variable <code class="docutils literal notranslate"><span class="pre">num</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch5_1_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(num, 316, "Testing that num value is assigned to correct value.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch5_2 data-question_label="5.15.5">
<div id=assess_ch5_2_question class="ac_question">
<p>We have provided a file called <code class="docutils literal notranslate"><span class="pre">emotion_words.txt</span></code> that contains lines of words that describe emotions. Find the total number of words in the file and assign this value to the variable <code class="docutils literal notranslate"><span class="pre">num_words</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch5_2_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(num_words, 48, "Testing that num_words was assigned to the correct value.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch5_3 data-question_label="5.15.6">
<div id=assess_ch5_3_question class="ac_question">
<p>Assign to the variable <code class="docutils literal notranslate"><span class="pre">num_lines</span></code> the number of lines in the file <code class="docutils literal notranslate"><span class="pre">school_prompt.txt</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch5_3_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(num_lines, 10, "Testing that num_lines has the correct value.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch5_4 data-question_label="5.15.7">
<div id=assess_ch5_4_question class="ac_question">
<p>Assign the first 30 characters of <code class="docutils literal notranslate"><span class="pre">school_prompt.txt</span></code> as a string to the variable <code class="docutils literal notranslate"><span class="pre">beginning_chars</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch5_4_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(len(beginning_chars), 30, "Testing that beginning_chars has the correct length.")
      self.assertEqual(beginning_chars, "Writing essays for school can ", "Testing that beginning_chars has the correct string.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch5_5 data-question_label="5.15.8">
<div id=assess_ch5_5_question class="ac_question">
<p><strong>Challenge:</strong> Using the file <code class="docutils literal notranslate"><span class="pre">school_prompt.txt</span></code>, assign the third word of every line to a list called <code class="docutils literal notranslate"><span class="pre">three</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch5_5_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(three, ['for', 'find', 'to', 'many', 'they', 'solid', 'for', 'have', 'some', 'ups,'], "Testing that three has the correct value.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch5_6 data-question_label="5.15.9">
<div id=assess_ch5_6_question class="ac_question">
<p><strong>Challenge:</strong> Create a list called <code class="docutils literal notranslate"><span class="pre">emotions</span></code> that contains the first word of every line in <code class="docutils literal notranslate"><span class="pre">emotion_words.txt</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch5_6_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(emotions, ['Sad', 'Angry', 'Happy', 'Confused', 'Excited', 'Scared', 'Nervous'], "Testing that emotions was created correctly.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch5_7 data-question_label="5.15.10">
<div id=assess_ch5_7_question class="ac_question">
<p>Assign the first 33 characters from the textfile, <code class="docutils literal notranslate"><span class="pre">travel_plans.txt</span></code> to the variable <code class="docutils literal notranslate"><span class="pre">first_chars</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch5_7_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testFive(self):
      self.assertEqual(first_chars, "This summer I will be travelling.", "Testing that first_chars is assigned to correct value.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch5_8 data-question_label="5.15.11">
<div id=assess_ch5_8_question class="ac_question">
<p><strong>Challenge:</strong> Using the file <code class="docutils literal notranslate"><span class="pre">school_prompt.txt</span></code>, if the character ‘p’ is in a word, then add the word to a list called <code class="docutils literal notranslate"><span class="pre">p_words</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch5_8_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(p_words, ['topic', 'point', 'papers,', 'ups,', 'scripts.'], "Testing that p_words has the correct list.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=assess_ch5_9 data-question_label="5.15.12">
<div id=assess_ch5_9_question class="ac_question">
<p>Read in the contents of the file <code class="docutils literal notranslate"><span class="pre">SP500.txt</span></code> which has monthly data for 2016 and 2017 about the S&amp;P 500 closing prices as well as some other financial indicators, including the “Long Term Interest Rate”, which is interest rate paid on 10-year U.S. government bonds.</p>
<p>Write a program that computes the average closing price (the second column, labeled SP500) and the highest long-term interest rate. Both should be computed only for the period from June 2016 through May 2017. Save the results in the variables <code class="docutils literal notranslate"><span class="pre">mean_SP</span></code> and <code class="docutils literal notranslate"><span class="pre">max_interest</span></code>.</p>

</div>
<textarea data-lang="python" id="assess_ch5_9_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertLess(abs(mean_SP - 2237), 0.5, "Testing that mean_SP is within 0.5 of the correct value. Make sure to use only the correct 12 month period.")
      self.assertEqual(max_interest, 2.49, "Testing the max_interest is correct. Make sure to use only the correct 12 month period.")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone datafile ">
<div class="datafile_caption">Data file: <code>SP500.txt</code></div>
<pre data-component="datafile" id=SP500.txt  data-edit="false" data-rows="20" data-cols="65">
Date,SP500,Dividend,Earnings,Consumer Price Index,Long Interest Rate,Real Price,Real Dividend,Real Earnings,PE10
1/1/2016,1918.6,43.55,86.5,236.92,2.09,2023.23,45.93,91.22,24.21
2/1/2016,1904.42,43.72,86.47,237.11,1.78,2006.62,46.06,91.11,24
3/1/2016,2021.95,43.88,86.44,238.13,1.89,2121.32,46.04,90.69,25.37
4/1/2016,2075.54,44.07,86.6,239.26,1.81,2167.27,46.02,90.43,25.92
5/1/2016,2065.55,44.27,86.76,240.23,1.81,2148.15,46.04,90.23,25.69
6/1/2016,2083.89,44.46,86.92,241.02,1.64,2160.13,46.09,90.1,25.84
7/1/2016,2148.9,44.65,87.64,240.63,1.5,2231.13,46.36,91,26.69
8/1/2016,2170.95,44.84,88.37,240.85,1.56,2251.95,46.51,91.66,26.95
9/1/2016,2157.69,45.03,89.09,241.43,1.63,2232.83,46.6,92.19,26.73
10/1/2016,2143.02,45.25,90.91,241.73,1.76,2214.89,46.77,93.96,26.53
11/1/2016,2164.99,45.48,92.73,241.35,2.14,2241.08,47.07,95.99,26.85
12/1/2016,2246.63,45.7,94.55,241.43,2.49,2324.83,47.29,97.84,27.87
1/1/2017,2275.12,45.93,96.46,242.84,2.43,2340.67,47.25,99.24,28.06
2/1/2017,2329.91,46.15,98.38,243.6,2.42,2389.52,47.33,100.89,28.66
3/1/2017,2366.82,46.38,100.29,243.8,2.48,2425.4,47.53,102.77,29.09
4/1/2017,2359.31,46.66,101.53,244.52,2.3,2410.56,47.67,103.74,28.9
5/1/2017,2395.35,46.94,102.78,244.73,2.3,2445.29,47.92,104.92,29.31
6/1/2017,2433.99,47.22,104.02,244.96,2.19,2482.48,48.16,106.09,29.75
7/1/2017,2454.1,47.54,105.04,244.79,2.32,2504.72,48.52,107.21,30
8/1/2017,2456.22,47.85,106.06,245.52,2.21,2499.4,48.69,107.92,29.91
9/1/2017,2492.84,48.17,107.08,246.82,2.2,2523.31,48.76,108.39,30.17
10/1/2017,2557,48.42,108.01,246.66,2.36,2589.89,49.05,109.4,30.92
11/1/2017,2593.61,48.68,108.95,246.67,2.35,2626.9,49.3,110.35,31.3
12/1/2017,2664.34,48.93,109.88,246.52,2.4,2700.13,49.59,111.36,32.09
</pre></div>

        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="5.15.14" id="assess_ch5_10_lh"  style="visibility: hidden;">
            <p>:</p>
<a class="reference internal image-reference" href="../_images/cdq5-1.png"><img alt="image of numerous variables being declared and assigned values" class="align-center" src="../_images/cdq5-1.png" style="width: 900px;" /></a>
<p>The hello function is in the <input type="text"> scope, while var_x is in the <input type="text"> scope.</p>

        <script type="application/json">
            [[{"regex": "^\\s*global\\s*$", "regexFlags": "i", "feedback": "<p>Correct! This function is in the global scope and avalaible for use everywhere.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "i", "feedback": "<p>Hmmm try again, look at the image to see the different frames</p>\n"}], [{"regex": "^\\s*local\\s*$", "regexFlags": "i", "feedback": "<p>Correct! Variables used in a local scope context are only available in that context</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "i", "feedback": "<p>Hmmm try again, look at the image to see the different frames?</p>\n"}]]
        </script>

        </div>
        </div>
            
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="5.15.15" data-multipleanswers="true"  id="assess_ch5_11_lh"  style="visibility: hidden;">
    <p>:</p>
<img alt="../_images/cdq5-2.png" src="../_images/cdq5-2.png" />
<p>What does the following code print? Be sure to examine the current state of the program what will happen on the next executable line. Select as many as apply.</p>

    <li data-component="answer" data-correct='yes' id="assess_ch5_11_lh_opt_a">an error</li><li data-component="feedback">This is a correct. The code reads top to bottom and we hit an error of hello() not being defined before any prints happen.</li>
    
    <li data-component="answer"  id="assess_ch5_11_lh_opt_b">5</li><li data-component="feedback">Hmm, this is inside the hello() function. Does our code get that far?</li>
    
    <li data-component="answer"  id="assess_ch5_11_lh_opt_c">Hello there, I am a method.</li><li data-component="feedback">Hmm, this is inside the hello() function. Does our code get that far?</li>
    
    <li data-component="answer"  id="assess_ch5_11_lh_opt_d">x</li><li data-component="feedback">x is a variable, not the value. You can tell because it doesn't have quotes</li>
    
    <li data-component="answer"  id="assess_ch5_11_lh_opt_e">10</li><li data-component="feedback">Hmm, this is at the bottom of our code. Does our code get that far?</li>
    

    </ul>
    </div>
    
<div class="runestone ">
<div data-component="clickablearea"  id="assess_ch5_12_lh" data-question_label="5.15.16"     style="visibility: hidden;">
<span data-question>: Given the image below, select the lines of code that call a function.</span><span data-feedback>Remember, the () indicate a function (even if there is something between them) and that lines that start with def are DEFINING a function, not calling one.</span><pre><span data-incorrect>import pickle as p</span>
<span data-incorrect>import random as r</span>
<span data-incorrect>import string as s</span>

<span data-incorrect>def say_hello():</span>
      <span data-incorrect>var_x = "Matilda":</span>
      <span data-correct>print("hi there ", var_x)</span>

<span data-incorrect>def snacks():</span>
  <span data-correct>say_hello()</span>
  <span data-correct>print("I love snacks")</span>

<span data-correct>say_hello()</span>
<span data-incorrect>type_y =</span> <span data-incorrect>"say_hello()"</span></pre>

</div>
</div>
<img alt="../_images/cdq5-3.png" src="../_images/cdq5-3.png" />
</section>


  

  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>

  <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 5.14. Exercises' data-toggle="tooltip">
  <a href="Exercises.html" aria-label="Previous - 5.14. Exercises">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 6. Conditionals' data-toggle="tooltip" >
    <a href="../Conditionals/toctree.html" aria-label="Next - 6. Conditionals">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2022 clatulip.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 7.2.11.
    </p>
  </div>
</footer>




<script>
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch();
  });
</script>


  </body>
</html>